---
id: a103376db3ba46b2d50db289
title: Implementator za konverziju spinalnog slučaja
challengeType: 26
dashedName: implement-a-spinal-case-converter
---

# --description--

Spinalni slučaj je format stringa gde su sve reči u malim slovima i razdvojene crticama. ``"this-is-spinal-tap"`` je primer stringa u spinalnom slučaju.

U ovom laboratorijumu, kreiraćete funkciju koja konvertuje dati string u spinalni slučaj.

**Cilj**: Ispunite korisničke priče ispod i učinite da svi testovi prođu kako biste završili laboratorijum.

**Korisničke priče:**

1. Trebalo bi da kreirate funkciju pod nazivom ``spinalCase``.
2. Funkcija ``spinalCase`` treba da primi jedan argument, string.
3. Funkcija ``spinalCase`` treba da vrati string u formatu spinalnog slučaja. Na primer, ako je argument ``ProductLanding page``, funkcija bi trebalo da vrati ``product-landing-page``.

# --hints--

You should create a function named `spinalCase`.

```js
assert.isFunction(spinalCase);
```

`spinalCase` should take a single argument.

```js
assert.lengthOf(spinalCase, 1);
```

`spinalCase("This Is Spinal Tap")` should return the string `this-is-spinal-tap`.

```js
assert.deepEqual(spinalCase('This Is Spinal Tap'), 'this-is-spinal-tap');
```

`spinalCase("thisIsSpinalTap")` should return the string `this-is-spinal-tap`.

```js
assert.strictEqual(spinalCase('thisIsSpinalTap'), 'this-is-spinal-tap');
```

`spinalCase("The_Andy_Griffith_Show")` should return the string `the-andy-griffith-show`.

```js
assert.strictEqual(
  spinalCase('The_Andy_Griffith_Show'),
  'the-andy-griffith-show'
);
```

`spinalCase("Teletubbies say Eh-oh")` should return the string `teletubbies-say-eh-oh`.

```js
assert.strictEqual(
  spinalCase('Teletubbies say Eh-oh'),
  'teletubbies-say-eh-oh'
);
```

`spinalCase("AllThe-small Things")` should return the string `all-the-small-things`.

```js
assert.strictEqual(spinalCase('AllThe-small Things'), 'all-the-small-things');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function spinalCase(str) {
  str = str.replace(/([a-z](?=[A-Z]))/g, '$1 ');
  return str.toLowerCase().replace(/\ |\_/g, '-');
}
```
