---
id: 6842a6cd9836f0114a5b7a8a
title: Izgradite generator brojnih obrazaca
challengeType: 27
dashedName: build-a-number-pattern-generator
---

# --description--

U ovom laboratorijumu ćete vežbati osnove Pythona kreiranjem male aplikacije koja stvara numerički obrazac.

**Cilj:** Ispunite priče korisnika ispod i učinite da svi testovi prođu kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da definišete funkciju pod nazivom `number_pattern` koja prima jedan parametar `n` (koji predstavlja pozitivan ceo broj).
1. `number_pattern` bi trebalo da koristi petlju `for`.
1. `number_pattern(n)` bi trebalo da vrati niz znakova sa svim celim brojevima počevši od 1 do `n` (uključujući) razdvojenima razmakom. Na primer, `number_pattern(4)` bi trebalo da vrati niz znakova `1 2 3 4`.
1. Ako argument prosleđen funkciji nije ceo broj, funkcija bi trebalo da vrati `Argument must be an integer value.`.
1. Ako je argument prosleđen funkciji manji od 1, funkcija bi trebalo da vrati `Argument must be an integer greater than 0.`.

# --hints--

You should have a `number_pattern` function.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('number_pattern')
    `)
})
```

The `number_pattern` function should have one parameter named `n`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('number_pattern').has_args('n')
    `)
})
```

`number_pattern(4)` should return `1 2 3 4`.

```js
({
    test: () => runPython(`
    assert number_pattern(4) == '1 2 3 4'
    `)
})
```

`number_pattern(12)` should return `1 2 3 4 5 6 7 8 9 10 11 12`.

```js
({
    test: () => runPython(`
    assert number_pattern(12) == '1 2 3 4 5 6 7 8 9 10 11 12'
    `)
})
```

`number_pattern` should return a space separated list of numbers for any positive integer.

```js
({
    test: () => runPython(`
    from random import randint
    n = randint(1,65)
    ls = ' '.join(map(lambda n: str(n), list(range(1, n+1))))
    assert number_pattern(n) == ls

    n2 = randint(34,122)
    ls2 = ' '.join(map(lambda n: str(n), list(range(1, n2+1))))
    assert number_pattern(n2) == ls2
    
    `)
})
```

`number_pattern` should return `Argument must be an integer value.` when passed a value that is not an integer.

```js
({
    test: () => runPython(`
    string = 'Argument must be an integer value.'
    assert number_pattern('a') == string
    assert number_pattern(3.5) == string
    assert number_pattern(lambda n: n) == string
    `)
})
```

`number_pattern` should return `Argument must be an integer greater than 0.` when passed a non-positive integer.

```js
({
    test: () => runPython(`
    from random import randint
    string = 'Argument must be an integer greater than 0.'
    assert number_pattern(0) == string
    assert number_pattern(-1185) == string
    assert number_pattern(-45) == string
    assert number_pattern(randint(-9999999, 0)) == string
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def number_pattern(n):
    result = "1"
    if isinstance(n, int):
        if n > 0:
            for i in range(2, n + 1):
                result += f" {i}"
            return result
        return "Argument must be an integer greater than 0."
    return "Argument must be an integer value."

```
