---
id: 68d275dd800f404d22a07564
title: Implementiraj algoritam pretrage dubine (Depth-First Search)
challengeType: 27
dashedName: lab-depth-first-search
---

# --description--

U ovom laboratorijumu, implementiraćete algoritam pretraživanja grafa koji se zove <dfn>pretraga u dubinu (depth-first search)</dfn>.

Dok <dfn>pretraga u širinu (breadth-first search)</dfn> pretražuje ivice sve veće dužine od izvornog čvora, <dfn>pretraga u dubinu</dfn> prvo ide niz ivica što je moguće dalje.

Kada dođe do jednog kraja puta, pretraga će se vratiti na poslednji čvor sa neposetom ivičnim putem i nastaviti pretraživanje.

Za razliku od pretrage u širinu, svaki put kada se poseti čvor, ne posećuje sve njegove susede. Umesto toga, prvo posećuje jednog od svojih suseda i nastavlja niz tim putem dok nema više čvorova za posetu na tom putu.

Da biste implementirali ovaj algoritam, želite da koristite stogu (niz gde je poslednji dodati element prvi uklonjen, prateći princip <dfn>Last-In-First-Out</dfn>). Stog je koristan u algoritmima pretrage u dubinu jer, dok dodajete susede na stog, želite da prvo posetite najnovije dodate susede i uklonite ih sa stoga.

Jednostavan izlaz ovog algoritma je lista čvorova koji su dostiživi od datog čvora. Stoga, takođe želite da pratite čvorove koje posetite.

**Cilj**: Ispunite korisničke priče ispod i prođite sve testove kako biste završili laboratoriju.

**Korisničke priče:**

1. Treba vam funkcija pod nazivom `dfs`.
1. Funkcija `dfs` treba da primi dva argumenta:
   - Nedirektivnu matricu susednosti (adjacency matrix).
   - Naziv čvora, koji je numerička vrednost čvora između `0` i `n - 1`, gde je `n` ukupan broj čvorova u grafu.
1. Funkcija `dfs` treba da implementira algoritam pretrage u dubinu kako bi izbacila listu svih čvorova dostižive od čvora koji joj je prosleđen.

# --hints--

You should have a function named `dfs` that takes two arguments.

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs)
  sig = inspect.signature(dfs)
  assert len(sig.parameters) == 2
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)` should return a list with `1`, `2`, `3`, and `0`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)
  assert isinstance(result, list)
  assert len(result) == 4
  assert set(result) == {1, 2, 3, 0}
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)` should return a list with `1`, `2`, `3`, and `0`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)
  assert isinstance(result, list)
  assert len(result) == 4
  assert set(result) == {3, 2, 1, 0}
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)` should return `[3]`.

```js
({ test: () => runPython(`
  assert dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3) == [3]
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)` should return a list with `3` and `2`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)
  assert result == [3, 2] or result == [2, 3]
`) })
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)` should return a list with `0` and `1`.

```js
({ test: () => runPython(`
  result = dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)
  assert result == [0, 1] or result == [1, 0]
`) })
```

The `dfs` function should return the correct results.

```js
({ test: () => runPython(`
  import random

  def solve(g, s):
      v, q = set(), [s]
      while q:
          u = q.pop()
          if u not in v:
              v.add(u)
              q.extend([i for i, x in enumerate(g[u]) if x == 1 and i not in v])
      return v

  random.seed(0)

  for _ in range(10):
      n = 6
      mat = [[0]*n for _ in range(n)]
      for i in range(n):
          for j in range(i+1, n):
              if random.random() > 0.5:
                  mat[i][j] = mat[j][i] = 1

      start = random.randint(0, n-1)
      assert set(dfs(mat, start)) == solve(mat, start), "Random graph check failed"
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs(graph, root):
    stack = []
    temp_v = None
    visited = []
    temp_v_neighbors = []
    stack.append(root)
    while stack:
        temp_v = stack.pop()
        if temp_v not in visited:
            visited.append(temp_v)
            temp_v_neighbors = graph[temp_v]
            for n, is_neighbor in enumerate(temp_v_neighbors):
                if is_neighbor == 1:
                    stack.append(n)
    return visited
```
