---
id: a77dbc43c33f39daa4429b4f
title: Izgradite booleovu funkciju za proveru
challengeType: 26
dashedName: build-a-boolean-check-function
---

# --description--

U ovom laboratorijumu, kreiraćete funkciju koja proverava da li je vrednost klasifikovana kao booleanski primitiv.

Booleanski primitivi su `true` i `false`.

**Cilj:** Ispunite priče korisnika ispod i učinite da svi testovi prođu kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da imate funkciju nazvanu `booWho` koja prima jedan argument.
1. Ako je primljeni argument booleanski primitiv, funkcija bi trebalo da vrati `true`.
1. Ako je argument bilo koja druga vrednost, funkcija bi trebalo da vrati `false`.

# --hints--

You should have a `booWho` function.

```js
assert.isFunction(booWho);
```

`booWho(true)` should return `true`.

```js
assert.isTrue(booWho(true));
```

`booWho(false)` should return `true`.

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` should return `false`.

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` should return `false`.

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` should return `false`.

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` should return `false`.

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` should return `false`.

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` should return `false`.

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` should return `false`.

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` should return `false`.

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

```
