---
id: 5688e62ea601b2482ff8422b
title: Izgradnja tražilice profila
challengeType:  26
dashedName: lab-profile-lookup
---

# --description--

U ovom laboratorijumu, kreiraćete pretragu profila koja traži informacije o ljudima u spisku kontakata.

Za ovaj primer, zamislite da postoji kontakt po imenu Akira Laine, i da bi `lookUpProfile("Akira", "lastName")` trebalo da vrati `Laine`. 

**Cilj:** Ispunite priče korisnika ispod i učinite da svi testovi prođu kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da kreirate funkciju pod nazivom `lookUpProfile` koja prima ime i svojstvo kao argumente. 
2. Trebaće vam da povučete informacije o kontaktu iz dostavljenog niza `contacts`. 
3. Ako funkcija primi ime kontakta i svojstvo postoji na vezi kontaktu, onda bi trebalo da se vrati vrednost tog svojstva.
4. Ako ime prosleđeno funkciji ne odgovara nijednom kontaktu u nizu kontakata, funkcija bi trebalo da vrati `"No such contact"`.
5. Ako svojstvo ne postoji na pronađenom kontaktu, funkcija bi trebalo da vrati `"No such property"`. 

# --before-each--

```js
const _contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];

contacts = structuredClone(_contacts); 
```

# --hints--

The `contacts` array should still be present in the global scope. Reset the lesson to recover it if you deleted it.

```js
assert.isArray(contacts);
```

You should have a function named `lookUpProfile`.

```js
assert.isFunction(lookUpProfile);
```

`lookUpProfile("Kristian", "lastName")` should return a string.

```js
assert.isString(lookUpProfile("Kristian", "lastName"));
```

`lookUpProfile("Kristian", "lastName")` should return the string `Vos`

```js
assert.strictEqual(lookUpProfile('Kristian', 'lastName'), 'Vos');
```

`lookUpProfile("Sherlock", "likes")` should return `["Intriguing Cases", "Violin"]`

```js
assert.deepEqual(lookUpProfile('Sherlock', 'likes'), [
  'Intriguing Cases',
  'Violin'
]);
```

`lookUpProfile("Harry", "likes")` should return an array

```js
assert.isArray(lookUpProfile('Harry', 'likes'));
```

`lookUpProfile("Bob", "number")` should return the string `No such contact`

```js
assert.strictEqual(lookUpProfile('Bob', 'number'),'No such contact');
```

`lookUpProfile("Bob", "potato")` should return the string `No such contact`

```js
assert.strictEqual(lookUpProfile('Bob', 'potato'), 'No such contact');
```

`lookUpProfile("Akira", "address")` should return the string `No such property`

```js
assert.strictEqual(lookUpProfile('Akira', 'address'), 'No such property');
```

The `contacts` array should remain unchanged after running `lookUpProfile` 

```js
lookUpProfile("Sherlock", "likes");
assert.deepEqual(contacts,_contacts);
```

# --seed--

## --seed-contents--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


```

# --solutions--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


function lookUpProfile(name, prop) {
  for (let i in contacts) {
    if (contacts[i].firstName === name) {
      return contacts[i][prop] || "No such property";
    }
  }
  return "No such contact";
}
```
