---
id: 68a23004e96a105df3a21fc4
title: Napravi asistenta za brojanje karata
challengeType: 26
dashedName: lab-counting-cards
---

# --description--

U igri u kasinu Blackjack, igrač može utvrditi da li ima prednost u sledećem rundu protiv kuće praćenjem broja visokih i niskih karata preostalihih u radi. Ovo se zove brojanje karata.

Imeći mnogo visokih karti ostalih na igri je korisno za igrača. Kada je račun pozitivan, igrač bi trebalo da stavi veći ulog. Kada je račun nula ili negativan, igrač bi trebalo da stavi manji ulog.

**Cilj:** Popunite priče korisnika ispod i prođite sve testove da biste završili laboratoriju.

Priče o korisniku:

1. Unapaswa kutumia `let` kutangaza kigezo cha kimataifa kinachoitwa `count` na kukiweka kuwa `0`.
1. Unapaswa kuwa na kitendakazi kinachoitwa `cardCounter`.
1. Kitendakazi cha `cardCounter` kinapaswa kupokea kigezo `card` ambacho kinaweza kuwa nambari au mfuatano wa herufi.
    - Kwa thamani kati ya `2` hadi `10`, kigezo cha `card` kitakuwa nambari.
    - Kwa thamani nyingine zote, kigezo cha `card` kitakuwa mfuatano wa herufi.
1. Kitendakazi cha `cardCounter` kinapaswa kubadilisha kigezo cha kimataifa `count` kulingana na vigezo fulani.
1. Kigezo cha kimataifa `count` kinapaswa kuongezwa kwa `1` kwa kadi `2`, `3`, `4`, `5`, au `6`.
1. Kigezo cha kimataifa `count` kinapaswa kubaki kisibadilike kwa kadi `7`, `8`, `9`.
1. Kigezo cha kimataifa `count` kinapaswa kupunguzwa kwa `1` kwa kadi `10`, `"J"`, `"Q"`, `"K"`, `"A"`.
1. Kitendakazi cha `cardCounter` kinapaswa kurudisha mfuatano wa herufi wenye hesabu ya sasa na mfuatano wa herufi `Bet` ikiwa hesabu ni chanya.
1. Kitendakazi cha `cardCounter` kinapaswa kurudisha mfuatano wa herufi wenye hesabu ya sasa na mfuatano wa herufi `Hold` ikiwa hesabu ni ndogo au sawa na `0`.
1. Katika matokeo ya kitendakazi, hesabu ya sasa na uamuzi wa mchezaji (`Bet` au `Hold`) vinapaswa kutenganishwa na nafasi moja. Kwa mfano, `-3 Hold`.

# --hints--

Morate koristiti `let` za deklarisanje globalne varijable nazvane `count` i postaviti je na `0`.

```js
assert.match(__helpers.removeJSComments(code), /^\s*let\s+count\s*=\s*0\s*;?\s*$/m);
```

Trebalo bi da imaš funkciju nazvanu `cardCounter`.

```js
assert.isFunction(cardCounter);
```

Vaša funkcija mora da vrati vrednost `count` i tekst (`Bet` ili `Hold`) na istoj poziciji između njih.

```js
count = 0;
let out = cardCounter(10);
const pattern = /^-?\d+ (Bet|Hold)$/;
assert.match(out, pattern);
```

Nakon kartica `2`, `3`, `4`, `5`, onda se naziva `cardCounter(6)` i treba da vrati sekvencu slova `5 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter(3);
cardCounter(4);
cardCounter(5);
const out = cardCounter(6);
assert.strictEqual(out, '5 Bet')
```

Nakon kartice `7`, `8`, zatim se poziva `cardCounter(9)` i mora vratiti sekvencu slova `0 Hold`.

```js
count = 0;
cardCounter(7);
cardCounter(8);
const out = cardCounter(9);
assert.strictEqual(out, '0 Hold');
```

Nakon kartice `10`, `"J"`, `"Q"`, `"K"`, onda se zove `cardCounter("A")` i mora da vrati sekvencu karaktera `-5 Hold`.

```js
count = 0;
cardCounter(10);
cardCounter('J');
cardCounter('Q');
cardCounter('K');
const out = cardCounter('A');
assert.strictEqual(out, '-5 Hold');
```

Nakon kartice `3`, `7`, `"Q"`, `8`, ona se zove `cardCounter("A")` i treba da vrati niz karaktera `-1 Hold`.

```js
count = 0;
cardCounter(3);
cardCounter(7);
cardCounter('Q');
cardCounter(8);
const out = cardCounter('A');
assert.strictEqual(out, '-1 Hold');
```

Nakon kartice `2`, `"J"`, `9`, `2`, kada se pozove `cardCounter(7)`, trebalo bi da vrati niz slova `1 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter('J');
cardCounter(9);
cardCounter(2);
const out = cardCounter(7);
assert.strictEqual(out, '1 Bet');
```

Nakon kartice `2`, `2`, kada se pozove `cardCounter(10)`, treba da vrati niz karaktera `1 Bet`.

```js
count = 0;
cardCounter(2);
cardCounter(2);
const out = cardCounter(10);
assert.strictEqual(out, '1 Bet')
```

Nakon kartica `3`, `2`, `"A"`, `10`, zatim nazvana `cardCounter("K")` mora da vrati sekvencu slova `-1 Hold`.

```js
count = 0;
cardCounter(3);
cardCounter(2);
cardCounter('A');
cardCounter(10);
const out = cardCounter('K');
assert.strictEqual(out, '-1 Hold');
```

# --seed--

## --seed-contents--


```js



```

# --solutions--

```js
let count = 0;
function cardCounter(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
