---
id: 69446f18fdcb11aa516f7c5c
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovom radionici, kreiraćete funkciju `countup`, koja kreira niz brojeva od 1 do zadatog broja koristeći rekurziju.

Počnite kreiranjem prazne funkcije pod nazivom `countup` koja prima jedan parametar.

# --hints--

You should have a function named `countup`.

```js
assert.isFunction(countup);
```

Your `countup` function should accept a single parameter.

```js
assert.strictEqual(countup.length, 1);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
