---
id: 69447054012a22f38c944ff0
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Dodajte ``if`` izjavu unutar vaše funkcije ``countup``.

Proverite da li je ``number`` manje od ``1`` i vratite prazan niz.

Ovo će služiti kao bazni slučaj za vašu rekurziju.

# --hints--

When `number` is less than `1` the function should return an empty array.

```js
assert.deepStrictEqual(countup(0), []);
```

When `number` is **not** less than `1` the function **should not** return anything.

```js
assert.isUndefined(countup(1)); 
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
