---
id: 694870bcf3c6874112b6f951
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Za sve brojeve veće ili jednake `1`, kreirajte blok `else` i unutar njega pozovite funkciju `countup` sa `number - 1` i dodelite rezultat `countArray`.

Zatim koristite `.push()` da sabirate `number` sa `countArray`, i na kraju vratite `countArray`.

Ovo osigurava da rekurzija konačno stigne do baznog slučaja i prekine.

# --hints--

You should assign the result of calling `countup` with `number - 1` to `countArray`.

```js
assert.match(countup.toString(), /countArray\s*=\s*countup\(\s*number\s*-\s*1\s*\)/);
```

You should push `number` into `countArray`.

```js
assert.match(countup.toString(), /countArray\.push\(\s*number\s*\)/);
```

You should return the `countArray`.

```js
assert.deepStrictEqual(countup(3), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
