---
id: 67051431a73f1ca25d9a6f25
title: Izgradnja Konceptne Tablice Projekta
challengeType: 26
dashedName: build-a-project-idea-board
---

# --description--

**Cilj:** Popunite korisničke priče ispod i prođite sve testove kako biste završili ovaj laboratorijski rad.

**Priče korisnika:**

1. Unapaswa kufafanua thabiti ya kitu kinachoitwa `projectStatus` yenye funguo tatu: `PENDING`, `SUCCESS`, na `FAILURE`. Kila hali inapaswa kupewa kitu chenye sifa ya `description` yenye thamani `Pending Execution`, `Executed Successfully`, na `Execution Failed`, kwa mpangilio huo.
2. Unapaswa kufafanua darasa la vitu linaloitwa `ProjectIdea` lenye `constructor` linalopokea `title` na `description` kama vigezo. Anzisha sifa za `title` na `description` kwa vigezo vilivyotolewa. Darasa hilo pia linapaswa kuwa na sifa inayoitwa `status` ambayo imewekwa thamani `projectStatus.PENDING` kwa chaguo-msingi.
3. Unapaswa kufafanua njia inayoitwa `updateProjectStatus` ndani ya darasa la `ProjectIdea`. Njia hii inapaswa kupokea kigezo cha `newStatus` na kusasisha sifa ya `status` kwa thamani iliyotolewa.
4. Unapaswa kufafanua darasa la `ProjectIdeaBoard` lenye `constructor` linalopokea `title` na kuanzisha safu ya data tupu inayoitwa `ideas` kuhifadhi matukio ya darasa la `ProjectIdea`.
5. Unapaswa kufafanua njia inayoitwa `pin` ndani ya darasa la `ProjectIdeaBoard` inayopokea tukio la darasa la `ProjectIdea` na kuingiza tukio hilo kwenye safu ya data ya `ideas`.
6. Unapaswa kufafanua njia inayoitwa `unpin` ndani ya darasa la `ProjectIdeaBoard`. Njia hii inapaswa kupokea tukio la darasa la `ProjectIdea` na kuliondoa kwenye safu ya data ya `ideas`.
7. Unapaswa kufafanua njia inayoitwa `count` inayorejesha idadi ya mawazo ya miradi katika safu ya data ya `ProjectIdeaBoard` iliyotolewa.
8. Unapaswa kufafanua njia inayoitwa `formatToString` inayorejesha majina ya mawazo ya miradi, maelezo yao na hali yao kwa muundo:

```js
<ProjectIdeaBoard title> has <ProjectIdeaBoard count> idea(s)
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
.
.
.
```

# --hints--

Morate definisati tabelu nečega pod nazivom `projectStatus`, sa ispravnim statusima projekata i njihovim opisima.

```js
assert.isObject(projectStatus);
assert.sameMembers(Object.keys(projectStatus), ["PENDING", "SUCCESS", "FAILURE"]);
assert.deepEqual(projectStatus.PENDING, { description: "Pending Execution" });
assert.deepEqual(projectStatus.SUCCESS, { description: "Executed Successfully" });
assert.deepEqual(projectStatus.FAILURE, { description: "Execution Failed" });
```

Trebalo bi da imaš klasu `ProjectIdea`.

```js
assert.isFunction(ProjectIdea);
assert.isObject(new ProjectIdea());
```

Tvoja klasa za `ProjectIdea` treba da inicijalizuje svojstva za `title` i `description` prema navedenim kriterijumima. Takođe mora da postavi stanje na `projectStatus.PENDING`.

```js
const test = new ProjectIdea("sample title", "sample description");
assert.isAbove(test.title.length, 0);
assert.isAbove(test.description.length, 0);
assert.strictEqual(test.status.description, "Pending Execution");
```

Vaša klasa za `ProjectIdea` treba da ima metod za `updateProjectStatus`.

```js
const projectIdea = new ProjectIdea();
assert.isFunction(projectIdea.updateProjectStatus);
```

Trebalo bi da imaš razred `ProjectIdeaBoard`.

```js
assert.isFunction(ProjectIdeaBoard);
assert.isObject(new ProjectIdeaBoard());
```

`ProjectIdeaBoard` treba da inicijalizuje svojstvo `title` prema datoj vrednosti i kreira prazan niz podataka nazvan `ideas` za skladištenje zapisa klase `ProjectIdea`.

```js
const test = new ProjectIdeaBoard("Sample Title");
assert.isArray(test.ideas);
assert.lengthOf(test.ideas, 0);
assert.exists(test.title);
```

Vaša klasa za `ProjectIdeaBoard` treba da ima metodu za `pin`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.pin);
```

Darasa lako za `ProjectIdeaBoard` treba da ima metodu `unpin`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.unpin);
```

Vaša klasa od `ProjectIdeaBoard` treba da ima metodu za `count`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.count);
```

Tvoja klasa za `ProjectIdeaBoard` treba da ima metodu `formatToString`.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.formatToString);
```

`new ProjectIdea("Smart Window Locks", "An automation project allowing users to lock, unlock windows automatically based on weather conditions.")` treba da vrati `{ title: 'Smart Window Locks', description: 'An automation project allowing users to lock, unlock windows automatically based on weather conditions.', status: { description: 'Pending Execution' } }`.

```js
const windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);
assert.deepEqual(windowLocks, {
  title: "Smart Window Locks",
  description:
    "An automation project allowing users to lock, unlock windows automatically based on weather conditions.",
  status: { description: "Pending Execution" },
});
```

Poziv za `updateProjectStatus(projectStatus.SUCCESS)` na `new ProjectIdea("Fitness Tracker App", "An app that tracks user workouts, diet, and sleep patterns.")` treba ažurirati status na `{ description: 'Executed Successfully' }`.

```js
let mobileApp = new ProjectIdea(
  "Fitness Tracker App",
  "An app that tracks user workouts, diet, and sleep patterns."
);

// Update the project status to SUCCESS
mobileApp.updateProjectStatus(projectStatus.SUCCESS);
assert.deepEqual(mobileApp, {
  title: "Fitness Tracker App",
  description: "An app that tracks user workouts, diet, and sleep patterns.",
  status: { description: "Executed Successfully" },
});
```

Poziv za `updateProjectStatus(projectStatus.FAILURE)` na `new ProjectIdea("Breakfast Chef Robot", "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.")` treba da ažurira status u `{ description: 'Execution Failed' }`.

```js
let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);

// Update the project status to FAILURE
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

assert.deepEqual(breakfastChef, {
  title: "Breakfast Chef Robot",
  description:
    "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.",
  status: { description: "Execution Failed" },
});
```

Poziv `updateProjectStatus(projectStatus.SUCCESS)` na `new ProjectIdea("Online Used Video Games Store", "An online platform where users can buy second hand physical copies of video games from other users.")` treba da ažurira status u `{ description: 'Executed Successfully' }`.

```js
let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);

// Update the project status set to SUCCESS
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

assert.deepEqual(usedVideoGamesStore, {
  title: "Online Used Video Games Store",
  description:
    "An online platform where users can buy second hand physical copies of video games from other users.",
  status: { description: "Executed Successfully" },
});
```

Trebalo bi da možeš postaviti nešto iz `ProjectIdea` u `ProjectIdeaBoard` koristeći način `pin`.

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin the idea
let techProjects = new ProjectIdeaBoard("Tech Projects Board");
techProjects.pin(smartHome);

// Correct the assertion to compare the first idea in the array
assert.deepEqual(techProjects.ideas[0], {
    title: 'Smart Home System',
    description: 'An integrated system to control lighting, temperature, and security devices remotely.',
    status: { description: 'Pending Execution' }
});
```

Trebalo bi da možete ukloniti nešto iz `ProjectIdea` na `ProjectIdeaBoard` koristeći metod `unpin`.

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

techProjects.pin(smartHome);
techProjects.unpin(smartHome);

assert.isEmpty(techProjects.ideas);
```

Kada je `new ProjectIdeaBoard("Empty Board")` prazno, `emptyBoard.formatToString()` treba da restaurira `Empty Board has 0 idea(s)\n`.

```js
let emptyBoard = new ProjectIdeaBoard("Empty Board");

assert.equal(emptyBoard.formatToString(), "Empty Board has 0 idea(s)\n");
```

Kada postavite `new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.")` u `new ProjectIdeaBoard("Tech Projects Board")`, `techProjects.formatToString()` treba da restaurira `Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n`.

```js
let techProjectsBoard = new ProjectIdeaBoard("Tech Projects Board");
let smartHome = new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.");
techProjectsBoard.pin(smartHome);

assert.equal(techProjectsBoard.formatToString(), "Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n"); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const projectStatus = {
  PENDING: { description: "Pending Execution" },
  SUCCESS: { description: "Executed Successfully" },
  FAILURE: { description: "Execution Failed" },
};

class ProjectIdea {
  constructor(title, description) {
    this.title = title;
    this.description = description;
    this.status = projectStatus.PENDING;
  }

  updateProjectStatus(newStatus) {
    this.status = newStatus;
  }
}

class ProjectIdeaBoard {
  constructor(title) {
    this.title = title;
    this.ideas = [];
  }

  pin(newIdea) {
    this.ideas.push(newIdea);
  }

  unpin(idea) {
    this.ideas.splice(this.ideas.indexOf(idea), 1);
  }

  count() {
    return this.ideas.length;
  }

  formatToString() {
    let representation = `${this.title} has ${this.count()} idea(s)\n`;

    this.ideas.forEach(
      (idea) =>
        (representation += `${idea.title} (${idea.status.description}) - ${idea.description}\n`),
    );
    return representation;
  }
}

let windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);

let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

let hardwareProjectIdeas = new ProjectIdeaBoard("Hardware Project Board");
hardwareProjectIdeas.pin(windowLocks);
hardwareProjectIdeas.pin(breakfastChef);

let softwareProjectIdeas = new ProjectIdeaBoard("Software Project Board");
softwareProjectIdeas.pin(usedVideoGamesStore);

console.log(softwareProjectIdeas.formatToString());

console.log(hardwareProjectIdeas.formatToString());
hardwareProjectIdeas.unpin(breakfastChef);
console.log(hardwareProjectIdeas.formatToString());
```
