---
id: 66e2df8c5676fe51cead8ec8
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Sada je vreme da testirate vašu `getVowelCount` funkciju.

Kreirajte varijablu `vowelCount` i dodelite joj rezultat pozivanja funkcije `getVowelCount` sa argumentom `"Apples are tasty fruits"`.

Nakon toga, logujte sledeće u konzolu: `"Vowel Count: [vowel count goes here]"`. Zamenite `[vowel count goes here]` stvarnim imenom varijable. Možete izabrati da koristite *template strings* ili konkatenaciju nizova sa operaterom `+` ovde.

# --hints--

You should create a `vowelCount` variable.

```js
assert.isNotNull(vowelCount)
```

Your `vowelCount` variable should be set to the result of `getVowelCount("Apples are tasty fruits")`.

```js
assert.equal(vowelCount, getVowelCount("Apples are tasty fruits"));
```

You should log the `vowelCount` variable to the console with the parameter `"Vowel Count: ${vowelCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Vowel\s+Count:\s+('|"|`)\s+\+\s+vowelCount|`Vowel\s+Count:\s+\${vowelCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
