---
id: 66e2eab8a5638f57b637b7cc
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Trebalo bi da sada prebrojite broj interpunkcijskih znakova.

Kreirajte funkciju `getPunctuationCount` sa parametrom `sentence`.

Unutar funkcije, kreirajte petlju za brojanje broja interpunkcijskih znakova u `sentence` koji će biti prosleđen funkciji kada se pozove. Za naše svrhe, interpunkcijski znak je bilo koji karakter koji nije razmak (`" "`) ili slovo.

Vaša funkcija `getPunctuationCount` mora vratiti broj.

# --hints--

You should create a `getPunctuationCount` function.

```js
assert.isFunction(getPunctuationCount);
```

Your `getPunctuationCount` function should have a `sentence` parameter.

```js
assert.match(getPunctuationCount.toString(), /sentence/);
```

Your `getPunctuationCount` function should return a number.

```js
assert.isNumber(getPunctuationCount("Coding is fun!"))
```

When the sentence is `"What's going on here?"`, the `getPunctuationCount` function should return `2`.

```js
assert.strictEqual(getPunctuationCount("What's going on here?"), 2);
```

When the sentence is `"What????!"`, the `getPunctuationCount` function should return `5`.

```js
assert.strictEqual(getPunctuationCount("What????!"), 5);
```

Your `getPunctuationCount` function should return the correct punctuation count for any sentence.

```js
assert.strictEqual(getPunctuationCount("Be quick, sign up! freeCodeCamp awaits, friend!!!"), 6);
assert.strictEqual(getPunctuationCount("Guess what? freeCodeCamp is launching a new cert soon!"), 2);
assert.strictEqual(getPunctuationCount("freeCodeCamp, again? It's incredible!"), 4);
assert.strictEqual(getPunctuationCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
