---
id: 66e2e99caa0cb3570f4822b4
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Sada je vreme da testirate vašu `getConsonantCount` funkciju.

Kreirajte varijablu `consonantCount` i dodelite joj rezultat pozivanja funkcije `getConsonantCount` sa argumentom `"Coding is fun"`.

Nakon toga, logujte sledeće u konzolu: `"Consonant Count: [Consonant count goes here]"`. Zamenite `[Consonant count goes here]` stvarnim imenom varijable. Možete izabrati da koristite *template strings* ili konkatenaciju nizova sa operaterom `+` ovde.

# --hints--

You should create a `consonantCount` variable.

```js
assert.isNotNull(consonantCount)
```

Your `consonantCount` variable should be set to the result of `getConsonantCount("Coding is fun")`.

```js
assert.equal(consonantCount, getConsonantCount("Coding is fun"));
```

You should log the `consonantCount` variable to the console with the parameter `"Consonant Count: ${consonantCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Consonant\s+Count:\s+('|"|`)\s+\+\s+consonantCount|`Consonant\s+Count:\s+\${consonantCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
