---
id: 66ed41f912d0bb1dc62da5dd
title: Izgradnja bota za pitanja na JavaScriptu
challengeType: 26
dashedName: lab-javascript-trivia-bot
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da napišete `"Hello! I'm your coding fun fact guide!"` u konzoli kao poruku pozdrava za korisnika.
1. Trebalo bi da deklarišete tri parametra: `botName`, `botLocation` i `favoriteLanguage`, koji čuvaju ime bota, njegovo poreklo i njegov omiljeni programski jezik, po redosledu.
1. Trebalo bi da napišete `"My name is (botName) and I live on (botLocation)."` u konzoli.
1. Trebalo bi da napišete `"My favorite programming language is (favoriteLanguage)."` u konzoli.
1. Trebalo bi da koristite `let` za deklarisanje novog parametra `codingFact` i postavljanje vrednosti alfanumeričke sekvence koja se tiče zanimljive činjenice o omiljenom jeziku bota, uključujući korišćenje parametra `favoriteLanguage`.
1. Trebalo bi da napišete `codingFact` u konzoli.
1. Trebalo bi da postavite novu vrednost za parametar `codingFact` o drugoj činjenici o omiljenom jeziku bota, koristeći ponovo parametar `favoriteLanguage`.
1. Trebalo bi da ponovo napišete `codingFact` u konzoli.
1. Trebalo bi da ponovo postavite novu vrednost za parametar `codingFact` o drugoj činjenici o omiljenom jeziku bota, koristeći parametar `favoriteLanguage`.
1. Trebalo bi da napišete `codingFact` u konzoli treći put.
1. Trebalo bi da napišete `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` u konzoli kao poruku za oproštaj od bota.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Treba da napišete `"Hello! I'm your coding fun fact guide!"` u konzolu.

```js
assert.equal(getLogs()[0], "Hello! I'm your coding fun fact guide!")
```

Trebaš objaviti parametar `botName`. Proveri da nema greške u pravopisu ili velikih/malih slova.

```js
assert.exists(botName);
```

Tvoj parametar `botName` mora biti niz slova.

```js
assert.isString(botName);
```

Treba da objaviš parametar `botLocation`. Proveri da nema grešaka u pravopisu ili velikih/malih slova.

```js
assert.exists(botLocation);
```

Vaš parametar `botLocation` mora biti sekvenca karaktera.

```js
assert.isString(botLocation);
```

Trebalo bi da objaviš parametar `favoriteLanguage`. Proveri da nema greške u pravopisu ili velikih/malih slova.

```js
assert.exists(favoriteLanguage);
```

Vaš parametar `favoriteLanguage` treba da bude niz slova/karaktera.

```js
assert.isString(favoriteLanguage);
```

Treba da kucate u konzolu `"My name is (botName) and I live on (botLocation)."` i dodate parametre u niz znakova.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[1], `My name is ${botName} and I live on ${botLocation}.`)
```

Treba da napišete u konzolu `"My favorite programming language is (favoriteLanguage)."` i dodate parametar na niz karaktera.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(getLogs()[2], `My favorite programming language is ${favoriteLanguage}.`)
```

Trebalo bi da koristite `let` za deklarisanje novog parametra `codingFact`.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/)
```

Unapaswa kumpa `codingFact` thamani inayojumuisha `favoriteLanguage`.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.match(code, /let\s+codingFact/);
assert.match(code, /codingFact\s*=\s*(("|'|`)?.+?\1?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|'|`)?.+?\3?))/);
assert.exists(codingFact);
assert.isNotEmpty(codingFact);
assert.include(String(codingFact), favoriteLanguage);
assert.include(getLogs()[3], favoriteLanguage);
```

Treba da napišete `codingFact` u konzoli.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[3], favoriteLanguage);
assert.isAtLeast(loggingCodingFacts.length, 1);
```

Trebalo bi da postavite novu vrednost za `codingFact` koja uključuje `favoriteLanguage`, i unesete u konzolu.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
const [first, second, third] = codeWithoutComments.match(/(let )?\s*codingFact\s*=\s*(("|')?.+?\2?\s*\+?\s*|favoriteLanguage\s*\+\s*(("|')?.+?\2?))/g);
assert.include(getLogs()[4], favoriteLanguage);
assert.notEqual(getLogs()[4], getLogs()[3]);
assert.isAtLeast(loggingCodingFacts.length, 2);
assert.exists(second);
assert.isNotEmpty(codingFact); 
```

Trebalo bi da unesete vrednost za `codingFact` treći put koja uključuje `favoriteLanguage`, i napišete u konzolu.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[5], favoriteLanguage);
assert.notEqual(getLogs()[5], getLogs()[4]);
assert.equal(getLogs()[5], codingFact);
assert.lengthOf(loggingCodingFacts, 3);
assert.isNotEmpty(codingFact); 
```

Trebalo bi da kucate u konzolu `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` i dodajte vrednosti parametara.

```js
assert.equal(getLogs()[6], `It was fun sharing these facts with you. Goodbye! - ${botName} from ${botLocation}.`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
console.log("Hello! I'm your coding fun fact guide!");

const botName = "JsBot";
const botLocation = "planet Eris";
const favoriteLanguage = "JavaScript";

console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

let codingFact = "Did you know that " + favoriteLanguage + " was created in just 10 days?";

console.log(codingFact);

codingFact = "Another fun fact: " + favoriteLanguage + "was originally called Mocha!";

console.log(codingFact);

codingFact = "Also, " + favoriteLanguage + " is the most popular programming language in the world.";

console.log(codingFact);

console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");

```
