---
id: 69899123889dcf5b19be1847
title: Implementirajte algoritam pretrage u dubinu
challengeType: 26
dashedName: implement-depth-first-search-algorithm
---

# --description--

U ovom laboratorijumu, implementiraćete algoritam pretraživanja grafa koji se zove <dfn>pretraga u dubinu</dfn> (DFS).

Dok <dfn>pretraga u širinu</dfn> (BFS) istražuje susjedima nivo po nivo, <dfn>pretraga u dubinu</dfn> se spušta što je moguće dublje duž jedne staze ivica prije povratka.

Kada algoritam dođe do kraja staze (čvora bez neposjećenih susjeda), on se vraća na najnoviji čvor koji još uvijek ima neiststražene ivice i nastavlja pretragu odatle.

Da biste implementirali ovaj algoritam iterativno, želite da koristite **stogu**. U JavaScriptu, standardni `Array` savršeno služi kao stog koristeći metode `.push()` i `.pop()`, slijedeći princip **Posljednji-Ušao-Prvi-Izlazi (LIFO)**. Ovo osigurava da je susjed najnovije dodan na stog sljedeći koji će biti istraživan.

Jednostavan izlaz ovog algoritma je lista čvorova koji su dostiživi od datog čvora. Stoga, također ćete želeti da pratite čvorove koje posjećujete.

**Cilj**: Ispunite korisničke priče ispod i osigurajte da svi testovi prođu kako biste završili laboratorijum.

**Korisničke Priče:**

1. Trebali biste definisati funkciju pod nazivom `dfs`.
1. Funkcija `dfs` treba da primi dva argumenta:

    - `graph`: Adjacency matricu (niz nizova).
    - `root`: Numeričku oznaku čvora (početni indeks) koja je numerička vrijednost čvora između `0` i `n - 1`, gdje je `n` ukupan broj čvorova u grafu.

1. Funkcija treba da vrati niz koji sadrži sve oznake čvorova dostižive od početnog čvora.

# --hints--

You should have a function named `dfs` that takes two arguments.

```js
assert.isFunction(dfs);
assert.lengthOf(dfs, 2);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)` should return a list with `1`, `2`, `3`, and `0`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)` should return a list with `1`, `2`, `3`, and `0`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)` should return `[3]`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3), [3]);
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)` should return a list with `3` and `2`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3), [3, 2]);

```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)` should return a list with `0` and `1`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0), [0, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfs(graph, root) {
  let stack = [root];
  let visited = [];

  while (stack.length > 0) {
    let current = stack.pop();

    if (!visited.includes(current)) {
      visited.push(current);
      let neighbors = graph[current];
      
      for (let i = 0; i < neighbors.length; i++) {
        if (neighbors[i] === 1) {
          stack.push(i);
        }
      }
    }
  }
  return visited;
}
```
