---
id: a3bfc1673c0526e06d3ac698
title: Izgraditi kalkulator za sumu premijnih brojeva
challengeType: 26
dashedName: build-a-prime-number-sum-calculator
---

# --description--

U ovom laboratorijumu, kreiraćete kalkulator koji prima broj i vraća zbir svih premijnih brojeva koji su manji ili jednaki tom broju.

**Cilj**: Ispunite priče korisnika ispod i osigurajte da svi testovi prođu kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da imate funkciju `sumPrimes` koja prima broj kao argument.
1. Funkcija `sumPrimes` treba da vrati zbir svih premijnih brojeva manjih ili jednakih datoj broju.
1. Ako je ulazni broj manji od `2`, funkcija bi trebalo da vrati `0`.

# --hints--

You should have a `sumPrimes` function.

```js
assert.isFunction(sumPrimes);
```

`sumPrimes(10)` should return `17`.

```js
assert.strictEqual(sumPrimes(10), 17);
```

`sumPrimes(5)` should return `10`.

```js
assert.strictEqual(sumPrimes(5), 10);
```

`sumPrimes(2)` should return `2`.

```js
assert.strictEqual(sumPrimes(2), 2);
```

`sumPrimes(0)` should return `0`.

```js
assert.strictEqual(sumPrimes(0), 0);
```

`sumPrimes(977)` should return `73156`.

```js
assert.strictEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumPrimes(num) {
  function isPrime(n) {
    if (n < 2) return false;
    for (let i = 2; i <= Math.sqrt(n); i++) {
      if (n % i === 0) {
        return false;
      }
    }
    return true;
  }

  let sum = 0;

  for (let i = 2; i <= num; i++) {
    if (isPrime(i)) {
      sum += i;
    }
  }
  return sum;
}
```
