---
id: 690ddb5e6ee94eb73ed172b0
title: Ispravljanje forme za donacije
challengeType: 25
dashedName: debug-donation-form
# Purposefully removed the demo type because this is a debugging project.
---

# --description--

Lokalna humanitarna organizacija je napravila veb-sajt za formular za donacije, ali postoji nekoliko problema koje treba popraviti. Formular nije pristupačan i ima neke greške u HTML sintaksi.

Vaš zadatak je da popravite sve greške kako bi formular radio ispravno i bio dostupan svim korisnicima. Popunite stavke u pričama korisnika (user stories) ispod i kliknite na "Check Your Code" da vidite da li ste popravili sve greške.

**Priče korisnika:**

1. Elementi `input` su prazni elementi (void elements) i ne bi trebalo da imaju zatvarajuće tagove. Uklonite sve zatvarajuće tagove `</input>` iz formulara.
2. Dodajte elemente `label` za svako polje unosa formulara kako bi korisnici znali čemu služi svako polje. Tekst etikete (label) treba da odgovara onome što se trenutno nalazi pored svakog unosa.
3. Tip unosa (input type) `Email Address:` treba da bude `email` umesto `text`.
4. Trebalo bi da povežete svaki element `label` sa njegovim odgovarajućim elementom `input` koristeći atribut `for` na `label` i odgovarajući atribut `id` na `input`.
5. Dodajte atribut `required` poljima za unos teksta, e-pošte i broja (ali ne i za polje sa čekboksom ili dugme za slanje) kako biste osigurali da korisnici popune potrebne informacije.

# --hints--

You should not have any `</input>` closing tags in your code.

```js
assert.notMatch(code, /<\/input>/);
```

You should have exactly five `input` elements in your form.

```js
assert.lengthOf(document.querySelectorAll("input"), 5);
```

You should have exactly four `label` elements in your form.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

Your first `label` should have the text `Full Name:`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[0]?.textContent, /Full\s*Name\s*:/i);
```

Your first `input` should have a `required` attribute.

```js
const firstInput = document.querySelector("input[name='name']");
assert.isTrue(firstInput?.hasAttribute("required"));
```

Your second `label` should have the text `Email Address:`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[1]?.textContent, /Email\s*Address\s*:/i);
```

Your second `input` should have a `required` attribute.

```js
const secondInput = document.querySelector("input[name='email']");
assert.isTrue(secondInput?.hasAttribute("required"));
```

Your third `label` should have the text `Donation Amount ($):`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[2]?.textContent, /Donation\s*Amount\s*\(\$\)\s*:/i);
```

Your third `input` should have a `required` attribute.

```js
const thirdInput = document.querySelector("input[name='amount']");
assert.isTrue(thirdInput?.hasAttribute("required"));
```

Your fourth `label` should have the text `Subscribe`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[3]?.textContent, /Subscribe/i);
```

You should associate every `input` element with a `label` element using the `for` and `id` attributes.

```js 
const els = document.querySelectorAll('input:not([type="submit"])'); 
assert.isNotEmpty(els); 
els.forEach(input => { 
  const label = document.querySelector(`label[for="${input.id}"]`); 
  assert.exists(label);  
  assert.equal(label.tagName, "LABEL"); 
})
``` 

Your `input` for the email address should be of type `email`.

```js
const emailInput = document.querySelector("input[name='email']");
assert.strictEqual(emailInput?.getAttribute("type"), "email");
```

Your checkbox `input` should not have a `required` attribute.

```js
const checkboxInput = document.querySelector("input[name='newsletter']");
assert.isFalse(checkboxInput?.hasAttribute("required"));
```

Your submit `button` should not have a `required` attribute.

```js
const submitButton = document.querySelector("input[type='submit']");
assert.isFalse(submitButton?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    
    Full Name:
    <input type="text" name="name"></input>

    Email Address:
    <input type="text" name="email">

    Donation Amount ($):
    <input type="number" name="amount"></input>

    <input type="checkbox" name="newsletter"></input>
    Subscribe

    <input type="submit" value="Send"></input>
  </form>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    <label for="full-name">Full Name:</label>
    <input id="full-name" type="text" name="name" required>

    <label for="email">Email Address:</label>
    <input id="email" type="email" name="email" required>

    <label for="amount">Donation Amount ($):</label>
    <input id="amount" type="number" name="amount" required>

    <label for="newsletter">Subscribe</label>
    <input id="newsletter" type="checkbox" name="newsletter">

    <input type="submit" value="Send">
  </form>
</body>
</html>
```
