id: 681cb1afdab50c87ddb2e516
title: "Vežba 5: Jbelmud niz slova"
challengeType: 28
dashedName: challenge-5
---

# --description--

Kada dobijete niz slova, vratite izmešaniu tu niza gde se svaka reč menja korišćenjem sledećih pravila:

- Prvo i poslednje slovo reči ostaju na svom mestu.
- Sva slova između prvog i poslednjeg slova poređana su po abecednom redosledu.
- Niz slova unosa neće imati interpunkcijske znakove i mora biti u celom malim slovima.

# --hints--

`jbelmu("hello world")` bi trebalo da vrati `"hello wlord"`.

```js
assert.equal(jbelmu("hello world"), "hello wlord");
```

`jbelmu("i love jumbled text")` bi trebalo da vrati `"i love jbelmud text"`.

```js
assert.equal(jbelmu("i love jumbled text"), "i love jbelmud text");
```

`jbelmu("freecodecamp is my favorite place to learn to code")` bi trebalo da vrati `"faccdeeemorp is my faiortve pacle to laern to cdoe"`.

```js
assert.equal(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe");
```

`jbelmu("the quick brown fox jumps over the lazy dog")` bi trebalo da vrati `"the qciuk borwn fox jmpus oevr the lazy dog"`.

```js
assert.equal(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog");
```

# --seed--

## --seed-contents--

```js
function jbelmu(text) {

  return text;
}
```

# --solutions--

```js
function jbelmu(text) {
  return text
    .split(' ')
    .map((word) => {
      if (word.length <= 3) return word;
      const first = word[0];
      const last = word[word.length - 1];
      const middle = word
        .slice(1, -1)
        .split('')
        .sort()
        .join('');
      return first + middle + last;
    })
    .join(' ');
}
```