id: 681cb1b0dab50c87ddb2e51b
title: "Izazov 10: Tri kvadrata"
challengeType: 28
dashedName: challenge-10
---

# --description--

Ako vam je dat ceo broj između 1 i 10,000, vratite brojač koliko brojeva od 1 do tog broja čiji kvadrat sadrži najmanje jednu cifru 3.

# --hints--

`squaresWithThree(1)` treba da vrati `0`.

```js
assert.equal(squaresWithThree(1), 0);
```

`squaresWithThree(10)` treba da vrati `1`.

```js
assert.equal(squaresWithThree(10), 1);
```

`squaresWithThree(100)` treba da vrati `19`.

```js
assert.equal(squaresWithThree(100), 19);
```

`squaresWithThree(1000)` treba da vrati `326`.

```js
assert.equal(squaresWithThree(1000), 326);
```

`squaresWithThree(10000)` treba da vrati `4531`.

```js
assert.equal(squaresWithThree(10000), 4531);
```

# --seed--

## --seed-contents--

```js
function squaresWithThree(n) {

  return n;
}
```

# --solutions--

```js
function squaresWithThree(n) {
  let count = 0;

  for (let i = 1; i <= n; i++) {
    const square = i * i;
    if (square.toString().includes('3')) {
      count++;
    }
  }

  return count;
}
```