id: 681cb1b0dab50c87ddb2e51a
title: "Vežba 9: Zbir kvadrata"
challengeType: 28
dashedName: challenge-9
---

# --description--

Ako vam je dat pozitivan ceo broj do 1.000, vratite zbir svih celih brojeva podignutih na kvadrat od 1 do tog broja.

# --hints--

`sumOfSquares(5)` treba da vrati `55`.

```js
assert.equal(sumOfSquares(5), 55);
```

`sumOfSquares(10)` treba da vrati `385`.

```js
assert.equal(sumOfSquares(10), 385);
```

`sumOfSquares(25)` treba da vrati `5525`.

```js
assert.equal(sumOfSquares(25), 5525);
```

`sumOfSquares(500)` treba da vrati `41791750`.

```js
assert.equal(sumOfSquares(500), 41791750);
```

`sumOfSquares(1000)` treba da vrati `333833500`.

```js
assert.equal(sumOfSquares(1000), 333833500);
```

# --seed--

## --seed-contents--

```js
function sumOfSquares(n) {

  return n;
}
```

# --solutions--

```js
function sumOfSquares(n) {
  let total = 1;
  for(let i = 2; i <= n; i++) {
      total += i * i;
  }
  return total;
}
```