id: 681cb1b0dab50c87ddb2e518
title: "Vežba 7: Cilkovita suma"
challengeType: 28
dashedName: challenge-7
---

# --description--

Ako dobijete niz podataka brojeva i ciljnu sumu, pronađite dva različita broja u nizu čija je suma jednaka vrednosti cilja. Vratite niz sa indeksima tih dva broja, ili `"Target not found"` ako ne postoji nijedna par brojeva koji čini ciljnu sumu.

- Niz podataka koji se vraća mora imati indekse u rastućem redosledu.

# --hints--

`findTarget([2, 7, 11, 15], 9)` bi trebalo da vrati `[0, 1]`.

```js
assert.deepEqual(findTarget([2, 7, 11, 15], 9), [0, 1]);
```

`findTarget([3, 2, 4, 5], 6)` bi trebalo da vrati `[1, 2]`.

```js
assert.deepEqual(findTarget([3, 2, 4, 5], 6), [1, 2]);
```

`findTarget([1, 3, 5, 6, 7, 8], 15)` bi trebalo da vrati `[4, 5]`.

```js
assert.deepEqual(findTarget([1, 3, 5, 6, 7, 8], 15), [4, 5]);
```

`findTarget([1, 3, 5, 7], 14)` bi trebalo da vrati `"Target not found"`.

```js
assert.equal(findTarget([1, 3, 5, 7], 14), "Target not found");
```

# --seed--

## --seed-contents--

```js
function findTarget(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findTarget(arr, target) {
  for (let i = 0; i < arr.length; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] + arr[j] === target) {
        return [i, j];
      }
    }
  }
  return 'Target not found'
}
```