id: 681cb1afdab50c87ddb2e517
title: "Vežba 6: Proveri anagrame"
challengeType: 28
dashedName: challenge-6
---

# --description--

Razmatrajući niz od dva slova, utvrdite da li su anagrame jedno drugog (imaju iste slova u bilo kom redosledu).

- Ignorišite razlike između velikih i malih slova, kao i praznim prostorima.

# --hints--

`areAnagrams("listen", "silent")` bi trebalo da vrati `true`.

```js
assert.isTrue(areAnagrams("listen", "silent"));
```

`areAnagrams("School master", "The classroom")` bi trebalo da vrati `true`.

```js
assert.isTrue(areAnagrams("School master", "The classroom"));
```

`areAnagrams("A gentleman", "Elegant man")` bi trebalo da vrati `true`.

```js
assert.isTrue(areAnagrams("A gentleman", "Elegant man"));
```

`areAnagrams("Hello", "World")` bi trebalo da vratiti `false`.

```js
assert.isFalse(areAnagrams("Hello", "World"));
```

`areAnagrams("apple", "banana")` bi trebalo da vrati `false`.

```js
assert.isFalse(areAnagrams("apple", "banana"));
```

`areAnagrams("cat", "dog")` bi trebalo da vrati `false`.

```js
assert.isFalse(areAnagrams("cat", "dog"));
```

# --seed--

## --seed-contents--

```js
function areAnagrams(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function areAnagrams(str1, str2) {
  const clean = (str) =>
    str.replace(/\s+/g, '').toLowerCase().split('').sort().join('');
    
  return clean(str1) === clean(str2);
}
```