id: 6821ebee237de8297eaee796
title: "Vežba 20: Duplikati u nizu podataka"
challengeType: 28
dashedName: challenge-20
---

# --description--

Datovan niz celih brojeva, vratite niz celih brojeva koji se pojavljuju više od jednom u originalnom nizu, sortirisan po rastućem redosledu. Ako nijedna vrednost ne ponavlja više od jednom, vratite prazan niz.

- Uključite samo jedan primer svake vrednosti u vraćeni niz.

# --hints--

`findDuplicates([1, 2, 3, 4, 5])` bi trebalo da vrati `[]`.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 5]), []);
```

`findDuplicates([1, 2, 3, 4, 1, 2])` bi trebalo da vrati `[1, 2]`.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 1, 2]), [1, 2]);
```

`findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])` bi trebalo da vrati `[-6, 0, 2, 4, 5, 23]`.

```js
assert.deepEqual(findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23]);
```

# --seed--

## --seed-contents--

```js
function findDuplicates(arr) {

  return arr;
}
```

# --solutions--

```js
function findDuplicates(arr) {
  const duplicates = [];
  
  for (let i = 0; i < arr.length - 1; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] === arr[j] && !duplicates.includes(arr[i])) {
        duplicates.push(arr[i]);
      }
    }
  }

  return duplicates.sort((a, b) => a - b);
}
```