id: 6821ebc9237de8297eaee78f
title: "Vežba 13: Neprirodan premij broj"
challengeType: 28
dashedName: challenge-13
---

# --description--

Za zadati ceo broj, utvrdite da li je taj broj premij broj ili negativan premij broj.

- Premij broj je pozitivan ceo broj veći od 1 koji se može deliti samo sa 1 i samim sobom.
- Negativan premij broj je negativna vrednost pozitivnog premija broja.
- `1` i `0` se ne smatraju premij brojevima.

# --hints--

`isUnnaturalPrime(1)` inapaswa kurudisha `false`.

```js
assert.isFalse(isUnnaturalPrime(1));
```

`isUnnaturalPrime(-1)` inapaswa kurudisha `false`.

```js
assert.isFalse(isUnnaturalPrime(-1));
```

`isUnnaturalPrime(19)` inapaswa kurudisha `true`.

```js
assert.isTrue(isUnnaturalPrime(19));
```

`isUnnaturalPrime(-23)` inapaswa kurudisha `true`.

```js
assert.isTrue(isUnnaturalPrime(-23));
```

`isUnnaturalPrime(0)` inapaswa kurudisha `false`.

```js
assert.isFalse(isUnnaturalPrime(0));
```

`isUnnaturalPrime(97)` inapaswa kurudisha `true`.

```js
assert.isTrue(isUnnaturalPrime(97));
```

`isUnnaturalPrime(-61)` inapaswa kurudisha `true`.

```js
assert.isTrue(isUnnaturalPrime(-61));
```

`isUnnaturalPrime(99)` inapaswa kurudisha `false`.

```js
assert.isFalse(isUnnaturalPrime(99));
```

`isUnnaturalPrime(-44)` inapaswa kurudisha `false`.

```js
assert.isFalse(isUnnaturalPrime(-44));
```

# --seed--

## --seed-contents--

```js
function isUnnaturalPrime(n) {

  return n;
}
```

# --solutions--

```js
function isUnnaturalPrime(n) {
  const abs = Math.abs(n);

  if (abs <= 1) return false;

  for (let i = 2; i <= Math.sqrt(abs); i++) {
    if (abs % i === 0) return false;
  }

  return true;
}
```