id: 6821ebea237de8297eaee795
title: "Vežba 19: Svetlost Sveća"
challengeType: 28
dashedName: challenge-19
---

# --description--

Dato vam je potpuno brojanje koje predstavlja broj sveća sa kojima počinjete, i potpuno brojanje koje predstavlja koliko izgorelih sveća je potrebno za napraviti novu sveću. Vratite ukupan broj sveća koje ćete potrošiti nakon što napravite i sagorete sve što možete.

Na primer, ako vam date 7 sveća, a da je potrebno 2 izgorele sveće za napravljenje nove:

1. Sagorite 7 sveća i dobijete 7 ostataka,
2. Vratite 6 ostataka u 3 nove sveće (ostaje 1 sveća),
3. Sagorite 3 sveće i dobijete još 3 ostatka (ukupno 4),
4. Vratite 4 ostatka u 2 nove sveće,
5. Sagorite 2 sveće i dobijete 2 ostataka,
6. Vratite 2 ostatka u 1 novu sveću,
7. Sagorite 1 sveću.

Ukupno je sagorelo 13 sveća u ovom primeru.

# --hints--

`burnCandles(7, 2)` should return `13`

```js
assert.equal(burnCandles(7, 2), 13);
```

`burnCandles(10, 5)` should return `12`

```js
assert.equal(burnCandles(10, 5), 12);
```

`burnCandles(20, 3)` should return `29`

```js
assert.equal(burnCandles(20, 3), 29);
```

`burnCandles(17, 4)` should return `22`

```js
assert.equal(burnCandles(17, 4), 22);
```

`burnCandles(2345, 3)` should return `3517`

```js
assert.equal(burnCandles(2345, 3), 3517);
```

# --seed--

## --seed-contents--

```js
function burnCandles(candles, leftoversNeeded) {

  return candles;
}
```

# --solutions--

```js
function burnCandles(candles, leftoversNeeded) {
  let totalBurned = 0;
  let unusedLeftovers = 0;

  while (candles > 0) {
    totalBurned += candles;
    const leftovers = candles + unusedLeftovers;
    candles = Math.floor(leftovers / leftoversNeeded);
    unusedLeftovers = leftovers % leftoversNeeded;
  }

  return totalBurned;
}
```