id: 6821ebce237de8297eaee790
title: "Vežba 14: Bitka slova"
challengeType: 28
dashedName: challenge-14
---

# --description--

Dva niza slova predstavljaju vašu vojsku i protivničku vojsku. Svako slovo iz vaše vojske bori se sa slovom na istoj poziciji iz protivničke vojske koristeći sledeća pravila:

- Slova `a-z` imaju snagu od `1-26`, po redosledu.
- Slova `A-Z` imaju snagu od `27-52`, po redosledu.
- Brojevi `0-9` imaju snagu jednaku njihovoj nominalnoj vrednosti.
- Sva druga slova imaju vrednost nula.
- Svako slovo može učestvovati u samo jednoj bitci.

U svakoj bitci pobeđuje slovo sa većom snagom. Vojska koja ima više pobeda osvaja rat. Vratite sledeće vrednosti:

- `"Opponent retreated"` ako vaša vojska ima više slova od protivničke vojske.
- `"We retreated"` ako protivnička vojska ima više slova od vaše.
- `"We won"` ako je vaša vojska osvojila više bitaka.
- `"We lost"` ako je protivnička vojska osvojila više bitaka.
- `"It was a tie"` ako su obe vojske osvojile isti broj bitaka.

# --hints--

`battle("Hello", "World")` bi trebalo da vrati `"We lost"`.

```js
assert.equal(battle("Hello", "World"), "We lost");
```

`battle("pizza", "salad")` bi trebalo da vrati `"We won"`.

```js
assert.equal(battle("pizza", "salad"), "We won");
```

`battle("C@T5", "D0G$")` bi trebalo da vrati `"We won"`.

```js
assert.equal(battle("C@T5", "D0G$"), "We won");
```

`battle("kn!ght", "orc")` bi trebalo da vrati `"Opponent retreated"`.

```js
assert.equal(battle("kn!ght", "orc"), "Opponent retreated");
```

`battle("PC", "Mac")` bi trebalo da vrati `"We retreated"`.

```js
assert.equal(battle("PC", "Mac"), "We retreated");
```

`battle("Wizards", "Dragons")` bi trebalo da vrati `"It was a tie"`.

```js
assert.equal(battle("Wizards", "Dragons"), "It was a tie");
```

`battle("Mr. Smith", "Dr. Jones")` bi trebalo da vrati `"It was a tie"`.

```js
assert.equal(battle("Mr. Smith", "Dr. Jones"), "It was a tie");
```

# --seed--

## --seed-contents--

```js
function battle(myArmy, opposingArmy) {

  return myArmy;
}
```

# --solutions--

```js
function getStrength(soldier) {
  const soldiers = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ';
  let strength = 0;

  if (/\d/.test(soldier)) {
    strength = parseInt(soldier);
  } else if (soldiers.includes(soldier)) {
    strength = soldiers.indexOf(soldier) + 1;
  }

  return strength;
}

function battle(myArmy, opposingArmy) {
  if (myArmy.length > opposingArmy.length) return 'Opponent retreated';
  if (opposingArmy.length > myArmy.length) return 'We retreated';

  let myWins = 0;
  let theirWins = 0;

  for (let i = 0; i < myArmy.length; i++) {
    const mySoldier = myArmy[i];
    const theirSoldier = opposingArmy[i];

    const myStrength = getStrength(mySoldier);
    const theirStrength = getStrength(theirSoldier);

    if (myStrength > theirStrength) myWins++;
    if (theirStrength > myStrength) theirWins++;
  }

  return myWins > theirWins ? 'We won' : theirWins > myWins ? 'We lost' : 'It was a tie';
}
```