id: 68216ef80f957572e7c340c5
title: "Vežba 12: Dekoder poruka"
challengeType: 28
dashedName: challenge-12
---

# --description--

Ako dobijete sekvencu slova iz tajne poruke i ceo broj koji predstavlja broj slova korišćenih za pomeranje poruke kako bi je sakrili, vratite dešifrovisanu sekvencu slova.

- Pozitivan broj znači da je poruka pomerena napred u alfabeti.
- Negativan broj znači da je poruka pomerena nazad u alfabeti.
- Razlika između velikih i malih slova je važna; dešifrovisana slova moraju sačuvati veličinu (velika ili mala) kao što su bila u originalnoj poruci.
- Nealfabetički znakovi ne bi trebalo dešifrovati.

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` bi trebalo da vrati `"This is a secret message."`

```js
assert.equal(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.");
```

`decode("Byffi Qilfx!", 20)` bi trebalo da vrati `"Hello World!"`

```js
assert.equal(decode("Byffi Qilfx!", 20), "Hello World!");
```

`decode("Zqd xnt njzx?", -1)` bi trebalo da vrati `"Are you okay?"`

```js
assert.equal(decode("Zqd xnt njzx?", -1), "Are you okay?");
```

`decode("oannLxmnLjvy", 9)` bi trebalo da vrati `"freeCodeCamp"`

```js
assert.equal(decode("oannLxmnLjvy", 9), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(message, shift) {

  return message;
}
```

# --solutions--

```js
function decode(message, shift) {
  return message.split('').map(char => {
    if (/[a-zA-Z]/.test(char)) {
      const base = char === char.toLowerCase() ? 'a'.charCodeAt(0) : 'A'.charCodeAt(0);
      const charCode = char.charCodeAt(0);
      const offset = (charCode - base - shift + 26) % 26;
      return String.fromCharCode(base + offset);
    } else {
      return char;
    }
  }).join('');
}
```