id: 6821ebd4237de8297eaee791
title: "Vežba 15: camelCase"
challengeType: 28
dashedName: challenge-15
---

# --description--

Izvući niz slova i vrati ga u format camelCase koristeći sledeće pravilo:

- Reči u ulazu (nizu slova) su razdvojene jednim ili više znakova iz sledećeg skupa: razmak (` `), crtica (`-`) ili donji vek (`_`). Svaka sekvenca ovih znakova se tretira kao granica reči.
- Prva reč mora biti u malim slovima (lowercase).
- Svaka sledeća reč mora početi velikim slovom, a ostatak mora biti u malim slovima.
- Svi razmaci i separatori moraju biti uklonjeni.

# --hints--

`toCamelCase("hello world")` treba da vrati `"helloWorld"`.

```js
assert.equal(toCamelCase("hello world"), "helloWorld");
```

`toCamelCase("HELLO WORLD")` treba da vrati `"helloWorld"`.

```js
assert.equal(toCamelCase("HELLO WORLD"), "helloWorld");
```

`toCamelCase("secret agent-X")` treba da vrati `"secretAgentX"`.

```js
assert.equal(toCamelCase("secret agent-X"), "secretAgentX");
```

`toCamelCase("FREE cODE cAMP")` treba da vrati `"freeCodeCamp"`.

```js
assert.equal(toCamelCase("FREE cODE cAMP"), "freeCodeCamp");
```

`toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")` treba da vrati `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"`.

```js
assert.equal(toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk");
```

# --seed--

## --seed-contents--

```js
function toCamelCase(s) {

  return s;
}
```

# --solutions--

```js
function toCamelCase(s) {
  const words = s.replace(/[_\- ]+/g, ' ').split(' ');

  return words.map((word, i) => {
    if (i === 0) {
      return word.toLowerCase();
    } else {
      const tempWord = word.split('');
      return tempWord.shift().toUpperCase() + tempWord.join('').toLowerCase();
    }
  }).join('')
}
```