id: 6821ebda237de8297eaee792
title: "Vežba 16: Dekodiranje zagrada"
challengeType: 28
dashedName: challenge-16
---

# --description--

Ako naiđete na niz znakova koji sadrži pravilno ugnježđene zagrade, dekodirajte rezultujući niz koristeći sledeća pravila:

- Sva slova unutar svakog para zagrada moraju biti obrnuta.
- Zagrade se moraju ukloniti iz konačnog rezultata.
- Ako postoje ugnježđene zagrade, najunutrašnji par mora prvo biti obrnut, a zatim njegov rezultat uključiti u obrnjavanje spoljašnjeg para.
- Pretpostavite da su sve zagrade pravilno uravnotežene i tačne. 

# --hints--

`decode("(f(b(dc)e)a)")` bi trebalo da vrati `"abcdef"`.

```js
assert.equal(decode("(f(b(dc)e)a)"), "abcdef");
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` bi trebalo da vrati `"Can you read this?"`.

```js
assert.equal(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?");
```

`decode("f(Ce(re))o((e(aC)m)d)p")` bi trebalo da vrati `"freeCodeCamp"`.

```js
assert.equal(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(s) {

  return s;
}
```

# --solutions--

```js
function decode(s) {
  while (s.includes(')')) {
    const closeIndex = s.indexOf(')');
    const openIndex = s.lastIndexOf('(', closeIndex);
    const before = s.slice(0, openIndex);
    const group = s.slice(openIndex + 1, closeIndex).split('').reverse().join('');
    const after = s.slice(closeIndex + 1);
    s = before + group + after;
  }

  return s;
}
```