id: 6821ebfd237de8297eaee799
title: "Vežba 23: RGB u Hex"
challengeType: 28
dashedName: challenge-23
---

# --description--

Kada je CSS niz boja `rgb(r, g, b)`, vratite njegov ekvivalent u heksadecimalnom formatu.

Evo nekoliko primera rezultata za određeni unos:

| Unos   | Rezultat   |
|---------|----------|
| `"rgb(255, 255, 255)"`| `"#ffffff"` |
| `"rgb(1, 2, 3)"` | `"#010203"` |

- Sve slova pretvorite u mala.
- Vratite `#` praćen šest slova. Nemojte koristiti skraćene vrednosti.

# --hints--

`rgbToHex("rgb(255, 255, 255)")` bi trebalo da vrati `"#ffffff"`.

```js
assert.equal(rgbToHex("rgb(255, 255, 255)"), "#ffffff");
```

`rgbToHex("rgb(1, 11, 111)")` bi trebalo da vrati `"#010b6f"`.

```js
assert.equal(rgbToHex("rgb(1, 11, 111)"), "#010b6f");
```

`rgbToHex("rgb(173, 216, 230)")` bi trebalo da vrati `"#add8e6"`.

```js
assert.equal(rgbToHex("rgb(173, 216, 230)"), "#add8e6");
```

`rgbToHex("rgb(79, 123, 201)")` bi trebalo da vrati `"#4f7bc9"`.

```js
assert.equal(rgbToHex("rgb(79, 123, 201)"), "#4f7bc9");
```

# --seed--

## --seed-contents--

```js
function rgbToHex(rgb) {

  return rgb;
}
```

# --solutions--

```js
function rgbToHex(rgb) {
  const match = rgb.match(/\d+/g);
  const [r, g, b] = match.map(num =>
    Math.max(0, Math.min(255, parseInt(num)))
      .toString(16)
      .padStart(2, '0')
  );

  return `#${r}${g}${b}`;
}
```