id: 6821ebdf237de8297eaee793
title: "Vežba 17: Ignorisanje redosleda operacija"
challengeType: 28
dashedName: challenge-17
---

# --description--

Dato vam je niz celih brojeva i niz string operatora. Primijenite operacije na brojeve jedan po jedan, od lijeva ka desno. Ponovite operacije koliko je potrebno sve dok se ne koriste svi brojevi. Vratite konačni rezultat.

Na primjer, dato `[1, 2, 3, 4, 5]` i `['+', '*']`, vratite rezultat evaluacije `1 + 2 * 3 + 4 * 5` od lijeva ka desno, bez obzira na standardni redosled operacija.

- Važeći operatori su `+`, `-`, `*`, `/`, i `%`.

# --hints--

`evaluate([5, 6, 7, 8, 9], ['+', '-'])` bi trebalo da vrati `3`

```js
assert.equal(evaluate([5, 6, 7, 8, 9], ['+', '-']), 3);
```

`evaluate([17, 61, 40, 24, 38, 14], ['+', '%'])` bi trebalo da vrati `38`

```js
assert.equal(evaluate([17, 61, 40, 24, 38, 14], ['+', '%']), 38);
```

`evaluate([20, 2, 4, 24, 12, 3], ['*', '/'])` bi trebalo da vrati `60`

```js
assert.equal(evaluate([20, 2, 4, 24, 12, 3], ['*', '/']), 60);
```

`evaluate([11, 4, 10, 17, 2], ['*', '*', '%'])` bi trebalo da vrati `30`

```js
assert.equal(evaluate([11, 4, 10, 17, 2], ['*', '*', '%']), 30);
```

`evaluate([33, 11, 29, 13], ['/', '-'])` bi trebalo da vrati `-2`

```js
assert.equal(evaluate([33, 11, 29, 13], ['/', '-']), -2);
```

# --seed--

## --seed-contents--

```js
function evaluate(numbers, operators) {

  return numbers;
}
```

# --solutions--

```js
function doMath(a, b, operator) {
  switch (operator) {
    case "+":
      return a + b;
    case "-":
      return a - b;
    case "*":
      return a * b;
    case "/":
      return a / b;
    default:
      return a % b;
  }
}

function evaluate(numbers, operators) { 
  let total = numbers[0];

  for (let i = 1; i < numbers.length; i++) {
    const operator = operators[(i - 1) % operators.length];
    total = doMath(total, numbers[i], operator);
  }

  return total;
}
```