id: 6821ebf8237de8297eaee798
title: "Vežba 22: Tribonacci niz"
challengeType: 28
dashedName: challenge-22
---

# --description--

Tribonacci niz je sekvenca brojeva gde je svaki broj zbir prethodnih tri broja. Počevši od `0`, `0` i `1`, prvih deset brojeva u nizu su `0`, `0`, `1`, `1`, `2`, `4`, `7`, `13`, `24`, `44`.

Dato niz sa prvih tri broja Tribonacci sekvence, i ceo broj koji predstavlja dužinu sekvence, vratite dato niz sa sekvencijom zadate dužine.

- Vaša funkcija treba da rukuje sa sekvencama bilo koje dužine veće ili jednake nuli.
- Ako je dužina nula, vratite prazan niz.
- Zapamtite da su početni brojevi deo sekvence.

# --hints--

`tribonacciSequence([0, 0, 1], 20)` bi trebalo da vrati `[0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]`.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 20), [0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]);
```

`tribonacciSequence([21, 32, 43], 1)` bi trebalo da vrati `[21]`.


```js
assert.deepEqual(tribonacciSequence([21, 32, 43], 1), [21]);
```

`tribonacciSequence([0, 0, 1], 0)` bi trebalo da vrati `[]`.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 0), []);
```

`tribonacciSequence([10, 20, 30], 2)` bi trebalo da vrati `[10, 20]`.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 2), [10, 20]);
```

`tribonacciSequence([10, 20, 30], 3)` bi trebalo da vrati `[10, 20, 30]`.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 3), [10, 20, 30]);
```

`tribonacciSequence([123, 456, 789], 8)` bi trebalo da vrati `[123, 456, 789, 1368, 2613, 4770, 8751, 16134]`.

```js
assert.deepEqual(tribonacciSequence([123, 456, 789], 8), [123, 456, 789, 1368, 2613, 4770, 8751, 16134]);
```

# --seed--

## --seed-contents--

```js
function tribonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function tribonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [startSequence[0], startSequence[1]];
  if (length === 3) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2] + + sequence[sequence.length - 3];
    sequence.push(nextValue);
  }
  return sequence;
}
```