id: 68adce01c0e1144d0a90295a
title: "Vežba 27: Rotacija matrice"
challengeType: 28
dashedName: challenge-27
---

# --description--

Ako vam je data matrica (niz redova podataka), rotirajte tu matricu za 90 stepeni u smeru kazaljke na sat i vratite je. Na primer, ako dobijete `[[1, 2], [3, 4]]`, koja izgleda ovako:

| 1 | 2 |
|---|---|
| 3 | 4 |

Trebalo bi da vratite `[[3, 1], [4, 2]]`, koja izgleda ovako:

| 3 | 1 |
|---|---|
| 4 | 2 |

# --hints--

Funkcija `rotate([[1]])` treba da vrati `[[1]]`.

```js
assert.deepEqual(rotate([[1]]), [[1]]);
```

Funkcija `rotate([[1, 2], [3, 4]])` treba da vrati `[[3, 1], [4, 2]]`.

```js
assert.deepEqual(rotate([[1, 2], [3, 4]]), [[3, 1], [4, 2]]);
```

Funkcija `rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]])` treba da vrati `[[7, 4, 1], [8, 5, 2], [9, 6, 3]]`.

```js
assert.deepEqual(rotate([[1, 2, 3], [4, 5, 6], [7, 8, 9]]), [[7, 4, 1], [8, 5, 2], [9, 6, 3]]);
```

Funkcija `rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]])` treba da vrati `[[0, 1, 0], [0, 0, 1], [0, 1, 0]]`.

```js
assert.deepEqual(rotate([[0, 1, 0], [1, 0, 1], [0, 0, 0]]), [[0, 1, 0], [0, 0, 1], [0, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function rotate(matrix) {

  return matrix;
}
```

# --solutions--

```js
function rotate(matrix) {
  const n = matrix.length;
  const result = Array.from({ length: n }, () => Array(n).fill(0));

  for (let i = 0; i < n; i++) {
    for (let j = 0; j < n; j++) {
      result[j][n - 1 - i] = matrix[i][j];
    }
  }

  return result;
}
```