id: 68adce01c0e1144d0a90295e
title: "Vežba 29: Generator akronima"
challengeType: 28
dashedName: challenge-29
---

# --description--

Ako vam je dat niz slova koji sadrži jedno ili više reči, vratite akronime tih reči koristeći sledeće uslove:

- Akronimi moraju biti velika slova, osim ako je drugačije navedeno.
- Akronimi ovih sledećih reči se ne smeju izostaviti osim ako nisu prva reč u datom nizu: `a`, `for`, `an`, `and`, `by`, i `of`.
- Akronimi moraju biti vraćeni u redosledu kojim su dati.
- Akronimi ne smeju imati razmak između njih.

# --hints--

`buildAcronym("Search Engine Optimization")` treba da vrati `"SEO"`.

```js
assert.equal(buildAcronym("Search Engine Optimization"), "SEO");
```

`buildAcronym("Frequently Asked Questions")` treba da vrati `"FAQ"`.

```js
assert.equal(buildAcronym("Frequently Asked Questions"), "FAQ");
```

`buildAcronym("National Aeronautics and Space Administration")` treba da vrati `"NASA"`.

```js
assert.equal(buildAcronym("National Aeronautics and Space Administration"), "NASA");
```

`buildAcronym("Federal Bureau of Investigation")` treba da vrati `"FBI"`.

```js
assert.equal(buildAcronym("Federal Bureau of Investigation"), "FBI");
```

`buildAcronym("For your information")` treba da vrati `"FYI"`.

```js
assert.equal(buildAcronym("Light Amplification by Stimulated Emission of Radiation"), "LASER");
```

`buildAcronym("By the way")` treba da vrati `"BTW"`.

```js
assert.equal(buildAcronym("By the way"), "BTW");
```

`buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily")` treba da vrati `"AUHWPOTIMSH"`.

```js
assert.equal(buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH");
```

# --seed--

## --seed-contents--

```js
function buildAcronym(str) {

  return str;
}
```

# --solutions--

```js
function buildAcronym(str) {
  const smallWords = ["a", "for", "an", "and", "by", "of"];
  const words = str.split(" ");
  let acronym = "";

  for (let i = 0; i < words.length; i++) {
    const word = words[i];
    const lowerWord = word.toLowerCase();

    if (i === 0 || !smallWords.includes(lowerWord)) {
      acronym += word[0].toUpperCase();
    }
  }

  return acronym;
}
```