id: 68adce01c0e1144d0a902956
title: "Vežba 25: Ponavljanje samoglasnika"
challengeType: 28
dashedName: challenge-25
---

# --description--

Kada dobijete niz znakova, vratite novu verziju tog niza gde je svaki samoglasnik ponovljen za jedan više puta nego prethodni susrećeni samoglasnik. Na primer, prvi samoglasnik u rečenici se ne menja. Drugi samoglasnik pojavljuje se dva puta zaredom. Treći samoglasnik pojavljuje se tri puta zaredom, i tako dalje.

- Slova `a`, `e`, `i`, `o`, i `u`, bilo da su velika ili mala slova, smatraju se samoglasnicima.
- Originalni slučaj (veličina) samoglasnika mora biti sačuvan.
- Ponavljajući samoglasnici moraju biti mala slova.
- Svi ne-samoglasnički znakovi moraju zadržati svoj originalni veličinski oblik.

# --hints--

`repeatVowels("hello world")` treba da vrati `"helloo wooorld"`.

```js
assert.equal(repeatVowels("hello world"), "helloo wooorld");
```

`repeatVowels("freeCodeCamp")` treba da vrati `"freeeCooodeeeeCaaaaamp"`.

```js
assert.equal(repeatVowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp");
```

`repeatVowels("AEIOU")` treba da vrati `"AEeIiiOoooUuuuu"`.

```js
assert.equal(repeatVowels("AEIOU"), "AEeIiiOoooUuuuu");
```

`repeatVowels("I like eating ice cream in Iceland")` treba da vrati `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`.

```js
assert.equal(repeatVowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand");
```

# --seed--

## --seed-contents--

```js
function repeatVowels(str) {

  return str;
}
```

# --solutions--

```js
function repeatVowels(str) {
  const vowels = "aeiouAEIOU";
  let count = 0;
  let result = "";

  for (let char of str) {
    result += char;
    if (vowels.includes(char)) {
      result += char.repeat(count).toLowerCase();
      count++;
    }
  }

  return result;
}
```