id: 68adce01c0e1144d0a902958
title: "Vežba 26: Provera validne IPv4 adrese"
challengeType: 28
dashedName: challenge-26
---

# --description--

Dato je niz karaktera; utvrdite da li je validna IPv4 adresa. Validna IPv4 adresa sastoji se od četiri celobrojne broja razdvojene tačkama (`.`). Svaki broj mora ispunjavati sledeće uslove:

- Mora biti između 0 i 255, uključujući oba.
- Ne sme imati vodeće nule (primer 0 je dozvoljeno, 01 nije).
- Dozvoljeni znakovi su samo cifre.

# --hints--

`isValidIPv4("192.168.1.1")` should return `true`.

```js
assert.isTrue(isValidIPv4("192.168.1.1"));
```

`isValidIPv4("0.0.0.0")` should return `true`.

```js
assert.isTrue(isValidIPv4("0.0.0.0"));
```

`isValidIPv4("255.01.50.111")` should return `false`.

```js
assert.isFalse(isValidIPv4("255.01.50.111"));
```

`isValidIPv4("255.00.50.111")` should return `false`.

```js
assert.isFalse(isValidIPv4("255.00.50.111"));
```

`isValidIPv4("256.101.50.115")` should return `false`.

```js
assert.isFalse(isValidIPv4("256.101.50.115"));
```

`isValidIPv4("192.168.101.")` should return `false`.

```js
assert.isFalse(isValidIPv4("192.168.101."));
```

`isValidIPv4("192168145213")` should return `false`.

```js
assert.isFalse(isValidIPv4("192168145213"));
```

# --seed--

## --seed-contents--

```js
function isValidIPv4(ipv4) {

  return ipv4;
}
```

# --solutions--

```js
function isValidIPv4(ipv4) {
  const parts = ipv4.split(".");
  
  if (parts.length !== 4) return false;

  for (let part of parts) {
    if (!/^\d+$/.test(part)) return false;    

    const num = Number(part);
    if (num < 0 || num > 255) return false;

    if (part.length > 1 && part.startsWith("0")) return false;
  }

  return true;
}
```