id: 68af0687ef34c76c28ffa54d
title: "Vežba 33: Vreme gledanja ekrana"
challengeType: 28
dashedName: challenge-33
---

# --description--

Dato je niz sedam celobrojnih broja koji predstavljaju vreme jedne nedelje, gde svaki ceo broj predstavlja broj sati koje ste koristili na telefonu tog dana. Odredite da li je to vreme gledanja ekrana previše u odnosu na sledeće kriterijume:

- Ako bilo koji dan ima 10 sati ili više, to je previše vremena.
- Ako je prosečna vrednost tri uzastopna dana veća ili jednaka 8 sati, to je previše vremena.
- Ako je prosečna vrednost sedam dana veća ili jednaka 6 sati, to je previše vremena.

# --hints--

`tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7])` treba da vrati `false`.

```js
assert.isFalse(tooMuchScreenTime([1, 2, 3, 4, 5, 6, 7]));
```

`tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3])` treba da vrati `false`.

```js
assert.isFalse(tooMuchScreenTime([7, 8, 8, 4, 2, 2, 3]));
```

`tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6])` treba da vrati `false`.

```js
assert.isFalse(tooMuchScreenTime([5, 6, 6, 6, 6, 6, 6]));
```

`tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4])` treba da vrati `true`.

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 11, 1, 3, 4]));
```

`tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0])` treba da vrati `true`.

```js
assert.isTrue(tooMuchScreenTime([1, 2, 3, 10, 2, 1, 0]));
```

`tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4])` treba da vrati `true`.

```js
assert.isTrue(tooMuchScreenTime([3, 3, 5, 8, 8, 9, 4]));
```

`tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6])` treba da vrati `true`.

```js
assert.isTrue(tooMuchScreenTime([3, 9, 4, 8, 5, 7, 6]));
```

# --seed--

## --seed-contents--

```js
function tooMuchScreenTime(hours) {

  return hours;
}
```

# --solutions--

```js
function tooMuchScreenTime(hours) {
  for (let h of hours) {
    if (h >= 10) return true;
  }

  for (let i = 0; i <= 4; i++) {
    const avg3 = (hours[i] + hours[i+1] + hours[i+2]) / 3;
    if (avg3 >= 8) return true;
  }

  const weeklyAvg = hours.reduce((sum, h) => sum + h, 0) / 7;
  if (weeklyAvg >= 6) return true;

  return false;
}
```