id: 68af0687ef34c76c28ffa54f
title: "Vežba 34: Nedostajuće brojeve"
challengeType: 28
dashedName: challenge-34
---

# --description--

Ako vam je dat niz celih brojeva od 1 do `n`, uključujući, vratite niz svih celih brojeva koji nedostaju između 1 i `n` (gde je `n` najveći broj u datoj niti).

- Data nit može biti neuređena i može sadržati duplikate.
- Vraćena data nit mora biti u rastućem redosledu.
- Ako nema nedostajućih celih brojeva, vratite prazan niz.

# --hints--

`findMissingNumbers([1, 3, 5])` treba vratiti `[2, 4]`.

```js
assert.deepEqual(findMissingNumbers([1, 3, 5]), [2, 4]);
```

`findMissingNumbers([1, 2, 3, 4, 5])` treba vratiti `[]`.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 10])` treba vratiti `[2, 3, 4, 5, 6, 7, 8, 9]`.

```js
assert.deepEqual(findMissingNumbers([1, 10]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([10, 1, 10, 1, 10, 1])` treba vratiti `[2, 3, 4, 5, 6, 7, 8, 9]`.

```js
assert.deepEqual(findMissingNumbers([10, 1, 10, 1, 10, 1]), [2, 3, 4, 5, 6, 7, 8, 9]);
```

`findMissingNumbers([3, 1, 4, 1, 5, 9])` treba vratiti `[2, 6, 7, 8]`.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5]), []);
```

`findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4])` treba vratiti `[11]`.

```js
assert.deepEqual(findMissingNumbers([1, 2, 3, 4, 5, 7, 8, 9, 10, 12, 6, 8, 9, 3, 2, 10, 7, 4]), [11]);
```

# --seed--

## --seed-contents--

```js
function findMissingNumbers(arr) {

  return arr;
}
```

# --solutions--

```js
function findMissingNumbers(arr) {
  if (arr.length === 0) return [];
  
  const n = Math.max(...arr);
  const set = new Set(arr);
  const result = [];

  for (let i = 1; i <= n; i++) {
    if (!set.has(i)) {
      result.push(i);
    }
  }
  return result;
}
```