id: 68af0687ef34c76c28ffa54b
title: "Vežba 32: Obratite rečenicu"
challengeType: 28
dashedName: challenge-32
---

# --description--

Dato je niz reči; vratite novi niz reči u obrnutom redosledu. Na primer, prva reč bi trebalo da bude na kraju vraćenog niza, a poslednja reč bi trebalo da bude na početku vraćenog niza.

- U datoj sekvenci, reči mogu biti razdvojene jednim ili više razmaka.
- Vraćeni niz mora imati samo jedan razmak između reči.

# --hints--

`reverseSentence("world hello")` bi trebalo da vrati `"hello world"`.

```js
assert.equal(reverseSentence("world hello"), "hello world");
```

`reverseSentence("push commit git")` bi trebalo da vrati `"git commit push"`.

```js
assert.equal(reverseSentence("push commit git"), "git commit push");
```

`reverseSentence("npm  install  sudo")` bi trebalo da vrati `"sudo install npm"`.

```js
assert.equal(reverseSentence("npm  install   apt    sudo"), "sudo apt install npm");
```

`reverseSentence("import    default   function  export")` bi trebalo da vrati `"export function default import"`.

```js
assert.equal(reverseSentence("import    default   function  export"), "export function default import");
```

# --seed--

## --seed-contents--

```js
function reverseSentence(sentence) {

  return sentence;
}
```

# --solutions--

```js
function reverseSentence(sentence) {
  return sentence.split(/\s+/).reverse().join(" ");
}
```