---
id: 68af0687ef34c76c28ffa547
title: "Vežba 30: Jedinstvena slova"
challengeType: 28
dashedName: challenge-30
---

# --description--

Kada vam je dat niz slova, utvrdite da li su sva slova u tom nizu jedinstvena.

- Velika i mala slova treba tretirati kao različita slova.

# --hints--

`allUnique("abc")` bi trebalo da vrati `true`.

```js
assert.isTrue(allUnique("abc"));
```

`allUnique("aA")` bi trebalo da vrati `true`.

```js
assert.isTrue(allUnique("aA"));
```

`allUnique("QwErTy123!@")` bi trebalo da vrati `true`.

```js
assert.isTrue(allUnique("QwErTy123!@"));
```

`allUnique("~!@#$%^&*()_+")` bi trebalo da vrati `true`.

```js
assert.isTrue(allUnique("~!@#$%^&*()_+"));
```

`allUnique("hello")` bi trebalo da vrati `false`.

```js
assert.isFalse(allUnique("hello"));
```

`allUnique("freeCodeCamp")` bi trebalo da vrati `false`.

```js
assert.isFalse(allUnique("freeCodeCamp"));
```

`allUnique("!@#*$%^&*()aA")` bi trebalo da vrati `false`.

```js
assert.isFalse(allUnique("!@#*$%^&*()aA"));
```

# --seed--

## --seed-contents--

```js
function allUnique(str) {

  return str;
}
```

# --solutions--

```js
function allUnique(str) {
  let seen = '';
  for (let char of str) {
    if (seen.includes(char)) {
      return false;
    }
    seen += char;
  }
  return true;
}
```