id: 68b06e589bf2273243814773
title: "Vežba 37: Kapitalizacija prvog slova rečenice"
challengeType: 28
dashedName: challenge-37
---

# --description--

Ako dobijete pasus, vratite novi pasus u kojem je prvo slovo svake rečenice napisano velikim slovom.

- Sva druga slova ostaju kao što jesu.
- Rečenice mogu se završiti tačkom (`.`), jednim ili više upitnih znakova (`?`), ili jednim ili više uzvičnika (`!`).

# --hints--

`capitalize("this is a simple sentence.")` treba da vrati `"This is a simple sentence."`.

```js
assert.equal(capitalize("this is a simple sentence."), "This is a simple sentence.");
```

`capitalize("hello world. how are you?")` treba da vrati `"Hello world. How are you?"`.

```js
assert.equal(capitalize("hello world. how are you?"), "Hello world. How are you?");
```

`capitalize("i did today's coding challenge... it was fun!!")` treba da vrati `"I did today's coding challenge... It was fun!!"`.

```js
assert.equal(capitalize("i did today's coding challenge... it was fun!!"), "I did today's coding challenge... It was fun!!");
```

`capitalize("crazy!!!strange???unconventional...sentences.")` treba da vrati `"Crazy!!!Strange???Unconventional...Sentences."`.

```js
assert.equal(capitalize("crazy!!!strange???unconventional...sentences."), "Crazy!!!Strange???Unconventional...Sentences.");
```

`capitalize("there's a space before this period . why is there a space before that period ?")` treba da vrati `"There's a space before this period . Why is there a space before that period ?"`.

```js
assert.equal(capitalize("there's a space before this period . why is there a space before that period ?"), "There's a space before this period . Why is there a space before that period ?");
```

# --seed--

## --seed-contents--

```js
function capitalize(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function capitalize(paragraph) {
  return paragraph.replace(/([.!?]+)(\s*)([a-z])/g, (match, punc, spaces, char) => {
    return punc + spaces + char.toUpperCase();
  })
  .replace(/^([a-z])/, (match, char) => char.toUpperCase());
}
```