id: 68b06e589bf2273243814777
title: "Vežba 39: Napuni rezervoar"
challengeType: 28
dashedName: challenge-39
---

# --description--

Dato je veličina gorivnog rezervoara, trenutni nivo goriva i cena po galonama; vratite trošak punjenja celog rezervoara.

- `tankSize` je ukupna kapacitet rezervoara u galonima.
- `fuelLevel` je trenutna količina goriva u rezervoaru u galonima.
- `pricePerGallon` je trošak jednog galona goriva.
- Vraćena vrednost treba da bude zaokružena na dve decimale u formatu: `"$d.dd"`.

# --hints--

`costToFill(20, 0, 4.00)` treba da vrati `"$80.00"`.

```js
assert.equal(costToFill(20, 0, 4.00), "$80.00");
```

`costToFill(15, 10, 3.50)` treba da vrati `"$17.50"`.

```js
assert.equal(costToFill(15, 10, 3.50), "$17.50");
```

`costToFill(18, 9, 3.25)` treba da vrati `"$29.25"`.

```js
assert.equal(costToFill(18, 9, 3.25), "$29.25");
```

`costToFill(12, 12, 4.99)` treba da vrati `"$0.00"`.

```js
assert.equal(costToFill(12, 12, 4.99), "$0.00");
```

`costToFill(15, 9.5, 3.98)` treba da vrati `"$21.89"`.

```js
assert.equal(costToFill(15, 9.5, 3.98), "$21.89");
```

# --seed--

## --seed-contents--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {

  return tankSize;
}
```

# --solutions--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {
  let gallonsNeeded = tankSize - fuelLevel;
  let cost = gallonsNeeded * pricePerGallon;
  return `$${cost.toFixed(2)}`;
}
```