id: 68b06e589bf2273243814771
title: "Vežba 36: Termostatski regulator"
challengeType: 28
dashedName: challenge-36
---

# --description--

Dato vam je trenutna temperatura u sobi i željena temperatura, vratite niz karaktera koji pokazuje kako podesiti temperaturu u sobi prema ovim kriterijumima:

- Vratite `"heat"` ako je trenutna temperatura niža od željene temperature.
- Vratite `"cool"` ako je trenutna temperatura viša od željene temperature.
- Vratite `"hold"` ako je trenutna temperatura jednaka željenoj temperaturi.

# --hints--

`adjustThermostat(68, 72)` treba da vrati `"heat"`.

```js
assert.equal(adjustThermostat(68, 72), "heat");
```

`adjustThermostat(75, 72)` treba da vrati `"cool"`.

```js
assert.equal(adjustThermostat(75, 72), "cool");
```

`adjustThermostat(72, 72)` treba da vrati `"hold"`.

```js
assert.equal(adjustThermostat(72, 72), "hold");
```

`adjustThermostat(-20.5, -10.1)` treba da vrati `"heat"`.

```js
assert.equal(adjustThermostat(-20.5, -10.1), "heat");
```

`adjustThermostat(100, 99.9)` treba da vrati `"cool"`.

```js
assert.equal(adjustThermostat(100, 99.9), "cool");
```

`adjustThermostat(0.0, 0.0)` treba da vrati `"hold"`.

```js
assert.equal(adjustThermostat(0.0, 0.0), "hold");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(temp, target) {

  return temp;
}
```

# --solutions--

```js
function adjustThermostat(temp, target) {
  if (temp < target) {
    return "heat";
  } else if (temp > target) {
    return "cool";
  } else {
    return "hold";
  }
}
```