id: 68b06e589bf2273243814775
title: "Vežba 38: Generator sluga (slug)"
challengeType: 28
dashedName: challenge-38
---

# --description--

Ako dobijete niz karaktera, vratite odgovarajuću verziju za URL tog niza koristeći sledeće uslove:

- Svi slovi moraju biti mala.
- Slova, brojevi ili razmaci koji nisu slova, brojevi ili razmaci moraju biti uklonjeni.
- Svi razmaci moraju se zameniti URL kodom za razmak `%20`.
- Uzastopni razmaci bi trebalo da se promene u samo jedan `%20`.
- Vraćeni niz ne sme počinjati niti završavati sa `%20`.

# --hints--

`generateSlug("helloWorld")` treba da vrati `"helloworld"`.

```js
assert.equal(generateSlug("helloWorld"), "helloworld");
```

`generateSlug("hello world!")` treba da vrati `"hello%20world"`.

```js
assert.equal(generateSlug("hello world!"), "hello%20world");
```

`generateSlug(" hello-world ")` treba da vrati `"helloworld"`.

```js
assert.equal(generateSlug(" hello-world "), "helloworld");
```

`generateSlug("hello  world")` treba da vrati `"hello%20world"`.

```js
assert.equal(generateSlug("hello  world"), "hello%20world");
```

`generateSlug("  ?H^3-1*1]0! W[0%R#1]D  ")` treba da vrati `"h3110%20w0r1d"`.

```js
assert.equal(generateSlug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d");
```

# --seed--

## --seed-contents--

```js
function generateSlug(str) {

  return str;
}
```

# --solutions--

```js
function generateSlug(str) {
  let cleaned = str.replace(/[^a-zA-Z0-9 ]+/g, "");
  cleaned = cleaned.replace(/\s+/g, " ").trim();
  return cleaned.toLowerCase().replace(/ /g, "%20");
}
```