id: 68b06e589bf227324381476f
title: "Vežba 35: Frekvencija reči"
challengeType: 28
dashedName: challenge-35
---

# --description--

Datot je pasus, vratite niz tri najčešće pojavljene reči.

- Reči u pasusu su razdvojene razmakom.
- Ignorišite razlike između malih i velikih slova u datom pasusu. Na primer, tretirajte `Hello` i `hello` kao jednu reč.
- Ignorišite interpunkciju u datom pasusu. Interpunkcija uključuje zareze (`,`), tačke (`.`) i uzvičnike (`!`).
- Vraćeni niz mora sadržati sve reči u malim slovima.
- Vraćeni niz mora biti sortiran po frekvenciji opadajućim redosledom, sa najčešće pojavljenom reči na prvom mestu.

# --hints--

`getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` bi trebalo da vrati `["coding", "python", "in"]`.

```js
assert.deepEqual(getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"]);
```

`getWords("I like coding. I like testing. I love debugging!")` bi trebalo da vrati `["i", "like", "coding"]`.

```js
assert.deepEqual(getWords("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"]);
```

`getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` bi trebalo da vrati `["debug", "test", "deploy"]`.

```js
assert.deepEqual(getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"]);
```

# --seed--

## --seed-contents--

```js
function getWords(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function getWords(paragraph) {
  const cleaned = paragraph.replace(/[.,!]/g, "").toLowerCase();
  const words = cleaned.split(/\s+/);

  const freq = {};
  for (const word of words) {
    if (word) {
      freq[word] = (freq[word] || 0) + 1;
    }
  }

  const sorted = Object.keys(freq).sort((a, b) => freq[b] - freq[a]);
  return sorted.slice(0, 3);
}
```