id: 68b1f72371a5ac895ac70a02
title: "Vežba 40: Čuvanje slika"
challengeType: 28
dashedName: challenge-40
---

# --description--

Uzimajući u obzir veličinu slike u megabajtima (MB) i kapacitet hard diska u gigabajtima (GB), vratite broj slika koje hard disk može da sačuva koristeći sledeće uslove:

- 1 gigabajt je jednak 1000 megabajta.
- Vratite celobrojni broj slika koje disk može da sačuva.

# --hints--

`numberOfPhotos(1, 1)` bi trebalo da vrati `1000`.

```js
assert.equal(numberOfPhotos(1, 1), 1000);
```

`numberOfPhotos(2, 1)` bi trebalo da vrati `500`.

```js
assert.equal(numberOfPhotos(2, 1), 500);
```

`numberOfPhotos(4, 256)` bi trebalo da vrati `64000`.

```js
assert.equal(numberOfPhotos(4, 256), 64000);
```

`numberOfPhotos(3.5, 750)` bi trebalo da vrati `214285`.

```js
assert.equal(numberOfPhotos(3.5, 750), 214285);
```

`numberOfPhotos(3.5, 5.5)` bi trebalo da vrati `1571`.

```js
assert.equal(numberOfPhotos(3.5, 5.5), 1571);
```

# --seed--

## --seed-contents--

```js
function numberOfPhotos(photoSizeMb, hardDriveSizeGb) {

  return photoSizeMb;
}
```

# --solutions--

```js
function numberOfPhotos(photoSizeMb, driveSizeGb) {
  const driveSizeMb = driveSizeGb * 1000;
  return Math.floor(driveSizeMb / photoSizeMb);
}
```