id: 68b1f72371a5ac895ac70a08
title: "Vežba 43: Brojevi naspram Slova"
challengeType: 28
dashedName: challenge-43
---

# --description--

Za datoj sekvenci slova, vratite `"digits"` ako sekvenca ima više brojeva nego slova, `"letters"` ako ima više slova nego brojeva, i `"tie"` ako ima jednak broj brojeva i slova.

- Brojevi uključuju `0-9`.
- Slova uključuju `a-z` u velikim ili malim slovima.
- Ignorišite bilo koja druga slova.

# --hints--

`digitsOrLetters("abc123")` bi trebalo da vrati `"tie"`.

```js
assert.equal(digitsOrLetters("abc123"), "tie");
```

`digitsOrLetters("a1b2c3d")` bi trebalo da vrati `"letters"`.

```js
assert.equal(digitsOrLetters("a1b2c3d"), "letters");
```

`digitsOrLetters("1a2b3c4")` bi trebalo da vrati `"digits"`.

```js
assert.equal(digitsOrLetters("1a2b3c4"), "digits");
```

`digitsOrLetters("abc123!@#DEF")` bi trebalo da vrati `"letters"`.

```js
assert.equal(digitsOrLetters("abc123!@#DEF"), "letters");
```

`digitsOrLetters("H3110 W0R1D")` bi trebalo da vrati `"digits"`.

```js
assert.equal(digitsOrLetters("H3110 W0R1D"), "digits");
```

`digitsOrLetters("P455W0RD")` bi trebalo da vrati `"tie"`.

```js
assert.equal(digitsOrLetters("P455W0RD"), "tie");
```

# --seed--

## --seed-contents--

```js
function digitsOrLetters(str) {

  return str;
}
```

# --solutions--

```js
function digitsOrLetters(str) {
  let digitCount = 0;
  let letterCount = 0;

  for (let char of str) {
    if (/[0-9]/.test(char)) digitCount++;
    else if (/[a-zA-Z]/.test(char)) letterCount++;
  }

  if (digitCount > letterCount) return "digits";
  if (letterCount > digitCount) return "letters";
  return "tie";
}
```