id: 68b1f72371a5ac895ac70a06
title: "Vežba 42: Skladištenje videa"
challengeType: 28
dashedName: challenge-42
---

# --description--

Uzimajući u obzir veličinu videa, jedinicu za veličinu videa, kapacitet diska i jedinicu za disk, vratite broj videa koje disk može da skladišti koristeći sljedeće ograničenje:

- Jedinica veličine videa može biti bajti ("B"), kilobajti ("KB"), megabajti ("MB") ili gigabajti ("GB").
- Ako nije navedena nijedna od gore spomenutih jedinica za video, vratite `"Invalid video unit"`.
- Jedinica kapaciteta diska može biti gigabajti ("GB") ili terabajti ("TB").
- Ako nije navedena nijedna od gore spomenutih jedinica za disk, vratite `"Invalid drive unit"`.
- Vratite broj punih videa koje disk može da skladišti.
- Koristiti sljedeće konverzije:

| Jedinica | Ekvivalent |
|:-------:|:--------:|
| 1 B     |   1 B    |
| 1 KB    |  1000 B  |
| 1 MB    |  1000 KB |
| 1 GB    |  1000 MB |
| 1 TB    |  1000 GB |

Na primjer, ako primite `500`, `"MB"`, `100`, i `"GB"` kao argumente, odredite koliko videa od 500 MB se može sačuvati na disku od 100 GB.

# --hints--

`numberOfVideos(500, "MB", 100, "GB")` bi trebalo da vrati `200`.

```js
assert.equal(numberOfVideos(500, "MB", 100, "GB"), 200);
```

`numberOfVideos(1, "TB", 10, "TB")` bi trebalo da vrati `"Invalid video unit"`.

```js
assert.equal(numberOfVideos(1, "TB", 10, "TB"), "Invalid video unit");
```

`numberOfVideos(2000, "MB", 100000, "MB")` bi trebalo da vrati `"Invalid drive unit"`.

```js
assert.equal(numberOfVideos(2000, "MB", 100000, "MB"), "Invalid drive unit");
```

`numberOfVideos(500000, "KB", 2, "TB")` bi trebalo da vrati `4000`.

```js
assert.equal(numberOfVideos(500000, "KB", 2, "TB"), 4000);
```

`numberOfVideos(1.5, "GB", 2.2, "TB")` bi trebalo da vrati `1466`.

```js
assert.equal(numberOfVideos(1.5, "GB", 2.2, "TB"), 1466);
```

# --seed--

## --seed-contents--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {

  return videoSize;
}
```

# --solutions--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {
  const videoUnits = { "KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000 };
  const driveUnits = { "GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000 };

  if (!videoUnits[videoUnit]) return "Invalid video unit";
  if (!driveUnits[driveUnit]) return "Invalid drive unit";

  const videoBytes = videoSize * videoUnits[videoUnit];
  const driveBytes = driveSize * driveUnits[driveUnit];

  return Math.floor(driveBytes / videoBytes);
}
```