id: 68b1f72371a5ac895ac70a0a
title: "Vežba 44: Ogledalo sekvence slova"
challengeType: 28
dashedName: challenge-44
---

# --description--

Uzimajući u obzir dve sekvence slova, utvrdite da li je druga sekvenca ogledalo prve.

- Sekvenca slova se smatra ogledalom ako ima iste slova u obrnutom redosledu.
- Velika i mala slova tretirajte kao različita.
- Ignorišite sve nealfabetičke znakove.

# --hints--

`isMirror("helloworld", "helloworld")` bi trebalo da vrati `false`.

```js
assert.isFalse(isMirror("helloworld", "helloworld"));
```

`isMirror("Hello World", "dlroW olleH")` bi trebalo da vrati `true`.

```js
assert.isTrue(isMirror("Hello World", "dlroW olleH"));
```

`isMirror("RaceCar", "raCecaR")` bi trebalo da vrati `true`.

```js
assert.isTrue(isMirror("RaceCar", "raCecaR"));
```

`isMirror("RaceCar", "RaceCar")` bi trebalo da vrati `false`.

```js
assert.isFalse(isMirror("RaceCar", "RaceCar"));
```

`isMirror("Mirror", "rorrim")` bi trebalo da vrati `false`.

```js
assert.isFalse(isMirror("Mirror", "rorrim"));
```

`isMirror("Hello World", "dlroW-olleH")` bi trebalo da vrati `true`.

```js
assert.isTrue(isMirror("Hello World", "dlroW-olleH"));
```

`isMirror("Hello World", "!dlroW !olleH")` bi trebalo da vrati `true`.

```js
assert.isTrue(isMirror("Hello World", "!dlroW !olleH"));
```

# --seed--

## --seed-contents--

```js
function isMirror(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirror(str1, str2) {
  const clean1 = str1.replace(/[^a-zA-Z]/g, "");
  const clean2 = str2.replace(/[^a-zA-Z]/g, "");
  return clean1.split("").reverse().join("") === clean2;
}
```