id: 68b7687dded630607aceccb3
title: "Vežba 49: Analizator zaglavlja CSV-a"
challengeType: 28
dashedName: challenge-49
---

# --description--

Kada uzmete prvi red fajla sa vrednostima razdvojenim zarezom (CSV), vratite niz podataka koji sadrži zaglavlja kolona.

- Prvi red CSV fajla ima zaglavlja kolona razdvojena zarezom.
- Uklonite sve praznine na početku ili kraju svakog zaglavlja kolone.

# --hints--

`getHeadings("name,age,city")` bi trebalo da vrati `["name", "age", "city"]`.

```js
assert.deepEqual(getHeadings("name,age,city"), ["name", "age", "city"]);
```

`getHeadings("first name,last name,phone")` bi trebalo da vrati `["first name", "last name", "phone"]`.

```js
assert.deepEqual(getHeadings("first name,last name,phone"), ["first name", "last name", "phone"]);
```

`getHeadings("username , email , signup date ")` bi trebalo da vrati `["username", "email", "signup date"]`.

```js
assert.deepEqual(getHeadings("username , email , signup date "), ["username", "email", "signup date"]);
```

# --seed--

## --seed-contents--

```js
function getHeadings(csv) {

  return csv;
}
```

# --solutions--

```js
function getHeadings(csv) {
  return csv
    .split(",")
    .map(h => h.trim());
}
```