id: 68b7687dded630607aceccab
title: "Vežba 45: Savršeno kvadratno число"
challengeType: 28
dashedName: challenge-45
---

# --description--

Ako vam je dat ceo broj, utvrdite da li je savršeno kvadratno число.

- Broj je savršeno kvadratno ako ga možete pomnožiti sa samim sobom i dobiti taj broj. Na primer, 9 je savršeno kvadratno jer možete pomnožiti 3 sa samim sobom i dobiti 9.

# --hints--

`isPerfectSquare(9)` bi trebalo da vrati `true`.

```js
assert.isTrue(isPerfectSquare(9));
```

`isPerfectSquare(49)` bi trebalo da vrati `true`.

```js
assert.isTrue(isPerfectSquare(49));
```

`isPerfectSquare(1)` bi trebalo da vrati `true`.

```js
assert.isTrue(isPerfectSquare(1));
```

`isPerfectSquare(2)` bi trebalo da vrati `false`.

```js
assert.isFalse(isPerfectSquare(2));
```

`isPerfectSquare(99)` bi trebalo da vrati `false`.

```js
assert.isFalse(isPerfectSquare(99));
```

`isPerfectSquare(-9)` bi trebalo da vrati `false`.

```js
assert.isFalse(isPerfectSquare(-9));
```

`isPerfectSquare(0)` bi trebalo da vrati `true`.

```js
assert.isTrue(isPerfectSquare(0));
```

`isPerfectSquare(25281)` bi trebalo da vrati `true`.

```js
assert.isTrue(isPerfectSquare(25281));
```

# --seed--

## --seed-contents--

```js
function isPerfectSquare(n) {

  return n;
}
```

# --solutions--

```js
function isPerfectSquare(n) {
  if (n < 0) return false;
  const root = Math.floor(Math.sqrt(n));
  return root * root === n;
}
```