id: 68b7687dded630607aceccb1
title: "Vežba 48: Detektor spama"
challengeType: 28
dashedName: challenge-48
---

# --description--

Ako je broj telefona u formatu `"+A (BBB) CCC-DDDD"`, gdje svaka slova predstavlja brojak kao što slijedi:

- `A` predstavlja državni kod i može imati bilo koji broj cifara.
- `BBB` predstavlja gradski kod i uvijek će imati tri cifre.
- `CCC` i `DDDD` predstavljaju lokalni broj i uvijek će imati po tri i četiri cifre, respektivno.

Odredite je li riječ o spam broju na temelju sljedećih kriterija:

- Državni kod ima više od 2 cifre ili ne počinje nulom (`0`).
- Gradski kod je veći od 900 ili manji od 200.
- Zbroj prvih tri cifre lokalnog broja pojavljuje se unutar zadnje četiri cifre lokalnog broja.
- Broj sadrži istu cifru četiri ili više puta uzastopno (bez obzira na znakove formata).

# --hints--

`isSpam("+0 (200) 234-0182")` bi trebao vratiti `false`.

```js
assert.isFalse(isSpam("+0 (200) 234-0182"));
```

`isSpam("+091 (555) 309-1922")` bi trebao vratiti `true`.

```js
assert.isTrue(isSpam("+091 (555) 309-1922"));
```

`isSpam("+1 (555) 435-4792")` bi trebao vratiti `true`.

```js
assert.isTrue(isSpam("+1 (555) 435-4792"));
```

`isSpam("+0 (955) 234-4364")` bi trebao vratiti `true`.

```js
assert.isTrue(isSpam("+0 (955) 234-4364"));
```

`isSpam("+0 (155) 131-6943")` bi trebao vratiti `true`.

```js
assert.isTrue(isSpam("+0 (155) 131-6943"));
```

`isSpam("+0 (555) 135-0192")` bi trebao vratiti `true`.

```js
assert.isTrue(isSpam("+0 (555) 135-0192"));
```

`isSpam("+0 (555) 564-1987")` bi trebao vratiti `true`.

```js
assert.isTrue(isSpam("+0 (555) 564-1987"));
```

`isSpam("+00 (555) 234-0182")` bi trebao vratiti `false`.

```js
assert.isFalse(isSpam("+00 (555) 234-0182"));
```

# --seed--

## --seed-contents--

```js
function isSpam(number) {

  return number;
}
```

# --solutions--

```js
function isSpam(number) {
  const match = number.match(/^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$/);
  const [, countryCode, areaCode, ccc, dddd] = match;
  const allDigits = countryCode + areaCode + ccc + dddd;
  
  if (countryCode.length > 2 || !countryCode.startsWith("0")) return true;

  const areaNum = parseInt(areaCode, 10);
  if (areaNum > 900 || areaNum < 200) return true;

  const sumCCC = ccc.split("").reduce((a, b) => a + parseInt(b), 0);
  if (dddd.includes(sumCCC.toString())) return true;

  if (/(\d)\1\1\1/.test(allDigits)) return true;

  return false;
}
```