id: 68b7cadffed0e75a517da675
title: "Vežba 53: Konverzija pozitivnog celog broja iz decimalnog u binarni sistem"
challengeType: 28
dashedName: challenge-53
---

# --description--

Ako vam je dat pozitivan ceo broj, vratite njegov binarni zapis kao niz karaktera.

Binarni broj koristi samo cifre `0` i `1` za predstavljanje bilo kog broja. Da biste pretvorili decimalni broj u binarni, stalno delite taj broj sa `2` i beležite ostatak. Ponovite dok broj ne postane nula. Pročitajte beležene ostatke od poslednjeg do prvog. Na primer, konverzija `12` u binarni sistem:

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12` u binarnom sistemu je `1100`.

# --hints--

Funkcija `toBinary(5)` bi trebalo da vrati `"101"`.

```js
assert.equal(toBinary(5), "101");
```

Funkcija `toBinary(12)` bi trebalo da vrati `"1100"`.

```js
assert.equal(toBinary(12), "1100");
```

Funkcija `toBinary(50)` bi trebalo da vrati `"110010"`.

```js
assert.equal(toBinary(50), "110010");
```

Funkcija `toBinary(99)` bi trebalo da vrati `"1100011"`.

```js
assert.equal(toBinary(99), "1100011");
```

# --seed--

## --seed-contents--

```js
function toBinary(decimal) {

  return decimal;
}
```

# --solutions--

```js
function toBinary(decimal) {
  if (decimal === 0) return "0";
  let binary = "";
  while (decimal > 0) {
    binary = (decimal % 2) + binary;
    decimal = Math.floor(decimal / 2);
  }
  return binary;
}
```