id: 68b7cadffed0e75a517da673
title: "Vežba 52: Iz binarnog broja u decimalni"
challengeType: 28
dashedName: challenge-52
---

# --description--

Ako vam je dat niz znakova koji predstavlja binarni broj, vratite njegov ekvivalent u decimalnom sistemu kao broj.

Binarni broj koristi samo cifre `0` i `1` za predstavljanje bilo kog broja. Da biste pretvorili binarno u decimalno, pomnožite svaku cifru sa stepenom `2`, a zatim ih saberi. Počnite množenjem najdesne cifre sa `2^0`, sledeće cifre nalevo sa `2^1` i tako dalje. Kada se sve cifre pomnože sa stepenima `2`, saberi rezultate.

Na primer, binarni broj `101` je jednak `5` u decimalnom sistemu jer:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`toDecimal("101")` bi trebalo da vrati `5`.

```js
assert.equal(toDecimal("101"), 5);
```

`toDecimal("1010")` bi trebalo da vrati `10`.

```js
assert.equal(toDecimal("1010"), 10);
```

`toDecimal("10010")` bi trebalo da vrati `18`.

```js
assert.equal(toDecimal("10010"), 18);
```

`toDecimal("1010101")` bi trebalo da vrati `85`.

```js
assert.equal(toDecimal("1010101"), 85);
```

# --seed--

## --seed-contents--

```js
function toDecimal(binary) {

  return binary;
}
```

# --solutions--

```js
function toDecimal(binary) {
  let decimal = 0;
  for (let i = 0; i < binary.length; i++) {
    const bit = parseInt(binary[binary.length - 1 - i], 10);
    decimal += bit * Math.pow(2, i);
  }
  return decimal;
}
```