id: 68b7687dded630607aceccaf
title: "Vežba 47: Detekcija prevelikih brzina"
challengeType: 28
dashedName: challenge-47
---

# --description--

Ako vam je dat niz podataka brojeva koji predstavljaju brzine sa kojima su vozila primljena u pokretu, i broj koji predstavlja ograničenje brzine, vratite niz podataka sa dva elementa: broj vozila koja su prebrzo vozila, praćen prosekom viška brzine iznad ograničenja za ta vozila.

- Ako nijedno vozilo nije prebrzo vozilo, vratite `[0, 0]`.

# --hints--

`speeding([50, 60, 55], 60)` bi trebalo da vrati `[0, 0]`.

```js
assert.deepEqual(speeding([50, 60, 55], 60), [0, 0]);
```

`speeding([58, 50, 60, 55], 55)` bi trebalo da vrati `[2, 4]`.

```js
assert.deepEqual(speeding([58, 50, 60, 55], 55), [2, 4]);
```

`speeding([61, 81, 74, 88, 65, 71, 68], 70)` bi trebalo da vrati `[4, 8.5]`.

```js
assert.deepEqual(speeding([61, 81, 74, 88, 65, 71, 68], 70), [4, 8.5]);
```

`speeding([100, 105, 95, 102], 100)` bi trebalo da vrati `[2, 3.5]`.

```js
assert.deepEqual(speeding([100, 105, 95, 102], 100), [2, 3.5]);
```

`speeding([40, 45, 44, 50, 112, 39], 55)` bi trebalo da vrati `[1, 57]`.

```js
assert.deepEqual(speeding([40, 45, 44, 50, 112, 39], 55), [1, 57]);
```

# --seed--

## --seed-contents--

```js
function speeding(speeds, limit) {

  return speeds;
}
```

# --solutions--

```js
function speeding(speeds, limit) {
  let speeding = 0;
  let totalOver = 0;

  for (let speed of speeds) {
    if (speed > limit) {
      speeding++;
      totalOver += (speed - limit);
    }
  }

  if (speeding === 0) return [0, 0];

  return [speeding, totalOver / speeding];
}
```