id: 68b7cadffed0e75a517da677
title: "Vežba 54: Snaga lozinke! (Password Strength!)"
challengeType: 28
dashedName: challenge-54
---

# --description--

Dato niz znakova lozinke, vratite `"weak"`, `"medium"` ili `"strong"` na osnovu snage te lozinke.

Lozinka se procenjuje prema sledećim pravilima:

- Mora imati najmanje osmi znak.
- Mora imati i velika i mala slova.
- Mora imati najmanje jednu cifru.
- Mora imati najmanje jedan specijalni znak iz ovog skupa: `!`, `@`, `#`, `$`, `%`, `^`, `&`, ili `*`.

Vratite `"weak"` ako lozinka ispunjava dva ili manje pravila.
Vratite `"medium"` ako lozinka ispunjava 2 ili 3 pravila.
Vratite `"strong"` ako lozinka ispunjava sva 4 pravila.

# --hints--

`checkStrength("123456")` bi trebalo da vrati `"weak"`.

```js
assert.equal(checkStrength("123456"), "weak");
```

`checkStrength("pass!!!")` bi trebalo da vrati `"weak"`.

```js
assert.equal(checkStrength("pass!!!"), "weak");
```

`checkStrength("Qwerty")` bi trebalo da vrati `"weak"`.

```js
assert.equal(checkStrength("Qwerty"), "weak");
```

`checkStrength("PASSWORD")` bi trebalo da vrati `"weak"`.

```js
assert.equal(checkStrength("PASSWORD"), "weak");
```

`checkStrength("PASSWORD!")` bi trebalo da vrati `"medium"`.

```js
assert.equal(checkStrength("PASSWORD!"), "medium");
```

`checkStrength("PassWord%^!")` bi trebalo da vrati `"medium"`.

```js
assert.equal(checkStrength("PassWord%^!"), "medium");
```

`checkStrength("qwerty12345")` bi trebalo da vrati `"medium"`.

```js
assert.equal(checkStrength("qwerty12345"), "medium");
```

`checkStrength("S3cur3P@ssw0rd")` bi trebalo da vrati `"strong"`.

```js
assert.equal(checkStrength("S3cur3P@ssw0rd"), "strong");
```

`checkStrength("C0d3&Fun!")` bi trebalo da vrati `"strong"`.

```js
assert.equal(checkStrength("C0d3&Fun!"), "strong");
```

# --seed--

## --seed-contents--

```js
function checkStrength(password) {

  return password;
}
```

# --solutions--

```js
function checkStrength(password) {
  let rulesMet = 0;

  if (password.length >= 8) rulesMet++;
  if (/[a-z]/.test(password) && /[A-Z]/.test(password)) rulesMet++;
  if (/\d/.test(password)) rulesMet++;
  if (/[!@#$%^&*]/.test(password)) rulesMet++;

  if (rulesMet < 2) return "weak";
  if (rulesMet < 4) return "medium";
  return "strong";
}
```