id: 68c1a929005bf54d342aa8d3
title: "Vežba 56: Nedelja u svemiru 2: Pronalaženje egzoplanete"
challengeType: 28
dashedName: challenge-56
---

# --description--

Za drugi dan Nedelje u svemiru, dobijate niz slova gde svako slovo predstavlja merenje svetline zvezde. Utvrdite da li su merenja otkrila egzoplanetu korišćenjem tranzitne metode (transit method). Tranzitna metoda je kada planeta prođe ispred zvezde, smanjujući vidljivu svetlinu.

- Merenja svetline sadrže samo slova `0-9` i `A-Z`, gde svako merenje odgovara sledećim numeričkim vrednostima:
- Slova `0-9` predstavljaju nivoe svetline `0-9`.
- Slova `A-Z` predstavljaju nivoe svetline `10-35`.

Zvezda se smatra da ima egzoplanetu ako je bilo koje merenje manje ili jednako 80 procenata proseka svih merenja. Na primer, ako je prosečna svetlina zvezde 10, smatraće se da ima egzoplanetu ako je bilo koje merenje 8 ili niže.

# --hints--

`hasExoplanet("665544554")` treba da vrati `false`.

```js
assert.isFalse(hasExoplanet("665544554"));
```

`hasExoplanet("FGFFCFFGG")` treba da vrati `true`.

```js
assert.isTrue(hasExoplanet("FGFFCFFGG"));
```

`hasExoplanet("MONOPLONOMONPLNOMPNOMP")` treba da vrati `false`.

```js
assert.isFalse(hasExoplanet("MONOPLONOMONPLNOMPNOMP"));
```

`hasExoplanet("FREECODECAMP")` treba da vrati `true`.

```js
assert.isTrue(hasExoplanet("FREECODECAMP"));
```

`hasExoplanet("9AB98AB9BC98A")` treba da vrati `false`.

```js
assert.isFalse(hasExoplanet("9AB98AB9BC98A"));
```

`hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` treba da vrati `true`.

```js
assert.isTrue(hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"));
```

# --seed--

## --seed-contents--

```js
function hasExoplanet(readings) {

  return readings;
}
```

# --solutions--

```js
function hasExoplanet(readings) {
  let total = 0;
  const values = readings.split('').map(c => parseInt(c, 36));
  
  total = values.reduce((sum, v) => sum + v, 0);
  const average = total / values.length;
  const threshold = average * 0.8;

  return values.some(v => v <= threshold);
}
```