id: 68c1a929005bf54d342aa8d2
title: "Vežba 55: Nedelja u svemiru Dan 1: Klasifikacija zvezda"
challengeType: 28
dashedName: challenge-55
---

# --description--

4. oktobar označava početak Svemirske nedelje na Zemlji. Sledećih sedam dana donosiće vam izazove programiranja sa astrološkom tematikom.

Za današnji izazov, dat je temperaturu površine zvezde u Kelvinima (K), a potrebno je odrediti njenu klasifikaciju zvezda prema sledećim standardima:

- `"O"`: 30,000 K ili više
- `"B"`: 10,000 K - 29,999 K
- `"A"`: 7,500 K - 9,999 K
- `"F"`: 6,000 K - 7,499 K
- `"G"`: 5,200 K - 5,999 K
- `"K"`: 3,700 K - 5,199 K
- `"M"`: 0 K - 3,699 K

- Vratite datiu klasifikaciju zvezde.

# --hints--

`classification(5778)` bi trebalo da vrati `"G"`.

```js
assert.equal(classification(5778), "G");
```

`classification(2400)` bi trebalo da vrati `"M"`.

```js
assert.equal(classification(2400), "M");
```

`classification(9999)` bi trebalo da vrati `"A"`.

```js
assert.equal(classification(9999), "A");
```

`classification(3700)` bi trebalo da vrati `"K"`.

```js
assert.equal(classification(3700), "K");
```

`classification(3699)` bi trebalo da vrati `"M"`.

```js
assert.equal(classification(3699), "M");
```

`classification(210000)` bi trebalo da vrati `"O"`.

```js
assert.equal(classification(210000), "O");
```

`classification(6000)` bi trebalo da vrati `"F"`.

```js
assert.equal(classification(6000), "F");
```

`classification(11432)` bi trebalo da vrati `"B"`.

```js
assert.equal(classification(11432), "B");
```

# --seed--

## --seed-contents--

```js
function classification(temp) {

  return temp;
}
```

# --solutions--

```js
function classification(temp) {
  if (temp >= 30000) return "O";
  if (temp >= 10000) return "B";
  if (temp >= 7500) return "A";
  if (temp >= 6000) return "F";
  if (temp >= 5200) return "G";
  if (temp >= 3700) return "K";
  return "M";
}
```