id: 68c1a929005bf54d342aa8d4
title: "Vežba 57: Tjedan u svemiru dan 3: Poziv kući"
challengeType: 28
dashedName: challenge-57
---

# --description--

Za treći dan Tjedna u svemiru, dobijate niz podataka brojeva koji predstavljaju udaljenost (u kilometrima) između vas, satelita i vaše matične planete na komunikacionoj liniji. Odredite koliko vremena će poruka potrebno da stigne do ciljne planete koristeći sljedeće uslove:

- Prva vrijednost u nizu podataka je udaljenost od vašeg položaja do prvog satelita.
- Svaka sljedeća vrijednost, osim posljednje, predstavlja udaljenost do sljedećeg satelita.
- Posljednja vrijednost u nizu podataka je udaljenost od zadnjeg satelita do vaše matične planete.
- Poruka putuje brzinom od 300.000 km/s.
- Svaki satelit koji poruka prolazi povećava kašnjenje prenosa za 0,5 sekunde.
- Vratite broj zaokružen na 4 decimalne mjesta, uklanjajući zadnje nule.

# --hints--

`sendMessage([300000, 300000])` inapaslo kuruditi `2.5`.

```js
assert.equal(sendMessage([300000, 300000]), 2.5);
```

`sendMessage([384400, 384400])` inapaslo kuruditi `3.0627`.

```js
assert.equal(sendMessage([384400, 384400]), 3.0627);
```

`sendMessage([54600000, 54600000])` inapaslo kuruditi `364.5`.

```js
assert.equal(sendMessage([54600000, 54600000]), 364.5);
```

`sendMessage([1000000, 500000000, 1000000])` inapaslo kuruditi `1674.3333`.

```js
assert.equal(sendMessage([1000000, 500000000, 1000000]), 1674.3333);
```

`sendMessage([10000, 21339, 50000, 31243, 10000])` inapaslo kuruditi `2.4086`.

```js
assert.equal(sendMessage([10000, 21339, 50000, 31243, 10000]), 2.4086);
```

`sendMessage([802101, 725994, 112808, 3625770, 481239])` inapaslo kuruditi `21.1597`.

```js
assert.equal(sendMessage([802101, 725994, 112808, 3625770, 481239]), 21.1597);
```

# --seed--

## --seed-contents--

```js
function sendMessage(route) {

  return route;
}
```

# --solutions--

```js
function sendMessage(route) {
  let totalDistance = route.reduce((a, d) => a += d, 0);

  const delay = (route.length - 1) * 0.5
  const time = totalDistance / 300000
  const total = time + delay;
  return Number(total.toFixed(4))
}
```