id: 68c1a929005bf54d342aa8d6
title: "Vežba 59: Tjedan u svemiru Dan 5: Goldilocks zona"
challengeType: 28
dashedName: challenge-59
---

# --description--

Za peti dan Tjedna u svemiru, izračunaćeš "Goldilocks zonu" zvijjezde – oblast oko zvijjezde gdje su uslovi "savršeni" za postojanje tekućih voda.

Uzimajući u obzir masu zvijjezde, vrati niz podataka sa početnom i krajnjom udaljenošću njene Goldilocks zone za svemirsko mjerilo.

Za izračunavanje Goldilocks zone:

1. Pronađi luminozitet zvijjezde podizanjem njene mase na stepen 3.5.
2. Početak zone je 0.95 puta kvadratni korijen njenog luminoziteta.
3. Kraj zone je 1.37 puta kvadratni korijen njenog luminoziteta.

- Vrati udaljenost zaokruženu na dvije decimale.

Na primjer, ako dobiješ `1` kao masu, vrati `[0.95, 1.37]`.

# --hints--

`goldilocksZone(1)` bi trebalo da vrati `[0.95, 1.37]`.

```js
assert.deepEqual(goldilocksZone(1), [0.95, 1.37]);
```

`goldilocksZone(0.5)` bi trebalo da vrati `[0.28, 0.41]`.

```js
assert.deepEqual(goldilocksZone(0.5), [0.28, 0.41]);
```

`goldilocksZone(6)` bi trebalo da vrati `[21.85, 31.51]`.

```js
assert.deepEqual(goldilocksZone(6), [21.85, 31.51]);
```

`goldilocksZone(3.7)` bi trebalo da vrati `[9.38, 13.52]`.

```js
assert.deepEqual(goldilocksZone(3.7), [9.38, 13.52]);
```

`goldilocksZone(20)` bi trebalo da vrati `[179.69, 259.13]`.

```js
assert.deepEqual(goldilocksZone(20), [179.69, 259.13]);
```

# --seed--

## --seed-contents--

```js
function goldilocksZone(mass) {

  return mass;
}
```

# --solutions--

```js
function goldilocksZone(mass) {
  const luminosity = Math.pow(mass, 3.5);
  const start = 0.95 * Math.sqrt(luminosity);
  const end = 1.37 * Math.sqrt(luminosity);
  return [Number(start.toFixed(2)), Number(end.toFixed(2))];
}
```