id: 68c1a929005bf54d342aa8d5
title: "Vežba 58: Nedelja u svemiru Dan 4: Mesto za slijetanje"
challengeType: 28
dashedName: challenge-58
---

# --description--

U četvrti dan Nedelje u svemiru, dobijate matricu brojeva (dvodimenzionalni niz podataka) koja predstavlja zone koje mogu biti korišćene za slijetanje vašeg rovera. Pronađite najbezbednije mesto za slijetanje na osnovu sledećih pravila:

- Svaki element u matrici će imati broj od `0-9`, uključujući i nulu.
- Svaka `0` predstavlja potencijalno mesto za slijetanje.
- Svaki broj osim `0` je veoma opasno mesto za slijetanje. Što veći je broj, to je veća opasnost.
- Najbezbednije mesto za slijetanje definiše se kao ćelija sa vrednošću `0` čije susedne ćelije (do 4 suseda, bez brojeg ivice) imaju najniži ukupan nivo opasnosti.
- Ignorišite susede koji su van granica (uglasi i ivice imaju samo nekoliko suseda).
- Vratite koordinate najbezbednijeg mesta za slijetanje. Uvek će u biti postojati samo jedno najbezbednije mesto.

Na primer, ako dobijete:

```js
[
  [1, 0],
  [2, 0]
]
```

Vratite `[0, 1]`, koordinate za `0` u prvom redu.

# --hints--

`findLandingSpot([[1, 0], [2, 0]])` bi trebalo da vrati `[0, 1]`.

```js
assert.deepEqual(findLandingSpot([[1, 0], [2, 0]]), [0, 1]);
```

`findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]])` bi trebalo da vrati `[1, 1]`.

```js
assert.deepEqual(findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]]), [1, 1]);
```

`findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]])` bi trebalo da vrati `[2, 2]`.

```js
assert.deepEqual(findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]]), [2, 2]);
```

`findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]])` bi trebalo da vrati `[2, 1]`.

```js
assert.deepEqual(findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]]), [2, 1]);
```

# --seed--

## --seed-contents--

```js
function findLandingSpot(matrix) {

  return matrix;
}
```

# --solutions--

```js
function findLandingSpot(matrix) {
  let bestSpot = null;
  let lowestNeighborSum = Infinity;

  for (let i = 0; i < matrix.length; i++) {
    for (let j = 0; j < matrix[i].length; j++) {

      if (matrix[i][j] === 0) {
        let currentNeighborSum = 0;

        if (i > 0) currentNeighborSum += matrix[i - 1][j];
        if (j < matrix[i].length - 1) currentNeighborSum += matrix[i][j + 1];
        if (i < matrix.length - 1) currentNeighborSum += matrix[i + 1][j];
        if (j > 0) currentNeighborSum += matrix[i][j - 1];

        if (currentNeighborSum < lowestNeighborSum) {
          lowestNeighborSum = currentNeighborSum;
          bestSpot = [i, j];
        }
      }
    }
  }

  return bestSpot;
}
```