id: 68c497f3aaefc9fd9f1b0e26
title: "Vežba 62: Hex na Decimal"
challengeType: 28
dashedName: challenge-62
---

# --description--

Dato je niz karaktera koji predstavlja heksadecimalni broj (osnova 16); vratite njegovu decimalnu vrednost (osnova 10) kao celobrojni broj.

Heksadecimalni sistem brojeva koristi 16 cifara:

- `0-9` predstavljaju vrednosti od `0` do `9`.
- `A-F` predstavljaju vrednosti od `10` do `15`.

Evo tabele za konverziju delova:

| Heksadecimal | Decimalni |
|:-----------:|:--------:|
| 0           | 0        |
| 1           | 1        |
| ...         | ...      |
| 9           | 9        |
| A           | 10       |
| ...         | ...      |
| F           | 15       |
| 10          | 16       |
| ...         | ...      |
| 9F          | 159      |
| A0          | 160      |
| ...         | ...      |
| FF          | 255      |
| 100         | 256      |

- Niz karaktera će sadržati samo karaktere `0–9` i `A–F`.

# --hints--

`hexToDecimal("A")` bi trebalo da vrati `10`.

```js
assert.equal(hexToDecimal("A"), 10);
```

`hexToDecimal("15")` bi trebalo da vrati `21`.

```js
assert.equal(hexToDecimal("15"), 21);
```

`hexToDecimal("2E")` bi trebalo da vrati `46`.

```js
assert.equal(hexToDecimal("2E"), 46);
```

`hexToDecimal("FF")` bi trebalo da vrati `255`.

```js
assert.equal(hexToDecimal("FF"), 255);
```

`hexToDecimal("A3F")` bi trebalo da vrati `2623`.

```js
assert.equal(hexToDecimal("A3F"), 2623);
```

# --seed--

## --seed-contents--

```js
function hexToDecimal(hex) {

  return hex;
}
```

# --solutions--

```js
function hexToDecimal(hex) {

  return parseInt(hex, 16);
}
```