id: 68c497f3aaefc9fd9f1b0e24
title: "Vežba 60: Nedelja u svemiru Dan 6: Faza Meseca"
challengeType: 28
dashedName: challenge-60
---

# --description--

Za šestog dana Nedelje u svemiru, biće vam dat datum u formatu `"YYYY-MM-DD"` i morate odrediti fazu meseca za taj dan koristeći sledeća pravila:

Koristite jednostavan ciklusa meseca od 28 dana, podeljen na četiri jednake faze:

- `"New"` (Novi): dan 1 - 7
- `"Waxing"` (Povećanje): dan 8 - 14
- `"Full"` (Pun): dan 15 - 21
- `"Waning"` (Smanjenje): dan 22 - 28

Nakon dana 28, ciklus se ponavlja sa danom 1, novi mesec.

- Koristite `"2000-01-06"` kao referencu za novi mesec (dan 1 ciklusa) da biste odredili fazu datuma.
- Neće vam biti dat nijedan datum pre referentnog datuma.
- Vratite tačnu fazu kao niz karaktera.

**Napomena:** Dan 1 predstavlja dan novog meseca, što znači da je prošlo 0 dana od prethodnog novog meseca.

# --hints--

`moonPhase("2000-01-12")` treba da vrati `"New"`.

```js
assert.equal(moonPhase("2000-01-12"), "New");
```

`moonPhase("2000-01-13")` treba da vrati `"Waxing"`.

```js
assert.equal(moonPhase("2000-01-13"), "Waxing");
```

`moonPhase("2014-10-15")` treba da vrati `"Full"`.

```js
assert.equal(moonPhase("2014-10-15"), "Full");
```

`moonPhase("2012-10-21")` treba da vrati `"Waning"`.

```js
assert.equal(moonPhase("2012-10-21"), "Waning");
```

`moonPhase("2022-12-14")` treba da vrati `"New"`.

```js
assert.equal(moonPhase("2022-12-14"), "New");
```

# --seed--

## --seed-contents--

```js
function moonPhase(dateString) {

  return dateString;
}
```

# --solutions--

```js
function moonPhase(dateString) {
  const ONE_DAY_IN_MS = 1000 * 60 * 60 * 24;
  const refDate = new Date("2000-01-06");
  const targetDate = new Date(dateString);
  const diffMs = targetDate - refDate;
  const diffDays = diffMs / ONE_DAY_IN_MS;
  const phaseDay = (diffDays % 28 + 1);

  if (phaseDay <= 7) {
    return "New";
  } else if (phaseDay <= 14) {
    return "Waxing";
  } else if (phaseDay <= 21) {
    return "Full"
  } else {
    return "Waning";
  }
}

```