id: 68cae5b538ff798bbd4da00a
title: "Vežba 72: Regulator termostata 2"
challengeType: 28
dashedName: challenge-72
---

# --description--

Uzimajući u obzir trenutnu temperaturu sobe u Fahrenheit i željenu temperaturu u Celzijus, vratite niz karaktera koji opisuje kako podesiti temperaturu sobe na osnovu ovih uslova:

- Vratite `"Heat: X degrees Fahrenheit"` ako je trenutna temperatura niža od ciljne. Gde je `X` broj stepeni Farenhita sa kojima se mora podignuti temperatura sobe do cilja, zaokružen na jednu decimalnu tačku.
- Vratite `"Cool: X degrees Fahrenheit"` ako je trenutna temperatura viša od ciljne. Gde je `X` broj stepeni Farenhita sa kojima se mora smanjiti temperatura sobe do cilja, zaokružen na jednu decimalnu tačku.
- Vratite `"Hold"` ako je trenutna temperatura jednaka ciljnoj.

Da biste pretvorili Celzijus u Fahrenheit, pomnožite temperaturu u Celzijusu sa 1.8, a zatim dodajte 32 rezultatu (`F = (C * 1.8) + 32`).

# --hints--

`adjustThermostat(32, 0)` bi trebalo da vrati `"Hold"`.

```js
assert.equal(adjustThermostat(32, 0), "Hold");
```

`adjustThermostat(70, 25)` bi trebalo da vrati `"Heat: 7.0 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(70, 25), "Heat: 7.0 degrees Fahrenheit");
```

`adjustThermostat(72, 18)` bi trebalo da vrati `"Cool: 7.6 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(72, 18), "Cool: 7.6 degrees Fahrenheit");
```

`adjustThermostat(212, 100)` bi trebalo da vrati `"Hold"`.

```js
assert.equal(adjustThermostat(212, 100), "Hold");
```

`adjustThermostat(59, 22)` bi trebalo da vrati `"Heat: 12.6 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(59, 22), "Heat: 12.6 degrees Fahrenheit");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(currentF, targetC) {

  return currentF;
}
```

# --solutions--

```js
function adjustThermostat(currentF, targetC) {
  const targetF = targetC * 1.8 + 32;
  const diff = Math.abs(targetF - currentF).toFixed(1);

  if (currentF < targetF) {
    return `Heat: ${diff} degrees Fahrenheit`;
  } else if (currentF > targetF) {
    return `Cool: ${diff} degrees Fahrenheit`;
  } else {
    return "Hold"
  }
}
```