id: 68cae5b538ff798bbd4da009
title: "Vežba 71: Kalkulator napojnica"
challengeType: 28
dashedName: challenge-71
---

# --description--

S obzirom na cenu vašeg obroka i željeni procenat napojnice, vratite niz od tri vrednosti napojnice; 15%, 20% i željeni prilagođeni iznos napojnice.

- Cene će biti date u formatu: `"$N.NN"`.
- Standardni procenati napojnica biće dati u ovom formatu: `"25%"`.
- Vratite iznos u istom formatu od `"$N.NN"`, zaokružen na dve decimalne za mestima.

Na primer, s obzirom na cenu obroka od `"$10.00"`, i standardnu vrednost napojnice od `"25%"`, vratite `["$1.50", "$2.00", "$2.50"]`.

# --hints--

`calculateTips("$10.00", "25%")` treba da vrati `["$1.50", "$2.00", "$2.50"]`.

```js
assert.deepEqual(calculateTips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"]);
```

`calculateTips("$89.67", "26%")` treba da vrati `["$13.45", "$17.93", "$23.31"]`.

```js
assert.deepEqual(calculateTips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"]);
```

`calculateTips("$19.85", "9%")` treba da vrati `["$2.98", "$3.97", "$1.79"]`.

```js
assert.deepEqual(calculateTips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"]);
```

# --seed--

## --seed-contents--

```js
function calculateTips(mealPrice, customTip) {

  return mealPrice;
}
```

# --solutions--

```js
function calculateTips(mealPrice, customTip) {
  const meal = parseFloat(mealPrice.slice(1));
  const customPercent = parseFloat(customTip.slice(0, -1));
  const tips = [15, 20, customPercent];

  return tips.map(percent => {
    const amount = (meal * percent / 100).toFixed(2);
    return `$${amount}`;
  });
}
```