id: 68cae5b538ff798bbd4da003
title: "Vežba 65: Brojanje sekvenci slova"
challengeType: 28
dashedName: challenge-65
---

# --description--

Uzimajući u obzir sekvencu od dva slova, odredite koliko puta se druga sekvenca pojavljuje u prvoj.

- Primer sekvenca može se preklapati u prvoj sekvenci slova. Na primer, `"aaa"` sadrži `"aa"` dva puta. Prva dva 'a' i poslednja dva 'a'.

# --hints--

`count('abcdefg', 'def')` bi trebalo da vrati `1`.

```js
assert.equal(count('abcdefg', 'def'), 1);
```

`count('hello', 'world')` bi trebalo da vrati `0`.

```js
assert.equal(count('hello', 'world'), 0);
```

`count('mississippi', 'iss')` bi trebalo da vrati `2`.

```js
assert.equal(count('mississippi', 'iss'), 2);
```

`count('she sells seashells by the seashore', 'sh')` bi trebalo da vrati `3`.

```js
assert.equal(count('she sells seashells by the seashore', 'sh'), 3);
```

`count('101010101010101010101', '101')` bi trebalo da vrati `10`.

```js
assert.equal(count('101010101010101010101', '101'), 10);
```

# --seed--

## --seed-contents--

```js
function count(text, pattern) {

  return text;
}
```

# --solutions--

```js
function count(text, pattern) {
  let occurrences = 0;

  for (let i = 0; i <= text.length - pattern.length; i++) {
    if (text.slice(i, i + pattern.length) === pattern) {
      occurrences++;
    }
  }

  return occurrences;
}
```