id: 68cae5b538ff798bbd4da002
title: "Vežba 64: sa 24 na 12"
challengeType: 28
dashedName: challenge-64
---

# --description--

Dato je niz znakova koji predstavlja vreme u 24-časlosnom formatu sa oznakom "HHMM"; pretvorite to vreme u ekvivalentni 12-časlosni format sa oznakama "H:MM AM" ili "H:MM PM".

- Unos će uvek biti niz od četiri cifre u 24-časlosnom formatu, opsežan od "0000" do "2359".

# --hints--

`to12("1124")` bi trebalo da vrati `"11:24 AM"`.

```js
assert.equal(to12("1124"), "11:24 AM");
```

`to12("0900")` bi trebalo da vrati `"9:00 AM"`.

```js
assert.equal(to12("0900"), "9:00 AM");
```

`to12("1455")` bi trebalo da vrati `"2:55 PM"`.

```js
assert.equal(to12("1455"), "2:55 PM");
```

`to12("2346")` bi trebalo da vrati `"11:46 PM"`.

```js
assert.equal(to12("2346"), "11:46 PM");
```

`to12("0030")` bi trebalo da vrati `"12:30 AM"`.

```js
assert.equal(to12("0030"), "12:30 AM");
```

# --seed--

## --seed-contents--

```js
function to12(time) {

  return time;
}
```

# --solutions--

```js
function convertHours(hours) {
  if (hours === 0) return 12;
  if (hours > 12) return hours - 12;
  return hours;
}

function to12(time) {
  const hours = parseInt(time.slice(0, 2), 10);
  const minutes = time.slice(2);
  const period = hours < 12 ? 'AM' : 'PM';

  return `${convertHours(hours)}:${minutes} ${period}`;
}
```