id: 68cae5b538ff798bbd4da006
title: "Vežba 68: Maskiranje kreditne kartice"
challengeType: 28
dashedName: challenge-68
---

# --description--

Kada vam je dat niz brojeva kreditne kartice, vratite maskirani format koristeći sledeće uslove:

- Niz će imati četiri grupe od po četiri broja svaka (`0-9`), a sve grupe su razdvojene jednom razmakom ili crticom (`-`).
- Zamijenite sve brojeve osim poslednje četiri sa zvezdicama (`*`).
- Ostavite preostale znakove nepromenjene.

Na primer, ako je dat `"4012-8888-8888-1881"`, vratite `"****-****-****-1881"`.

# --hints--

`mask("4012-8888-8888-1881")` bi trebalo da vrati `"****-****-****-1881"`.

```js
assert.equal(mask("4012-8888-8888-1881"), "****-****-****-1881");
```

`mask("5105 1051 0510 5100")` bi trebalo da vrati `"**** **** **** 5100"`.

```js
assert.equal(mask("5105 1051 0510 5100"), "**** **** **** 5100");
```

`mask("6011 1111 1111 1117")` bi trebalo da vrati `"**** **** **** 1117"`.

```js
assert.equal(mask("6011 1111 1111 1117"), "**** **** **** 1117");
```

`mask("2223-0000-4845-0010")` bi trebalo da vrati `"****-****-****-0010"`.

```js
assert.equal(mask("2223-0000-4845-0010"), "****-****-****-0010");
```

# --seed--

## --seed-contents--

```js
function mask(card) {

  return card;
}
```

# --solutions--

```js
function mask(card) {
  const split = card.split(card.includes('-') ? '-' : ' ');
  
  return card.includes('-') ?
    `****-****-****-${split[3]}` :
    `**** **** **** ${split[3]}`;
}
```