id: 68cae5b538ff798bbd4da008
title: "Vežba 70: Ekstraktor atributa HTML elementa"
challengeType: 28
dashedName: challenge-70
---

# --description--

Kada dobijete niz znakova validnog HTML elementa, vratite atribute tog elementa koristeći sledeće kriterijume:

- Biće vam dat samo jedan element.
- Atributi će biti u formatu: `attribute="value"`.
- Vratite niz podataka sa svakim atributom i njegovom vrednošću, razdvojenima zarezom, u ovom formatu: `["attribute1, value1", "attribute2, value2"]`.
- Atribute vratiti u redosledu kojim su dati.
- Ako nema atributa, vratite prazan niz podataka.

# --hints--

`extractAttributes('<span class="red"></span>')` bi trebalo da vrati `["class, red"]`.

```js
assert.deepEqual(extractAttributes('<span class="red"></span>'), ["class, red"]);
```

`extractAttributes('<meta charset="UTF-8" />')` bi trebalo da vrati `["charset, UTF-8"]`.

```js
assert.deepEqual(extractAttributes('<meta charset="UTF-8" />'), ["charset, UTF-8"]);
```

`extractAttributes("<p>Lorem ipsum dolor sit amet</p>")` bi trebalo da vrati `[]`.

```js
assert.deepEqual(extractAttributes("<p>Lorem ipsum dolor sit amet</p>"), []);
```

`extractAttributes('<input name="email" type="email" required="true" />')` bi trebalo da vrati `["name, email", "type, email", "required, true"]`.

```js
assert.deepEqual(extractAttributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"]);
```

`extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>')` bi trebalo da vrati `["id, submit", "class, btn btn-primary"]`.

```js
assert.deepEqual(extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"]);
```

# --seed--

## --seed-contents--

```js
function extractAttributes(element) {

  return element;
}
```

# --solutions--

```js
function extractAttributes(element) {
  const regex = /(\w[\w-]*)="([^"]*)"/g;
  const result = [];
  let match;

  while ((match = regex.exec(element)) !== null) {
    const attr = match[1];
    const val = match[2];
    result.push(`${attr}, ${val}`);
  }

  return result;
}
```