id: 68cae5b538ff798bbd4da004
title: "Vežba 66: Uklanjanje HTML oznaka"
challengeType: 28
dashedName: challenge-66
---

# --description--

Dato vam je niz HTML koda; uklonite oznake i vratite običan tekstualni sadržaj.

- Ulazni niz će sadržati samo validan HTML.
- HTML oznake mogu biti ugnježđene.
- Uklonite sve oznake i atribute elemenata.

Na primer, `'<a href="#">Click here</a>'` bi trebalo da vrati `"Click here"`.

# --hints--

`stripTags('<a href="#">Click here</a>')` bi trebalo da vrati `"Click here"`.

```js
assert.equal(stripTags('<a href="#">Click here</a>'), "Click here");
```

`stripTags('<p class="center">Hello <b>World</b>!</p>')` bi trebalo da vrati `"Hello World!"`.

```js
assert.equal(stripTags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!");
```

`stripTags('<img src="cat.jpg" alt="Cat">')` bi trebalo da vrati prazan niz znakova (`""`).

```js
assert.equal(stripTags('<img src="cat.jpg" alt="Cat">'), "");
```

`stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` bi trebalo da vrati `sectionsection`.

```js
assert.equal(stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection");
```

# --seed--

## --seed-contents--

```js
function stripTags(html) {

  return html;
}
```

# --solutions--

```js
function stripTags(html) {
  return html.replace(/<[^>]*>/g, '');
}
```