id: 68cae5b538ff798bbd4da005
title: "Vežba 67: Validacija e-pošte"
challengeType: 28
dashedName: challenge-67
---

# --description--

Izvučena niz karaktera, utvrdite da li je validna e-mail adresa koristeći sledeće kriterijume:

- Mora sadržati samo jedan znak `@`.
- Lokalni deo (pre `@`):
  - Može sadržati samo slova (`a-z`, `A-Z`), brojeve (`0-9`), tačku (`.`), donje crte (`_`) ili tire (`-`).
  - Ne sme počinjati niti završavati sa tačkom.
- Domen deo (nakon `@`):
  - Mora sadržati najmanje jednu tačku.
  - Mora se završavati tačkom praćenom najmanje dva slova.
- Lokalni deo ni domen deo ne sme sadržati dve uzastopne tačke.

# --hints--

`validate("a@b.cd")` bi trebalo da vrati `true`.

```js
assert.isTrue(validate("a@b.cd"));
```

`validate("hell.-w.rld@example.com")` bi trebalo da vrati `true`.

```js
assert.isTrue(validate("hell.-w.rld@example.com"));
```

`validate(".b@sh.rc")` bi trebalo da vrati `false`.

```js
assert.isFalse(validate(".b@sh.rc"));
```

`validate("example@test.c0")` bi trebalo da vrati `false`.

```js
assert.isFalse(validate("example@test.c0"));
```

`validate("freecodecamp.org")` bi trebalo da vrati `false`.

```js
assert.isFalse(validate("freecodecamp.org"));
```

`validate("develop.ment_user@c0D!NG.R.CKS")` bi trebalo da vrati `true`.

```js
assert.isTrue(validate("develop.ment_user@c0D!NG.R.CKS"));
```

`validate("hello.@wo.rld")` bi trebalo da vrati `false`.

```js
assert.isFalse(validate("hello.@wo.rld"));
```

`validate("hello@world..com")` bi trebalo da vrati `false`.

```js
assert.isFalse(validate("hello@world..com"));
```

`validate("develop..ment_user@c0D!NG.R.CKS")` bi trebalo da vrati `false`.

```js
assert.isFalse(validate("develop..ment_user@c0D!NG.R.CKS"));
```

`validate("git@commit@push.io")` bi trebalo da vrati `false`.

```js
assert.isFalse(validate("git@commit@push.io"));
```

# --seed--

## --seed-contents--

```js
function validate(email) {

  return email;
}
```

# --solutions--

```js
function validate(email) {
  if (email.includes('..')) return false;

  const parts = email.split('@');
  if (parts.length !== 2) return false;

  const [local, domain] = parts;

  if (local.startsWith('.') || local.endsWith('.')) return false;
  if (!/^[a-zA-Z0-9._-]+$/.test(local)) return false;
  if (!domain.includes('.')) return false;

  const tld = domain.split('.').pop();
  if (tld.length < 2 || !/^[a-zA-Z]+$/.test(tld)) return false;

  return true;
}
```