id: 68cae5b538ff798bbd4da007
title: "Vežba 69: Izgubljeni čarape"
challengeType: 28
dashedName: challenge-69
---

# --description--

Uzimajući u obzir pun broj koji predstavlja početni broj para čarapa, i drugi broj koji predstavlja koliko je ciklusa pranja urađeno, vratite trenutni broj kompletnih parova čarapa koristeći sledeće uslove:

- Svakih 2 ciklusa pranja, gubiš jednu čaraupu.
- Svakih 3 ciklusa pranja, dobijaš jednu izgubljenu čaraupu.
- Svakih 5 ciklusa pranja, jedna čaraupa se ošteti i mora se baciti.
- Svakih 10 ciklusa pranja, kupuješ novi par čarapa.
- Ne možeš imati manje od nulte čaraupe.
- Pravila mogu biti kombinovana. Na primer, u 10. ciklusu pranja, gubiš jednu čaraupu, bacaš jednu i kupuješ novi par čarapa.
- Vrati broj kompletnih parova čarapa.

# --hints--

`sockPairs(2, 5)` bi trebalo da vrati `1`.

```js
assert.equal(sockPairs(2, 5), 1);
```

`sockPairs(1, 2)` bi trebalo da vrati `0`.

```js
assert.equal(sockPairs(1, 2), 0);
```

`sockPairs(5, 11)` bi trebalo da vrati `4`.

```js
assert.equal(sockPairs(5, 11), 4);
```

`sockPairs(6, 25)` bi trebalo da vrati `3`.

```js
assert.equal(sockPairs(6, 25), 3);
```

`sockPairs(1, 8)` bi trebalo da vrati `0`.

```js
assert.equal(sockPairs(1, 8), 0);
```

# --seed--

## --seed-contents--

```js
function sockPairs(pairs, cycles) {

  return pairs;
}
```

# --solutions--

```js
function sockPairs(pairs, cycles) {
  let socks = pairs * 2;

  for (let i = 1; i <= cycles; i++) {
    if (i % 2 === 0) socks -= 1;
    if (i % 3 === 0) socks += 1;
    if (i % 5 === 0) socks -= 1;
    if (i % 10 === 0) socks += 2;

    if (socks < 0) socks = 0;
  }

  return Math.floor(socks / 2);
}
```