id: 68d2ba1468508398389487ce
title: "Vežba 73: Govori mudro, moraš to učiniti"
challengeType: 28
dashedName: challenge-73
---

# --description--

Kada dobijete rečenicu, vratite njenu verziju koja zvuči kao savet od mudrog učitelja koristeći sledeća pravila:

- Reči su razdvojene samo jednim razmakom.
- Pronađite prvi navrh pomenutih reči u rečenici: `"have"`, `"must"`, `"are"`, `"will"`, `"can"`.
- Premestite sve reči pre i zajedno sa tom reči na kraj rečenice, pazeći da:
  - Sačuvate redosled reči prilikom pomeranja.
  - Sve pretvorite u mala slova.
  - I dodate zarez i razmak ispred njih.
- Veliko slovo pišete prvog reči nove rečenice.
- Sve izlazne rečenice završavaju sa jednim tačkom. Sačuvajte originalni znak interpunkcije rečenice i premestite ga na kraj nove rečenice.
- Vratite novu rečenicu, osiguravajući da postoji samo jedan razmak između svake reči i da nema razmaka na početku ili kraju rečenice.

Na primer, ako dobijete `"You must speak wisely."`, vratite `"Speak wisely, you must."`

# --hints--

`wiseSpeak("You must speak wisely.")` treba da vrati `"Speak wisely, you must."`

```js
assert.equal(wiseSpeak("You must speak wisely."), "Speak wisely, you must.");
```

`wiseSpeak("You can do it!")` treba da vrati `"Do it, you can!"`

```js
assert.equal(wiseSpeak("You can do it!"), "Do it, you can!");
```

`wiseSpeak("Do you think you will complete this?")` treba da vrati `"Complete this, do you think you will?"`

```js
assert.equal(wiseSpeak("Do you think you will complete this?"), "Complete this, do you think you will?");
```

`wiseSpeak("All your base are belong to us.")` treba da vrati `"Belong to us, all your base are."`

```js
assert.equal(wiseSpeak("All your base are belong to us."), "Belong to us, all your base are.");
```

`wiseSpeak("You have much to learn.")` treba da vrati `"Much to learn, you have."`

```js
assert.equal(wiseSpeak("You have much to learn."), "Much to learn, you have.");
```

# --seed--

## --seed-contents--

```js
function wiseSpeak(sentence) {

  return sentence;
}
```

# --solutions--

```js
function wiseSpeak(sentence) {
  const triggers = ["have", "must", "are", "will", "can"];
  const punctuation = sentence[sentence.length - 1];
  const words = sentence.split(" ");

  const triggerIndex = words.findIndex(w => triggers.includes(w));

  const toMove = words.slice(0, triggerIndex + 1).map(w => w.toLowerCase());
  const remaining = words.slice(triggerIndex + 1);

  const newStart = remaining.map((w, i) => {
    return i === 0 ? w[0].toUpperCase() + w.slice(1) :
      i === remaining.length - 1 ? w.slice(0, w.length - 1) : w;
  })

  return newStart.join(" ") + ", " + toMove.join(" ") + punctuation;
}
```