id: 68d2ba1468508398389487cf
title: "Vežba 74: Tvoje najdraže pesme"
challengeType: 28
dashedName: challenge-74
---

# --description--

Da li se sećate iPodsa? Prvi model je izašao pre 24 godine danas, 23. oktobar 2001.

S obziма na niz podataka objekata pesama koji predstavljaju vaš iPod playlist, vratite niz podataka sa naslovima dve najčešće puštene pesme, počevši od pesme koja je puštena najviše puta.

- Svaki objekat će imati svojstvo `"title"` (string/niz karaktera) i svojstvo `"plays"` (ceo broj).

# --hints--

`favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ])` bi trebalo da vrati `["Sync or Swim", "Earbud Blues"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Sync or Swim", "plays": 3}, {"title": "Byte Me", "plays": 1}, {"title": "Earbud Blues", "plays": 2} ]), ["Sync or Swim", "Earbud Blues"]);
```

`favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ])` bi trebalo da vrati `["Clickwheel Love", "99 Downloads"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Skip Track", "plays": 98}, {"title": "99 Downloads", "plays": 99}, {"title": "Clickwheel Love", "plays": 100} ]), ["Clickwheel Love", "99 Downloads"]);
```

`favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ])` bi trebalo da vrati `["Song B", "Song C"]`.

```js
assert.deepEqual(favoriteSongs([{"title": "Song A", "plays": 42}, {"title": "Song B", "plays": 99}, {"title": "Song C", "plays": 75} ]), ["Song B", "Song C"]);
```

# --seed--

## --seed-contents--

```js
function favoriteSongs(playlist) {

  return playlist;
}
```

# --solutions--

```js
function favoriteSongs(playlist) {
  const sorted = [...playlist].sort((a, b) => b.plays - a.plays);
  return sorted.slice(0, 2).map(song => song.title);
}
```