id: 68d2ba1468508398389487d0
title: "Vežba 75: Skriveno blago"
challengeType: 28
dashedName: challenge-75
---

# --description--

Dati je 2D niz podataka koji predstavlja mapu dna mora sa skrivenim blagom, i niz podataka sa koordinatama (`[red, kolona]`) za ronjenje u potragu za tim blagom. Vratite `"Empty"`, `"Found"` ili `"Recovered"` koristeći sledeća pravila:

- Dati 2D niz će sadržavati samo jedno neotkriveno blago koje zauzima više ćelija.
- Svaka ćelija u 2D nizu podataka ima jednu od sledećih vrednosti:
  - `"-"`: Nema blaga.
  - `"O"`: Deo blaga koji nije pronađen.
  - `"X"`: Deo blaga koji je već pronađen.

- Ako područje ronjenja nema blago, vratite `"Empty"`.
- Ako područje ronjenja pronalazi blago, ali barem deo tog blaga još uvek nije pronađen, vratite `"Found"`.
- Ako područje ronjenja pronalazi poslednji deo blaga koji nije pronađen, vratite `"Recovered"`.

Na primer, ako je dato:

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

I `[2, 1]` za koordinate područja ronjenja, vratite `"Recovered"` jer je ronjenje pronašlo poslednji deo blaga koji nije pronađen.

# --hints--

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` bi trebalo da vrati `"Recovered"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Recovered");
```

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` bi trebalo da vrati `"Empty"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Empty");
```

`dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` bi trebalo da vrati `"Found"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Found");
```

`dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` bi trebalo da vrati `"Found"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Found");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` bi trebalo da vrati `"Recovered"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Recovered");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` bi trebalo da vrati `"Empty"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Empty");
```

# --seed--

## --seed-contents--

```js
function dive(map, coordinates) {

  return map;
}
```

# --solutions--

```js
function dive(map, coordinates) {
  const [row, col] = coordinates;
  const target = map[row][col];

  if (target === "-") {
    return "Empty";
  }

  if (target === "O") {
    for (let r = 0; r < map.length; r++) {
      for (let c = 0; c < map[r].length; c++) {
        if ((r !== row || c !== col) && map[r][c] === "O") {
          return "Found";
        }
      }
    }
    return "Recovered";
  }

  return "Found";
}
```