id: 68d30845cc08266018fc46bc
title: "Vežba 76: Komplementarna DNA"
challengeType: 28
dashedName: challenge-76
---

# --description--

Dati je niz slova koji predstavlja DNA sekvencu; vratite njen komplementarni niz koristeći sledeća pravila:

- DNA sadrži slova "A", "C", "G" i "T".
- Slovo "A" i "T" su komplementarno jedno drugom.
- Slovo "C" i "G" su komplementarno jedno drugom.

Na primer, ako dobijete "ACGT", vratite "TGCA".

# --hints--

`complementaryDNA("ACGT")` bi trebalo da vrati `"TGCA"`.

```js
assert.equal(complementaryDNA("ACGT"), "TGCA");
```

`complementaryDNA("ATGCGTACGTTAGC")` bi trebalo da vrati `"TACGCATGCAATCG"`.

```js
assert.equal(complementaryDNA("ATGCGTACGTTAGC"), "TACGCATGCAATCG");
```

`complementaryDNA("GGCTTACGATCGAAG")` bi trebalo da vrati `"CCGAATGCTAGCTTC"`.

```js
assert.equal(complementaryDNA("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC");
```

`complementaryDNA("GATCTAGCTAGGCTAGCTAG")` bi trebalo da vrati `"CTAGATCGATCCGATCGATC"`.

```js
assert.equal(complementaryDNA("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC");
```

# --seed--

## --seed-contents--

```js
function complementaryDNA(strand) {

  return strand;
}
```

# --solutions--

```js
function complementaryDNA(strand) {
  const complements = { A: "T", T: "A", C: "G", G: "C" };
  return strand
    .split("")
    .map(c => complements[c])
    .join("");
}
```