id: 68d30845cc08266018fc46be
title: "Vežba 78: Sekvenca celih brojeva"
challengeType: 28
dashedName: challenge-78
---

# --description--

Ako dobijete pozitivan ceo broj, vratite niz znakova koji prikazuje sve cele brojeve od `1` do, i uključujući, taj broj, u redosledu brojeva.

Na primer, ako dobijete `5`, vratite `"12345"`.

# --hints--

`sequence(5)` bi trebalo da vrati `"12345"`.

```js
assert.equal(sequence(5), "12345");
```

`sequence(10)` bi trebalo da vrati `"12345678910"`.

```js
assert.equal(sequence(10), "12345678910");
```

`sequence(1)` bi trebalo da vrati `"1"`.

```js
assert.strictEqual(sequence(1), "1");
```

`sequence(27)` bi trebalo da vrati `"123456789101112131415161718192021222324252627"`.

```js
assert.equal(sequence(27), "123456789101112131415161718192021222324252627");
```

# --seed--

## --seed-contents--

```js
function sequence(n) {

  return n;
}
```

# --solutions--

```js
function sequence(n) {
  let result = "";
  for (let i = 1; i <= n; i++) {
    result += i;
  }
  return result;
}
```