id: 68d30fc57588d97fd3027b30
title: "Vežba 79: Navigator"
challengeType: 28
dashedName: challenge-79
---

# --description--

28. oktobar 1994. godine lansiran je Netscape Navigator, pomažući milionima ljudi da istražuju ranu mrežu.

Na osnovu niza podataka o komandama pretraživača koje ste izvršili u Netscape Navigator-u, vratite trenutnu stranicu koju imate nakon izvršenja svih komandi koristeći sledeća pravila:

- Uvek počinjete na stranici `"Home"`, koja nije uključena u niz komandi.
- Validne komande su:
  - `"Visit Page"`: Gde je `"Page"` naziv posete stranice. Na primer, `"Visit About"` vodi vas na stranicu `"About"`. Kada posetite novu stranicu, uverite se da brišete bilo koju istoriju napred koju imate.
  - `"Back"`: Vodi vas na prethodno posetenu stranicu u vašoj istoriji ili ostaje na trenutnoj stranici ako nema takve.
  - `"Forward"`: Voči vas napred u istoriji do stranice od koje ste otišli ili ostaje na trenutnoj stranici ako nema takve.

Na primer, za `["Visit About Us", "Back", "Forward"]`, vratite `"About Us"`.

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` bi trebalo da vrati `"About Us"`.

```js
assert.equal(navigate(["Visit About Us", "Back", "Forward"]), "About Us");
```

`navigate(["Forward"])` bi trebalo da vrati `"Home"`.

```js
assert.equal(navigate(["Forward"]), "Home");
```

`navigate(["Back"])` bi trebalo da vrati `"Home"`.

```js
assert.equal(navigate(["Back"]), "Home");
```

`navigate(["Visit About Us", "Visit Gallery"])` bi trebalo da vrati `"Gallery"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery"]), "Gallery");
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` bi trebalo da vrati `"Home"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home");
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` bi trebalo da vrati `"Contact"`.

```js
assert.equal(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact");
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` bi trebalo da vrati `"Visit Us"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us");
```

# --seed--

## --seed-contents--

```js
function navigate(commands) {

  return commands;
}
```

# --solutions--

```js
function navigate(commands) {
  const history = ["Home"];
  let currentPageIndex = 0;

  for (const command of commands) {
    if (command.startsWith("Visit")) {
      history.splice(currentPageIndex + 1);
      history.push(command.slice(6));
      currentPageIndex++;
    } else if(command === "Back" && currentPageIndex > 0) {
      currentPageIndex--;
    } else if (command === "Forward" && currentPageIndex < history.length - 1) {
      currentPageIndex++;
    }
  }

  return history[currentPageIndex];
}
```