id: 68e39ed6106dac2f0a98fd64
title: "Vežba 82: SpOoKy~CaSe"
challengeType: 28
dashedName: challenge-82
---

# --description--

Ako vam je dat niz slova koji predstavlja ime promenljive, pretvorite ga u "spooky case" koristeći sledeće uslove:

- Zamenite sve donje crte (`_`) i crtice (`-`) sa tilde (`~`).
- Prvo slovo niza slova pišite velikim slovom, a svako drugo uzastopno slovo takođe velikim slovom. Ignorišite tilde prilikom brojanja. Sve ostale slova pišite malim slovima.

Primer: Ako primite `hello_world`, vratite `HeLlO~wOrLd`.

# --hints--

`spookify("hello_world")` bi trebalo da vrati `"HeLlO~wOrLd"`.

```js
assert.equal(spookify("hello_world"), "HeLlO~wOrLd");
```

`spookify("Spooky_Case")` bi trebalo da vrati `"SpOoKy~CaSe"`.

```js
assert.equal(spookify("Spooky_Case"), "SpOoKy~CaSe");
```

`spookify("TRICK-or-TREAT")` bi trebalo da vrati `"TrIcK~oR~tReAt"`.

```js
assert.equal(spookify("TRICK-or-TREAT"), "TrIcK~oR~tReAt");
```

`spookify("c_a-n_d-y_-b-o_w_l")` bi trebalo da vrati `"C~a~N~d~Y~~b~O~w~L"`.

```js
assert.equal(spookify("c_a-n_d-y_-b-o_w_l"), "C~a~N~d~Y~~b~O~w~L");
```

`spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts")` bi trebalo da vrati `"ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS"`.

```js
assert.equal(spookify("thE_hAUntEd-hOUsE-Is-fUll_Of_ghOsts"), "ThE~hAuNtEd~HoUsE~iS~fUlL~oF~gHoStS");
```

# --seed--

## --seed-contents--

```js
function spookify(boo) {

  return boo;
}
```

# --solutions--

```js
function spookify(boo) {
  const replaced = boo.replace(/[_-]/g, "~");
  
  let result = "";
  let capitalize = true;

  for (let char of replaced) {
    if (char === "~") {
      result += char;
    } else {
      result += capitalize ? char.toUpperCase() : char.toLowerCase();
      capitalize = !capitalize;
    }
  }

  return result;
}
```