id: 68e39ed6106dac2f0a98fd63
title: "Vežba 81: nth premično broju"
challengeType: 28
dashedName: challenge-81
---

# --description--

Premični brojevi su prirodni pozitivni brojevi veći od 1 koji se dele samo sa 1 i sobom. Prvih pet premičnih brojeva su `2`, `3`, `5`, `7`, i `11`.

Ako unesete prirodni pozitivni broj `n`, vratite nth premično broju. Na primer, ako unesete `5`, vratite 5. premični broj: `11`.

# --hints--

`nthPrime(5)` treba da vrati `11`.

```js
assert.equal(nthPrime(5), 11);
```

`nthPrime(10)` treba da vrati `29`.

```js
assert.equal(nthPrime(10), 29);
```

`nthPrime(16)` treba da vrati `53`.

```js
assert.equal(nthPrime(16), 53);
```

`nthPrime(99)` treba da vrati `523`.

```js
assert.equal(nthPrime(99), 523);
```

`nthPrime(1000)` treba da vrati `7919`.

```js
assert.equal(nthPrime(1000), 7919);
```

# --seed--

## --seed-contents--

```js
function nthPrime(n) {

  return n;
}
```

# --solutions--

```js
function nthPrime(n) {
  const primes = [];
  let num = 2;

  while (primes.length < n) {
    let isPrime = true;
    for (let i = 2; i * i <= num; i++) {
      if (num % i === 0) {
        isPrime = false;
        break;
      }
    }
    if (isPrime) primes.push(num);
    num++;
  }

  return primes[n - 1];
}
```