id: 68e39ed6106dac2f0a98fd66
title: "Vežba 84: Infekcija"
challengeType: 28
dashedName: challenge-84
---

# --description--

Na 2. novembra 1988, prvi veliki internet crv je pušten, ostavljajući približno 10% povezanih računara na internetu zaraženim nakon samo jednog dana.

U ovom izazovu, dat vam je broj dana od kada je internet crv pušten, i morate utvrditi koliko je računara zarazeno koristeći sledeća pravila:

- Dan 0, prvi računar je zaražen.
- Svaki naredni dan, broj zaraženih računara se udvostručuje.
- Na svakom trećem danu, primenjuje se zaštita nakon što virus proširi i smanji broj zaraženih računara za 20%. Zaokružite broj računara koji su dobili zaštitu na sledeći ceo broj.

Na primer, dana:

- Dan 0: ukupno je zaraženo 1 računar.
- Dan 1: ukupno je zaraženo 2 računara.
- Dan 2: ukupno je zaraženo 4 računara.
- Dan 3: ukupno je zaraženo 8 računara. Zatim, primenite zaštitu: 8 zaraženih računara * 20% = 1.6 računara dobilo zaštitu. Zaokružite 1.6 na 2. 8 zaraženih računara - 2 dobila zaštitu = ukupno 6 zaraženih računara nakon dana 3.

Vratite ukupan broj zaraženih računara nakon datoog broja dana.

# --hints--

`infected(1)` treba da vrati `2`.

```js
assert.equal(infected(1), 2);
```

`infected(3)` treba da vrati `6`.

```js
assert.equal(infected(3), 6);
```

`infected(8)` treba da vrati `152`.

```js
assert.equal(infected(8), 152);
```

`infected(17)` treba da vrati `39808`.

```js
assert.equal(infected(17), 39808);
```

`infected(25)` treba da vrati `5217638`.

```js
assert.equal(infected(25), 5217638);
```

# --seed--

## --seed-contents--

```js
function infected(days) {

  return days;
}
```

# --solutions--

```js
function infected(days) {
  let infected = 1;

  for (let day = 1; day <= days; day++) {
    infected *= 2;

    if (day % 3 === 0) {
      let patched = Math.ceil(infected * 0.2);
      infected -= patched;
    }
  }

  return infected;
}
```