id: 68e39ed6106dac2f0a98fd62
title: "Vežba 80: Sortiranje e-pošte"
challengeType: 28
dashedName: challenge-80
---

# --description--

Prvi e-mail je poslat 29. oktobra 1971. godine, uspostavljajući format `korisničko_ime@domena` koji i danas koristimo. Sada postoji milijarde adresa e-pošte.

U ovom izazovu, dat vam je spisak adresa e-pošte i morate ih sortirati po domeni (delu nakon `@`) prvo, a zatim po korisničkom imenu (delu pre `@`).

- Sortiranje mora biti nezavisno od veličine slova (case-insensitive).
- Ako više e-mailova imaju istu domenu, sortirajte ih po korisničkom imenu.
- Vratite niz sortiranih adresa.
- Vraćene adrese moraju zadržati svoja originalna velika i mala slova.

Na primer, ako dobijete `["jill@mail.com", "john@example.com", "jane@example.com"]`, vratite `["jane@example.com", "john@example.com", "jill@mail.com"]`.

# --hints--

`sort(["jill@mail.com", "john@example.com", "jane@example.com"])` bi trebalo da vrati `["jane@example.com", "john@example.com", "jill@mail.com"]`.

```js
assert.deepEqual(sort(["jill@mail.com", "john@example.com", "jane@example.com"]), ["jane@example.com", "john@example.com", "jill@mail.com"]);
```

`sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"])` bi trebalo da vrati `["bob@mail.com", "carol@mail.com", "alice@zoo.com"]`.

```js
assert.deepEqual(sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"]), ["bob@mail.com", "carol@mail.com", "alice@zoo.com"]);
```

`sort(["user@z.com", "user@y.com", "user@x.com"])` bi trebalo da vrati `["user@x.com", "user@y.com", "user@z.com"]`.

```js
assert.deepEqual(sort(["user@z.com", "user@y.com", "user@x.com"]), ["user@x.com", "user@y.com", "user@z.com"]);
```

`sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"])` bi trebalo da vrati `["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]`.

```js
assert.deepEqual(sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"]), ["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]);
```

`sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"])` bi trebalo da vrati `["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]`.

```js
assert.deepEqual(sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"]), ["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]);
```

# --seed--

## --seed-contents--

```js
function sort(emails) {

  return emails;
}
```

# --solutions--

```js
function sort(emails) {
  return emails.slice().sort((a, b) => {
    const [userA, domainA] = a.split('@');
    const [userB, domainB] = b.split('@');

    const domainCompare = domainA.toLowerCase().localeCompare(domainB.toLowerCase());
    if (domainCompare !== 0) return domainCompare;

    return userA.toLowerCase().localeCompare(userB.toLowerCase());
  });
}
```