id: 68ee9e3066cfd4eb2328e8a7
title: "Vežba 88: Pronalazač dana u sedmici"
challengeType: 28
dashedName: challenge-88
---

# --description--

Ako vam je dat niz znakova datuma u formatu `YYYY-MM-DD`, vratite dan u sedmici.

Validni dana koje treba vratiti su:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

Uverite se da ne uzimate u obzir vremonske zone.

# --hints--

`getWeekday("2025-11-06")` bi trebalo da vrati `Thursday`.

```js
assert.equal(getWeekday("2025-11-06"), "Thursday");
```

`getWeekday("1999-12-31")` bi trebalo da vrati `Friday`.

```js
assert.equal(getWeekday("1999-12-31"), "Friday");
```

`getWeekday("1111-11-11")` bi trebalo da vrati `Saturday`.

```js
assert.equal(getWeekday("1111-11-11"), "Saturday");
```

`getWeekday("2112-12-21")` bi trebalo da vrati `Wednesday`.

```js
assert.equal(getWeekday("2112-12-21"), "Wednesday");
```

`getWeekday("2345-10-01")` bi trebalo da vrati `Monday`.

```js
assert.equal(getWeekday("2345-10-01"), "Monday");
```

# --seed--

## --seed-contents--

```js
function getWeekday(dateString) {

  return dateString;
}
```

# --solutions--

```js
function getWeekday(dateString) {
  const days = ["Sunday", "Monday", "Tuesday", "Wednesday", "Thursday", "Friday", "Saturday"];
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  return days[date.getUTCDay()];
}
```