id: 68ee9e3066cfd4eb2328e8a5
title: "Vežba 86: Pretraga slika"
challengeType: 28
dashedName: challenge-86
---

# --description--

4. novembra 2001. godine, Google je lansirao svoju pretragu slika, omogućavajući ljudima da pronađu slike koristeći ključne reči za pretraživanje. U ovom izazovu, simuliraćete pretragu slika.

Dato vam je niz imena slika i reč za pretraživanje; vratite niz imena slika koje sadrže reč za pretraživanje.

- Ignorišite razlike između velikih i malih slova prilikom poređenja ključnih reči za pretraživanje.
- Vratite slike u istom redosledu kako su prikazane u nizu podataka unosa.

# --hints--

`imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog")` bi trebalo da vrati `["dog.png"]`.

```js
assert.deepEqual(imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog"), ["dog.png"]);
```

`imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun")` bi trebalo da vrati `["Sunset.jpg", "sunflower.jpeg"]`.

```js
assert.deepEqual(imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun"), ["Sunset.jpg", "sunflower.jpeg"]);
```

`imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG")` bi trebalo da vrati `["Moon.png", "stars.png"]`.

```js
assert.deepEqual(imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG"), ["Moon.png", "stars.png"]);
```

`imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat")` bi trebalo da vrati `["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]`.

```js
assert.deepEqual(imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat"), ["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]);
```

# --seed--

## --seed-contents--

```js
function imageSearch(images, term) {

  return images;
}
```

# --solutions--

```js
function imageSearch(images, term) {
  const lowerTerm = term.toLowerCase();
  return images.filter(img => img.toLowerCase().includes(lowerTerm));
}
```