id: 68ee9e3066cfd4eb2328e8a6
title: "Vežba 87: Izgradnja matrice"
challengeType: 28
dashedName: challenge-87
---

# --description--

Ako dobijete dva celobrojna broja (broj redova i broj kolona), vratite matricu (niz podataka redova sa podacima) ispunjenu nulama (`0`) dimenzija koje su date.

Na primer, ako dobijete `2` i `3`, vratiti:

```json
[
  [0, 0, 0],
  [0, 0, 0]
]
```

# --hints--

`buildMatrix(2, 3)` bi trebalo da vrati `[[0, 0, 0], [0, 0, 0]]`.

```js
assert.deepEqual(buildMatrix(2, 3), [[0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(3, 2)` bi trebalo da vrati `[[0, 0], [0, 0], [0, 0]]`.

```js
assert.deepEqual(buildMatrix(3, 2), [[0, 0], [0, 0], [0, 0]]);
```

`buildMatrix(4, 3)` bi trebalo da vrati `[[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]`.

```js
assert.deepEqual(buildMatrix(4, 3), [[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]);
```

`buildMatrix(9, 1)` bi trebalo da vrati `[[0], [0], [0], [0], [0], [0], [0], [0], [0]]`.

```js
assert.deepEqual(buildMatrix(9, 1), [[0], [0], [0], [0], [0], [0], [0], [0], [0]]);
```

# --seed--

## --seed-contents--

```js
function buildMatrix(rows, cols) {

  return rows;
}
```

# --solutions--

```js
function buildMatrix(rows, cols) {
  const matrix = [];
  for (let i = 0; i < rows; i++) {
    const row = new Array(cols).fill(0);
    matrix.push(row);
  }
  return matrix;
}
```