id: 68ee9e3066cfd4eb2328e8a4
title: "Vežba 85: Brojanje reči"
challengeType: 28
dashedName: challenge-85
---

# --description--

Kada dobijete niz znakova rečenice, vratite broj reči u toj rečenici.

- Reč je bilo koji niz ne-razmak znakova i razdvojena je samo jednim razmakom.

# --hints--

`countWords("Hello world")` bi trebalo da vrati `2`.

```js
assert.equal(countWords("Hello world"), 2);
```

`countWords("The quick brown fox jumps over the lazy dog.")` bi trebalo da vrati `9`.

```js
assert.equal(countWords("The quick brown fox jumps over the lazy dog."), 9);
```

`countWords("I like coding challenges!")` bi trebalo da vrati `4`.

```js
assert.equal(countWords("I like coding challenges!"), 4);
```

`countWords("Complete the challenge in JavaScript and Python.")` bi trebalo da vrati `7`.

```js
assert.equal(countWords("Complete the challenge in JavaScript and Python."), 7);
```

`countWords("The missing semi-colon crashed the entire internet.")` bi trebalo da vrati `7`.

```js
assert.equal(countWords("The missing semi-colon crashed the entire internet."), 7);
```

# --seed--

## --seed-contents--

```js
function countWords(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countWords(sentence) {

  return sentence.split(' ').length;
}
```