id: 68ee9e3066cfd4eb2328e8a8
title: "Vežba 89: Računanje karti"
challengeType: 28
dashedName: challenge-89
---

# --description--

Standardni paket igračkih karata ima 13 jedinstvenih karata u svakoj šarki. Ako dobijete ceo broj koji predstavlja ukupan broj karata iz paketa, vratite broj jedinstvenih kombinacija karata koje možete odabrati.

- Redos ne utiče na rezultat. Odabir karte A pa karte B isto je kao i izbor karte B pa karte A.

Na primer, ako dobijete `52`, vratite `1`. Postoji samo jedna kombinacija od 52 karata iz paketa od 52 karata. I ako dobijete `2`, vratite `1326`. Postoji 1326 kombinacija karata koje možete dobiti kada odaberete 2 karte iz paketa.

# --hints--

`combinations(52)` treba da vrati `1`.

```js
assert.equal(combinations(52), 1);
```

`combinations(1)` treba da vrati `52`.

```js
assert.equal(combinations(1), 52);
```

`combinations(2)` treba da vrati `1326`.

```js
assert.equal(combinations(2), 1326);
```

`combinations(5)` treba da vrati `2598960`.

```js
assert.equal(combinations(5), 2598960);
```

`combinations(10)` treba da vrati `15820024220`.

```js
assert.equal(combinations(10), 15820024220);
```

`combinations(50)` treba da vrati `1326`.

```js
assert.equal(combinations(50), 1326);
```

# --seed--

## --seed-contents--

```js
function combinations(cards) {

  return cards;
}
```

# --solutions--

```js
function combinations(cards) {
  const n = 52;

  function factorial(x) {
    let result = 1;
    for (let i = 2; i <= x; i++) {
      result *= i;
    }
    return result;
  }

  return factorial(n) / (factorial(cards) * factorial(n - cards));
}
```