id: 68f6587287ad1f4ad39b0c7c
title: "Vežba 90: Karakter limit"
challengeType: 28
dashedName: challenge-90
---

# --description--

U ovom izazovu, dat vam je niz karaktera i morate utvrditi da li odgovara objavi na društvenim mrežama. Vratite sledeći niz karaktera prema zadatim pravilima:

- `"short post"` ako odgovara karakter limitu od 40.
- `"long post"` ako je duži od 40 karaktera i odgovara karakter limitu od 80.
- `"invalid post"` ako je preduga da bi odgovarala bilo kom od dva limita.

# --hints--

`canPost("Hello world")` treba da vrati `"short post"`.

```js
assert.equal(canPost("Hello world"), "short post");
```

`canPost("This is a longer message but still under eighty characters.")` treba da vrati `"long post"`.

```js
assert.equal(canPost("This is a longer message but still under eighty characters."), "long post");
```

`canPost("This message is too long to fit into either of the character limits for a social media post.")` treba da vrati `"invalid post"`.

```js
assert.equal(canPost("This message is too long to fit into either of the character limits for a social media post."), "invalid post");
```

# --seed--

## --seed-contents--

```js
function canPost(message) {

  return message;
}
```

# --solutions--

```js
function canPost(message) {
  if (message.length <= 40) {
    return "short post"
  } else if (message.length <= 80) {
    return "long post"
  } else {
    return "invalid post"
  }
}
```