id: 68f6587287ad1f4ad39b0c7f
title: "Vežba 93: Vokali i Konsonanti"
challengeType: 28
dashedName: challenge-93
---

# --description--

Kada dobijete niz slova, vratite niz podataka sa brojem vokala i brojem konsonanti u tom nizu.

- Vokali su `a`, `e`, `i`, `o`, `u` u bilo kom slučaju.
- Konsonanti su sva ostala slova u bilo kom slučaju.
- Ignorišite nije-alfabetička slova.

Na primer, ako dobijete `"Hello World"`, vratite `[3, 7]`.

# --hints--

`count("Hello World")` bi trebalo da vrati `[3, 7]`.

```js
assert.deepEqual(count("Hello World"), [3, 7]);
```

`count("JavaScript")` bi trebalo da vrati `[3, 7]`.

```js
assert.deepEqual(count("JavaScript"), [3, 7]);
```

`count("Python")` bi trebalo da vrati `[1, 5]`.

```js
assert.deepEqual(count("Python"), [1, 5]);
```

`count("freeCodeCamp")` bi trebalo da vrati `[5, 7]`.

```js
assert.deepEqual(count("freeCodeCamp"), [5, 7]);
```

`count("Hello, World!")` bi trebalo da vrati `[3, 7]`.

```js
assert.deepEqual(count("Hello, World!"), [3, 7]);
```

`count("The quick brown fox jumps over the lazy dog.")` bi trebalo da vrati `[11, 24]`.

```js
assert.deepEqual(count("The quick brown fox jumps over the lazy dog."), [11, 24]);
```

# --seed--

## --seed-contents--

```js
function count(str) {

  return str;
}
```

# --solutions--

```js
function count(str) {
  const vowels = 'aeiou';
  const consonants = 'bcdfghjklmnpqrstvwxyz';
  let v = 0, c = 0;

  for (let i=0; i<str.length; i++) {
    if (vowels.includes(str[i].toLowerCase())) {
      v++;
    }
    if (consonants.includes(str[i].toLowerCase())) {
      c++;
    }
  }

  return [v, c];
}
```