id: 68f6587287ad1f4ad39b0c7e
title: "Vežba 92: Ekstraktor ekstenzija"
challengeType: 28
dashedName: challenge-92
---

# --description--

Kada je dato niz slova koji predstavlja ime fajla, vrati ekstenziju tog fajla.

- Ekstenzija je deo imena fajla koji sledi nakon poslednje tačke (`.`).
- Ako ime fajla nema tačku ili se završava tačkom, vrati `"none"`.
- Ekstenzija mora biti vraćena onakva kakva jeste, sa očuvanim velikim i malim slovima.

# --hints--

`getExtension("document.txt")` bi trebalo da vrati `"txt"`.

```js
assert.equal(getExtension("document.txt"), "txt");
```

`getExtension("README")` bi trebalo da vrati `"none"`.

```js
assert.equal(getExtension("README"), "none");
```

`getExtension("image.PNG")` bi trebalo da vrati `"PNG"`.

```js
assert.equal(getExtension("image.PNG"), "PNG");
```

`getExtension(".gitignore")` bi trebalo da vrati `"gitignore"`.

```js
assert.equal(getExtension(".gitignore"), "gitignore");
```

`getExtension("archive.tar.gz")` bi trebalo da vrati `"gz"`.

```js
assert.equal(getExtension("archive.tar.gz"), "gz");
```

`getExtension("final.draft.")` bi trebalo da vrati `"none"`.

```js
assert.equal(getExtension("final.draft."), "none");
```

# --seed--

## --seed-contents--

```js
function getExtension(filename) {

  return filename;
}
```

# --solutions--

```js
function getExtension(filename) {
  const lastDot = filename.lastIndexOf('.');
  
  if (lastDot === -1 || lastDot === filename.length - 1) {
    return "none";
  }
  
  return filename.slice(lastDot + 1);
}
```