id: 68f6587287ad1f4ad39b0c84
title: "Vežba 98: Računanje pravougaonika"
challengeType: 28
dashedName: challenge-98
---

# --description--

Ako dobijete dva pozitivna celobrojka koja predstavljaju širinu i visinu pravougaonika, utvrdite koliko je pravougaonika može biti uloženo unutar datog.

- Izračunajte samo broj pravougaonika sa celobrojnom širinom i visinom.

Na primer, ako dobijete `1` i `3`, vratite `6`. Tri pravougaonika od 1x1, dva pravougaonika od 1x2, i jedan pravougaonik od 1x3.

# --hints--

`countRectangles(1, 3)` bi trebalo da vrati `6`.

```js
assert.equal(countRectangles(1, 3), 6);
```

`countRectangles(3, 2)` bi trebalo da vrati `18`.

```js
assert.equal(countRectangles(3, 2), 18);
```

`countRectangles(1, 2)` bi trebalo da vrati `3`.

```js
assert.equal(countRectangles(1, 2), 3);
```

`countRectangles(5, 4)` bi trebalo da vrati `150`.

```js
assert.equal(countRectangles(5, 4), 150);
```

`countRectangles(11, 19)` bi trebalo da vrati `12540`.

```js
assert.equal(countRectangles(11, 19), 12540);
```

# --seed--

## --seed-contents--

```js
function countRectangles(width, height) {

  return width;
}
```

# --solutions--

```js
function countRectangles(width, height) {
  let count = 0;

  for (let w = 1; w <= width; w++) {
    for (let h = 1; h <= height; h++) {
      count += (width - w + 1) * (height - h + 1);
    }
  }

  return count;
}
```