id: 68f6587287ad1f4ad39b0c7d
title: "Vežba 91: Pronađi reč"
challengeType: 28
dashedName: challenge-91
---

# --description--

Ako vam je data matrica (niz redova sa pojedinačnim slovima) i reč koju treba pronaći, vratite početne i krajnje koordinate te reči u matrici.

- Data matrica je popunjena svim malim slovima (`a-z`).
- Reč koju treba pronaći će se pojaviti samo jednom u matrici.
- Reč koju treba pronaći će biti vidljiva na pravoj liniji u jednom od sledećih pravaca:
  - od leve ka desnoj
  - od desne ka levoj
  - odozgo nadole
  - odozdo nagore

Na primer, ako dobijete matricu:

```md
[
  ["a", "c", "t"],
  ["t", "a", "t"],
  ["c", "t", "c"]
]
```

Za reč `"cat"`, vratite:

```md
[[0, 1], [2, 1]]
```

Gde su `[0, 1]` početne koordinate za `"c"` (početak reči), a `[2, 1]` krajnje koordinate za `"t"` (kraj reči).

# --hints--

`findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat")` bi trebalo da vrati `[[0, 1], [2, 1]]`.

```js
assert.deepEqual(findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat"), [[0, 1], [2, 1]]);
```

`findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog")` bi trebalo da vrati `[[0, 0], [0, 2]]`.

```js
assert.deepEqual(findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog"), [[0, 0], [0, 2]]);
```

`findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish")` bi trebalo da vrati `[[3, 3], [0, 3]]`.

```js
assert.deepEqual(findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish"), [[3, 3], [0, 3]]);
```

`findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox")` bi trebalo da vrati `[[1, 3], [1, 1]]`.

```js
assert.deepEqual(findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox"), [[1, 3], [1, 1]]);
```

# --seed--

## --seed-contents--

```js
function findWord(matrix, word) {

  return matrix;
}
```

# --solutions--

```js
function findWord(matrix, word) {
  const rows = matrix.length;
  const cols = matrix[0].length;
  const len = word.length;

  const directions = [
    [0, 1],
    [0, -1],
    [1, 0],
    [-1, 0]
  ];

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      for (let [dr, dc] of directions) {
        let match = true;
        for (let i = 0; i < len; i++) {
          const nr = r + dr * i;
          const nc = c + dc * i;
          if (nr < 0 || nr >= rows || nc < 0 || nc >= cols || matrix[nr][nc] !== word[i]) {
            match = false;
            break;
          }
        }
        if (match) {
          return [
            [r, c],
            [r + dr * (len - 1), c + dc * (len - 1)]
          ];
        }
      }
    }
  }
}
```