id: 68f6587287ad1f4ad39b0c83
title: "Vežba 97: GCD"
challengeType: 28
dashedName: challenge-97
---

# --description--

Kada se uzmu dva pozitivna celobrojka, vratite njihov najveći zajednički delilitelj (GCD).

- GCD dva celobrojka je najveći broj koji deli oba broja bez ostataka.

Na primer, delitelji `4` su `1`, `2` i `4`. Delitelji `6` su `1`, `2`, `3` i `6`. Stoga, kada uzmete u obzir `4` i `6`, vratite `2`, najveći broj koji postoji u oba skupa delitelja.

# --hints--

`gcd(4, 6)` bi trebalo da vrati `2`.

```js
assert.equal(gcd(4, 6), 2);
```

`gcd(20, 15)` bi trebalo da vrati `5`.

```js
assert.equal(gcd(20, 15), 5);
```

`gcd(13, 17)` bi trebalo da vrati `1`.

```js
assert.equal(gcd(13, 17), 1);
```

`gcd(654, 456)` bi trebalo da vrati `6`.

```js
assert.equal(gcd(654, 456), 6);
```

`gcd(3456, 4320)` bi trebalo da vrati `864`.

```js
assert.equal(gcd(3456, 4320), 864);
```

# --seed--

## --seed-contents--

```js
function gcd(x, y) {

  return x;
}
```

# --solutions--

```js
function gcd(x, y) {
  while (y !== 0) {
    const temp = y;
    y = x % y;
    x = temp;
  }
  return x;
}
```