id: 68f6587287ad1f4ad39b0c81
title: "Vežba 95: Pomeranje niza podataka"
challengeType: 28
dashedName: challenge-95
---

# --description--

Ako vam je dat niz podataka i ceo broj koji predstavlja koliko pozicija treba pomeriti taj niz, vratite pomereni niz.

- Pozitivan ceo broj pomeraće niz ulevo (levo).
- Negativan ceo broj pomeraće niz desno (desno).
- Pomeranje se vrši unutar niza podataka (ciklično).

Na primer, ako dobijete `[1, 2, 3]` i `1`, pomerite niz za 1 ulevo, vraćajući `[2, 3, 1]`.

# --hints--

`shiftArray([1, 2, 3], 1)` treba da vrati `[2, 3, 1]`.

```js
assert.deepEqual(shiftArray([1, 2, 3], 1), [2, 3, 1]);
```

`shiftArray([1, 2, 3], -1)` treba da vrati `[3, 1, 2]`.

```js
assert.deepEqual(shiftArray([1, 2, 3], -1), [3, 1, 2]);
```

`shiftArray(["alpha", "bravo", "charlie"], 5)` treba da vrati `["charlie", "alpha", "bravo"]`.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], 5), ["charlie", "alpha", "bravo"]);
```

`shiftArray(["alpha", "bravo", "charlie"], -11)` treba da vrati `["bravo", "charlie", "alpha"]`.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], -11), ["bravo", "charlie", "alpha"]);
```

`shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15)` treba da vrati `[5, 6, 7, 8, 9, 0, 1, 2, 3, 4]`.

```js
assert.deepEqual(shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15), [5, 6, 7, 8, 9, 0, 1, 2, 3, 4]);
```

# --seed--

## --seed-contents--

```js
function shiftArray(arr, n) {

  return arr;
}
```

# --solutions--

```js
function shiftArray(arr, n) {
  const len = arr.length;
  n = n % len;
  if (n < 0) n += len;

  return arr.slice(n).concat(arr.slice(0, n));
}
```