id: 68f6587287ad1f4ad39b0c80
title: "Vežba 94: Generator e-mail potpisa"
challengeType: 28
dashedName: challenge-94
---

# --description--

Kada vam je dat niz slova imena osobe, titula i kompanija, vratite e-mail potpis kao jedan niz slova koristeći sledeće pravilo:

- Ime mora biti navedeno prvo, praćeno prefiksom koji zavisi od prvog slova imena. Za imena koja počinju sa (bez obzira na veliko ili malo slovo):
  - `A-I`: Koristite `>>` kao prefiks.
  - `J-R`: Koristite `--` kao prefiks.
  - `S-Z`: Koristite `::` kao prefiks.
- Zarez i razmak (`, `) mora slediti ime.
- Titula i kompanija moraju slediti zarez i razmak, razdvojeni sa `" at "` (sa razmacima na obe strane).

Na primer, ako vam je dat `"Quinn Waverly"`, `"Founder and CEO"` i `"TechCo"`, vratite `"--Quinn Waverly, Founder and CEO at TechCo"`.

# --hints--

`generateSignature("Quinn Waverly", "Founder and CEO", "TechCo")` bi trebalo da vrati `"--Quinn Waverly, Founder and CEO at TechCo"`.

```js
assert.equal(generateSignature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo");
```

`generateSignature("Alice Reed", "Engineer", "TechCo")` bi trebalo da vrati `">>Alice Reed, Engineer at TechCo"`.

```js
assert.equal(generateSignature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo");
```

`generateSignature("Tina Vaughn", "Developer", "example.com")` bi trebalo da vrati `"::Tina Vaughn, Developer at example.com"`.

```js
assert.equal(generateSignature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com");
```

`generateSignature("B. B.", "Product Tester", "AcmeCorp")` bi trebalo da vrati `">>B. B., Product Tester at AcmeCorp"`.

```js
assert.equal(generateSignature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp");
```

`generateSignature("windstorm", "Cloud Architect", "Atmospheronics")` bi trebalo da vrati `"::windstorm, Cloud Architect at Atmospheronics"`.

```js
assert.equal(generateSignature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics");
```

# --seed--

## --seed-contents--

```js
function generateSignature(name, title, company) {

  return name;
}
```

# --solutions--

```js
function generateSignature(name, title, company) {
  const firstLetter = name[0].toUpperCase();
  let prefix;
  if ("ABCDEFGHI".includes(firstLetter)) {
    prefix = ">>";
  } else if ("JKLMNOPQR".includes(firstLetter)) {
    prefix = "--";
  } else {
    prefix = "::";
  }

  return `${prefix}${name}, ${title} at ${company}`;
}
```