id: 68f6587287ad1f4ad39b0c82
title: "Vežba 96: Da li je vikend?"
challengeType: 28
dashedName: challenge-96
---

# --description--

Ako primite datum u formatu `"YYYY-MM-DD"`, vratite broj dana pre kraja nedelje.

- Nedelja počinje u subotu.
- Ako je dato traženi datum subota ili nedelja, vratite `"It's the weekend!"`.
- Inače, vratite `"X days until the weekend."`, gde je `X` broj dana do subote.
- Ako je `X` jednako `1`, koristite `"day"` (jednina) umesto `"days"` (množina).
- Uverite se da proračun ne uzima u obzir vremenske zone.

# --hints--

`daysUntilWeekend("2025-11-14")` bi trebalo da vrati `"1 day until the weekend."`.

```js
assert.equal(daysUntilWeekend("2025-11-14"), "1 day until the weekend.");
```

`daysUntilWeekend("2025-01-01")` bi trebalo da vrati `"3 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2025-01-01"), "3 days until the weekend.");
```

`daysUntilWeekend("2025-12-06")` bi trebalo da vrati `"It's the weekend!"`.

```js
assert.equal(daysUntilWeekend("2025-12-06"), "It's the weekend!");
```

`daysUntilWeekend("2026-01-27")` bi trebalo da vrati `"4 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2026-01-27"), "4 days until the weekend.");
```

`daysUntilWeekend("2026-09-07")` bi trebalo da vrati `"5 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2026-09-07"), "5 days until the weekend.");
```

`daysUntilWeekend("2026-11-29")` bi trebalo da vrati `"It's the weekend!"`.

```js
assert.equal(daysUntilWeekend("2026-11-29"), "It's the weekend!");
```

# --seed--

## --seed-contents--

```js
function daysUntilWeekend(dateString) {

  return dateString;
}
```

# --solutions--

```js
function daysUntilWeekend(dateString) {
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  const dayOfWeek = date.getUTCDay();

  if (dayOfWeek === 6 || dayOfWeek === 0) {
    return "It's the weekend!";
  }

  const daysUntilSaturday = (6 - dayOfWeek);
  return `${daysUntilSaturday} day${daysUntilSaturday === 1 ? "" : "s"} until the weekend.`;
}
```