id: 68ffb91507a5b645769328cb
title: "Vežba 108: BuzzFizz"
challengeType: 28
dashedName: challenge-108
---

# --description--

Iz niza podataka, proverite da li je ispravan FizzBuzz niz od 1 do poslednog elementa u nizu. Niz je ispravan ako:

- Savršeni brojevi deljivi sa 3 zamenjeni su sa "Fizz"
- Savršeni brojevi deljivi sa 5 zamenjeni su sa "Buzz"
- Savršeni brojevi deljivi i sa 3 i sa 5 zamenjeni su sa "FizzBuzz"
- Svi ostali savršeni brojevi ostaju kao potpuni brojevi u rastućem redosledu, počevši od 1.
- Niz mora početi sa 1 i ne sme imati nedostajuće ili višak elemente.

# --hints--

`isFizzBuzz([1, 2, "Fizz", 4])` bi trebalo da vrati `true`.

```js
assert.isTrue(isFizzBuzz([1, 2, "Fizz", 4]));
```

`isFizzBuzz([1, 2, 3, 4])` bi trebalo da vrati `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, 3, 4]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", 7])` bi trebalo da vrati `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", 7]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "FizzBuzz"])` bi trebalo da vrati `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "FizzBuzz"]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Fizz"])` bi trebalo da vrati `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Fizz"]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Buzz"])` bi trebalo da vrati `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Buzz"]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"])` bi trebalo da vrati `true`.

```js
assert.isTrue(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"]));
```

# --seed--

## --seed-contents--

```js
function isFizzBuzz(sequence) {

  return sequence;
}
```

# --solutions--

```js
function isFizzBuzz(sequence) {
  for (let i = 0; i < sequence.length; i++) {
    const n = i + 1;
    if (n % 15 === 0 && sequence[i] !== "FizzBuzz") return false;
    if (n % 3 === 0 && n % 5 !== 0 && sequence[i] !== "Fizz") return false;
    if (n % 5 === 0 && n % 3 !== 0 && sequence[i] !== "Buzz") return false;
    if (n % 3 !== 0 && n % 5 !== 0 && sequence[i] !== n) return false;
  }
  return true;
}
```