id: 68ffb91507a5b645769328cc
title: "Vežba 109: Koliko sam star?"
challengeType: 28
dashedName: challenge-109
---

# --description--

Date rođenja osobe u formatu `YYYY-MM-DD`, vratite godine te osobe do dana 27. novembra 2025.

- Pretpostavite da su svi datumi rođenja važeći datumi pre 27. novembra 2025.
- Vratite godine kao ceo broj.
- Osigurajte da uzmete u obzir da li je osoba već proslavila rođendan u 2025. godini.

# --hints--

`calculateAge("2000-11-20")` treba da vrati `25`.

```js
assert.equal(calculateAge("2000-11-20"), 25);
```

`calculateAge("2000-12-01")` treba da vrati `24`.

```js
assert.equal(calculateAge("2000-12-01"), 24);
```

`calculateAge("2014-10-25")` treba da vrati `11`.

```js
assert.equal(calculateAge("2014-10-25"), 11);
```

`calculateAge("1994-01-06")` treba da vrati `31`.

```js
assert.equal(calculateAge("1994-01-06"), 31);
```

`calculateAge("1994-12-14")` treba da vrati `30`.

```js
assert.equal(calculateAge("1994-12-14"), 30);
```

# --seed--

## --seed-contents--

```js
function calculateAge(birthday) {

  return birthday;
}
```

# --solutions--

```js
function calculateAge(birthday) {
  const today = new Date("2025-11-27");
  const [year, month, day] = birthday.split("-").map(Number);
  const birthDate = new Date(year, month - 1, day);

  let age = today.getFullYear() - birthDate.getFullYear();

  if (
    today.getMonth() < birthDate.getMonth() ||
    (today.getMonth() === birthDate.getMonth() && today.getDate() < birthDate.getDate())
  ) {
    age--;
  }

  return age;
}
```