id: 68ffb91507a5b645769328c3
title: "Vežba 100: 100 slova"
challengeType: 28
dashedName: challenge-100
---

# --description--

Dobrodošli u dnevni izazov kodiranja od 100 znakova!

Ako dobijete niz slova, ponavljajte ta slova sve dok rezultat ne sadrži tačno 100 znakova. Ako vaše ponavljanje pređe 100 znakova, isecite višak kako bi ostalo samo 100 znakova.

# --hints--

`oneHundred("One hundred ")` treba da vrati `"One hundred One hundred One hundred One hundred One hundred One hundred One hundred One "`.

```js
assert.equal(oneHundred("One hundred "), "One hundred One hundred One hundred One hundred One hundred One hundred One hundred One ");
```

`oneHundred("freeCodeCamp ")` treba da vrati `"freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC"`.

```js
assert.equal(oneHundred("freeCodeCamp "), "freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeCamp freeCodeC");
```

`oneHundred("daily challenges ")` treba da vrati `"daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge"`.

```js
assert.equal(oneHundred("daily challenges "), "daily challenges daily challenges daily challenges daily challenges daily challenges daily challenge");
```

`oneHundred("!")` treba da vrati `"!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!"`.

```js
assert.equal(oneHundred("!"), "!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!!");
```

# --seed--

## --seed-contents--

```js
function oneHundred(chars) {

  return chars;
}
```

# --solutions--

```js
function oneHundred(chars) {
  let result = "";
  let i = 0;

  while (result.length < 100) {
    result += chars[i % chars.length];
    i++;
  }

  return result.slice(0, 100);
}
```