id: 68ffb91507a5b645769328c9
title: "Vežba 106: Validator poruke"
challengeType: 28
dashedName: challenge-106
---

# --description--

Kada dobijete niz slova poruke i niz slova validatora, odredite da li je poruka validna.

- Poruka je validna ako svaka reč u poruci počinje slovom koje odgovara sekvenci slova validatora, po redu.
- Slova ne razmatraju razliku između velikih i malih slova.
- Reči u poruci su razdvojene pojedinačnim razmakom.

# --hints--

`isValidMessage("hello world", "hw")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidMessage("hello world", "hw"));
```

`isValidMessage("ALL CAPITAL LETTERS", "acl")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidMessage("ALL CAPITAL LETTERS", "acl"));
```

`isValidMessage("Coding challenge are boring.", "cca")` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidMessage("Coding challenge are boring.", "cca"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT")` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT"));
```

# --seed--

## --seed-contents--

```js
function isValidMessage(message, validator) {

  return message;
}
```

# --solutions--

```js
function isValidMessage(message, validation) {
  const words = message.split(" ");
  
  if (words.length !== validation.length) return false;

  for (let i = 0; i < words.length; i++) {
    if (words[i][0].toLowerCase() !== validation[i].toLowerCase()) {
      return false;
    }
  }

  return true;
}
```