id: 68ffb91507a5b645769328c5
title: "Vežba 102: Najduža reč"
challengeType: 28
dashedName: challenge-102
---

# --description--

Iz sekvence znakova u rečenici, vratite najdužu reč iz te rečenice.

- Reči su razdvojene samo jednim razmakom.
- Računaju se samo slova (`a-z`, bez obzira na veliko ili malo).
- Ako postoji više reči iste dužine, vratite prvu naiđenu.
- Vratite reč kako je data u sekvenci, bez interpunkcijskih znakova.

# --hints--

`longestWord("The quick red fox")` bi trebalo da vrati `"quick"`.

```js
assert.equal(longestWord("The quick red fox"), "quick");
```

`longestWord("Hello coding challenge.")` bi trebalo da vrati `"challenge"`.

```js
assert.equal(longestWord("Hello coding challenge."), "challenge");
```

`longestWord("Do Try This At Home.")` bi trebalo da vrati `"This"`.

```js
assert.equal(longestWord("Do Try This At Home."), "This");
```

`longestWord("This sentence... has commas, ellipses, and an exclamation point!")` bi trebalo da vrati `"exclamation"`.

```js
assert.equal(longestWord("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation");
```

`longestWord("A tie? No way!")` bi trebalo da vrati `"tie"`.

```js
assert.equal(longestWord("A tie? No way!"), "tie");
```

`longestWord("Wouldn't you like to know.")` bi trebalo da vrati `"Wouldnt"`.

```js
assert.equal(longestWord("Wouldn't you like to know."), "Wouldnt");
```

# --seed--

## --seed-contents--

```js
function longestWord(sentence) {

  return sentence;
}
```

# --solutions--

```js
function longestWord(sentence) {
  const words = sentence.split(" ");
  let longest = "";

  for (let word of words) {
    const cleaned = word.replace(/[^a-z]/gi, "");
    if (cleaned.length > longest.length) {
      longest = cleaned;
    }
  }

  return longest;
}
```