id: 68ffb91507a5b645769328c6
title: "Vežba 103: LCM"
challengeType: 28
dashedName: challenge-103
---

# --description--

Kada dva celobrojna broja, vratite najmanji zajednički sadržalj (LCM) tih brojeva.

LCM dva broja je najmanji pozitivan ceo broj koji je sadržalj oba broja. Na primer, kada razmotrite `4` i `6`, vratite `12` zato što:

- Sadržalji za `4` su `4`, `8`, `12` i tako dalje.
- Sadržalji za `6` su `6`, `12`, `18` i tako dalje.
- `12` je najmanji broj koji je sadržalj oba.

# --hints--

`lcm(4, 6)` bi trebalo da vrati `12`.

```js
assert.equal(lcm(4, 6), 12);
```

`lcm(9, 6)` bi trebalo da vrati `18`.

```js
assert.equal(lcm(9, 6), 18);
```

`lcm(10, 100)` bi trebalo da vrati `100`.

```js
assert.equal(lcm(10, 100), 100);
```

`lcm(13, 17)` bi trebalo da vrati `221`.

```js
assert.equal(lcm(13, 17), 221);
```

`lcm(45, 70)` bi trebalo da vrati `630`.

```js
assert.equal(lcm(45, 70), 630);
```

# --seed--

## --seed-contents--

```js
function lcm(a, b) {

  return a;
}
```

# --solutions--

```js
function lcm(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  return Math.abs(a * b) / gcd(a, b);
}
```