id: 68ffb91507a5b645769328c4
title: "Vežba 101: Konverter za naslov sekcije u Markdown"
challengeType: 28
dashedName: challenge-101
---

# --description--

Ako vam je dat niz znakova koji predstavlja Markdown naslov sekcije, vratite odgovarajući HTML naslov sekcije.

Validan Markdown naslov sekcije mora:

- Početi sa nula ili više razmaka, praćenih
- 1 do 6 uzastopnih hash znaka (`#`), zatim
- Bar jednim razmakom. I na kraju,
- Tekstom naslova sekcije.

Broj hash znakova određuje nivo naslova sekcije. Na primer, jedan hash znak odgovara tagu `h1`, a šest hash znakova odgovara tagu `h6`.

Ako dati niz znakova ne odgovara tačno strukturi gore navedeni, vratite `"Invalid format"`.

Na primer, ako primite `"# My level 1 heading"`, vratite `"<h1>My level 1 heading</h1>"`.

Napomena: Konzola možda neće prikazati HTML tagove unutar niza znakova prilikom logovanja poruka. Proverite konzolu pretraživača da vidite zapise sa ugrađenim tagovima.

# --hints--

`convert("# My level 1 heading")` bi trebalo da vrati `"<h1>My level 1 heading</h1>"`.

```js
assert.equal(convert("# My level 1 heading"), "<h1>My level 1 heading</h1>");
```

`convert("My heading")` bi trebalo da vrati `"Invalid format"`.

```js
assert.equal(convert("My heading"), "Invalid format");
```

`convert("##### My level 5 heading")` bi trebalo da vrati `"<h5>My level 5 heading</h5>"`.

```js
assert.equal(convert("##### My level 5 heading"), "<h5>My level 5 heading</h5>");
```

`convert("#My heading")` bi trebalo da vrati `"Invalid format"`.

```js
assert.equal(convert("#My heading"), "Invalid format");
```

`convert("  ###  My level 3 heading")` bi trebalo da vrati `"<h3>My level 3 heading</h3>"`.

```js
assert.equal(convert("  ###  My level 3 heading"), "<h3>My level 3 heading</h3>");
```

`convert("####### My level 7 heading")` bi trebalo da vrati `"Invalid format"`.

```js
assert.equal(convert("####### My level 7 heading"), "Invalid format");
```

`convert("## My #2 heading")` bi trebalo da vrati `"<h2>My #2 heading</h2>"`.

```js
assert.equal(convert("## My #2 heading"), "<h2>My #2 heading</h2>");
```

# --seed--

## --seed-contents--

```js
function convert(heading) {

  return heading;
}
```

# --solutions--

```js
function convert(heading) {
  const trimmed = heading.trim()

  if (!/^#{1,6}\s+\S+/.test(trimmed)) return "Invalid format"

  const hashCount = trimmed.indexOf(' ');
  const text = trimmed.slice(hashCount).trimStart();

  return `<h${hashCount}>${text}</h${hashCount}>`;
}
```