id: 69162d64f96574d9bb629efd
title: "Vežba 112: Detektor AI"
challengeType: 28
dashedName: challenge-112
---

# --description--

Danasni izazov potiče iz lansiranja ChatGPT-a 30. novembra 2022.

Analizirajući niz slova jedne ili više rečenica, utvrdite da li je verovatno kreiran od strane AI-ja koristeći sledeće pravilo:

- Ima dve ili više crtica (`-`).
- Ima dve ili više parazisa `()`. Tekst može biti unutar zagrada.
- Ima tri ili više reči sa 7 slova ili više.

- Reči su razdvojene samo jednim razmakom i sadrže samo slova (`A-Z`). Ne uključujte interpunkcijske znakove ili ne-slovočne karaktere kao deo reči.

Ako zadana rečenica ispunjava bilo koje od ovih pravila, vratite `"AI"`, u suprotnom, vratite `"Human"`.

# --hints--

`detectAI("The quick brown fox jumped over the lazy dog.")` bi trebalo da vrati `"Human"`.

```js
assert.equal(detectAI("The quick brown fox jumped over the lazy dog."), "Human");
```

`detectAI("The hypersonic brown fox - jumped (over) the lazy dog.")` bi trebalo da vrati `"Human"`.

```js
assert.equal(detectAI("The hypersonic brown fox - jumped (over) the lazy dog."), "Human");
```

`detectAI("Yes - you're right! I made a mistake there - let me try again.")` bi trebalo da vrati `"AI"`.

```js
assert.equal(detectAI("Yes - you're right! I made a mistake there - let me try again."), "AI");
```

`detectAI("The extraordinary students were studying vivaciously.")` bi trebalo da vrati `"AI"`.

```js
assert.equal(detectAI("The extraordinary students were studying vivaciously."), "AI");
```

`detectAI("The (excited) student was (coding) in the library.")` bi trebalo da vrati `"AI"`.

```js
assert.equal(detectAI("The (excited) student was (coding) in the library."), "AI");
```

# --seed--

## --seed-contents--

```js
function detectAI(text) {

  return text;
}
```

# --solutions--

```js
function detectAI(text) {
  const dashCount = (text.match(/-/g) || []).length;
  if (dashCount >= 2) return "AI";

  const parenCount = (text.match(/\([^)]*\)/g) || []).length;
  if (parenCount >= 2) return "AI";

  const words = text.split(" ");
  const longWordCount = words.filter(w => w.replace(/[^A-Za-z]/g, "").length >= 7).length;
  if (longWordCount >= 3) return "AI";

  return "Human";
}
```