id: 69162d64f96574d9bb629f01
title: "Vežba 116: Izračunavanje permutacija"
challengeType: 28
dashedName: challenge-116
---

# --description--

Ako vam je dat niz slova, vratite broj jedinstvenih permutacija koje se mogu stvoriti iz tih slova.

- Permutacija je bilo kakvo preuređenje slova u nizu.
- Ne računajte ponavljajuće permutacije.
- Ako niz sadrži ponavljajuća slova, ponavljajući aranžmani moraju biti uračunati samo jednom.
- Niz će sadržavati samo slova (`A-Z`, `a-z`).

Na primer, ako dobijete `"abb"`, vratite `3` jer postoji tri jedinstvena načina za poređanje slova: `"abb"`, `"bab"` i `"bba"`.

# --hints--

`countPermutations("abb")` treba da vrati `3`.

```js
assert.equal(countPermutations("abb"), 3);
```

`countPermutations("abc")` treba da vrati `6`.

```js
assert.equal(countPermutations("abc"), 6);
```

`countPermutations("racecar")` treba da vrati `630`.

```js
assert.equal(countPermutations("racecar"), 630);
```

`countPermutations("freecodecamp")` treba da vrati `39916800`.

```js
assert.equal(countPermutations("freecodecamp"), 39916800);
```

# --seed--

## --seed-contents--

```js
function countPermutations(str) {

  return str;
}
```

# --solutions--

```js
function countPermutations(str) {
  const freq = {};
  for (const ch of str) {
    freq[ch] = (freq[ch] || 0) + 1;
  }

  function factorial(n) {
    let result = 1;
    for (let i = 2; i <= n; i++) result *= i;
    return result;
  }

  const n = str.length;
  let result = factorial(n);

  for (const ch in freq) {
    result /= factorial(freq[ch]);
  }

  return result;
}
```