id: 69162d64f96574d9bb629efe
title: "Vežba 113: Milje u kilometre"
challengeType: 28
dashedName: challenge-113
---

# --description--

Kada unesete udaljenost u miljama kao broj, vratite istu udaljenost u kilometrima.

- Ulaz će uvek biti pozitivan broj.
- 1 milja je jednaka 1.60934 kilometra.
- Zaokruži rezultat na dve decimalne za mestu.
- Uklonite nepotrebne nule sa kraja zaokruženog rezultata.

# --hints--

`convertToKm(1)` bi trebalo da vrati `1.61`.

```js
assert.equal(convertToKm(1), 1.61);
```

`convertToKm(21)` bi trebalo da vrati `33.8`.

```js
assert.equal(convertToKm(21), 33.8);
```

`convertToKm(3.5)` bi trebalo da vrati `5.63`.

```js
assert.equal(convertToKm(3.5), 5.63);
```

`convertToKm(0)` bi trebalo da vrati `0`.

```js
assert.equal(convertToKm(0), 0);
```

`convertToKm(0.621371)` bi trebalo da vrati `1`.

```js
assert.equal(convertToKm(0.621371), 1);
```

# --seed--

## --seed-contents--

```js
function convertToKm(miles) {

  return miles;
}
```

# --solutions--

```js
function convertToKm(miles) {
  const km = miles * 1.60934;
  return Math.round(km * 100) / 100;
}
```