id: 69162d64f96574d9bb629efb
title: "Izazov 110: Prediktor reči"
challengeType: 28
dashedName: challenge-110
---

# --description--

Kada uzmete niz od slova iste dužine, tajnu reč i pogađanje (guess), uporedite pogađanje sa tajnom reči koristeći sledeće pravilo:

- Tajna i pogađanja će sadržati samo velika slova (`"A"` do `"Z"`).
- Za svako slovo u pogađanju, zamijenite ga brojem ovisno o tome kako odgovara tajnoj reči:
  - `"2"` ako je slovo prisutno u tajnoj reči i na ispravnom mjestu.
  - `"1"` ako je slovo prisutno u tajnoj reči, ali na pogrešnom mjestu.
  - `"0"` ako slovo nije prisutno u tajnoj reči.
- Svako slovo u tajnoj reči može se koristiti samo jednom.
- Prvo se dodjeljuje potpuna podudarnost (`"2"`), a zatim djelimična podudarnost (`"1"`) se dodjeljuje od lijeva ka desno za preostala slova.
- Ako se slovo često pojavljuje u pogađanju, može se podudarati samo onoliko puta koliko se pojavljuje u tajnoj reči.

Na primjer, uzimajući u obzir tajnu riječ `"APPLE"` i pogađanje `"POPPA"`, vratite `"10201"`:

Prvo slovo `"P"` nije na pogrešnom mjestu (`"1"`), `"O"` nije prisutno u tajnoj reči (`"0"`), drugo slovo `"P"` je na ispravnom mjestu (`"2"`), treće slovo `"P"` je nula (`"0"`) jer su dva slova `"P"` u tajnoj reči korištena, a `"A"` nije na pogrešnom mjestu (`"1"`).

# --hints--

`compare("APPLE", "POPPA")` treba da vrati `"10201"`.

```js
assert.equal(compare("APPLE", "POPPA"), "10201");
```

`compare("REACT", "TRACE")` treba da vrati `"11221"`.

```js
assert.equal(compare("REACT", "TRACE"), "11221");
```

`compare("DEBUGS", "PYTHON")` treba da vrati `"000000"`.

```js
assert.equal(compare("DEBUGS", "PYTHON"), "000000");
```

`compare("JAVASCRIPT", "TYPESCRIPT")` treba da vrati `"0000222222"`.

```js
assert.equal(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222");
```

`compare("ORANGE", "ROUNDS")` treba da vrati `"110200"`.

```js
assert.equal(compare("ORANGE", "ROUNDS"), "110200");
```

`compare("WIRELESS", "ETHERNET")` treba da vrati `"10021000"`.

```js
assert.equal(compare("WIRELESS", "ETHERNET"), "10021000");
```

# --seed--

## --seed-contents--

```js
function compare(word, guess) {

  return word;
}
```

# --solutions--

```js
function compare(word, guess) {
  const n = word.length;
  let result = Array(n).fill("0");
  let used = Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    if (guess[i] === word[i]) {
      result[i] = "2";
      used[i] = true;
    }
  }

  for (let i = 0; i < n; i++) {
    if (result[i] === "0") {
      for (let j = 0; j < n; j++) {
        if (!used[j] && guess[i] === word[j]) {
          result[i] = "1";
          used[j] = true;
          break;
        }
      }
    }
  }

  return result.join("");
}
```