id: 69162d64f96574d9bb629f04
title: "Vežba 119: Kompresija sekvence reči"
challengeType: 28
dashedName: challenge-119
---

# --description--

Ako dobijete rečenicu sekvence reči, vratite komprimovanu verziju te rečenice gde se uzastopno ponavljaju reči zamenjuju sledećom reči i brojem puta koliko se ponavlja unutar uglatih zagrada.

- Komprimuje se samo uzastopno ponavljaju reči.
- Reči su razdvojene pojedinačnim razmakom.

Na primer, ako dobijete `"yes yes yes please"`, vratite `"yes(3) please"`.

# --hints--

`compressString("yes yes yes please")` bi trebalo da vrati `"yes(3) please"`.

```js
assert.equal(compressString("yes yes yes please"), "yes(3) please");
```

`compressString("I have have have apples")` bi trebalo da vrati `"I have(3) apples"`.

```js
assert.equal(compressString("I have have have apples"), "I have(3) apples");
```

`compressString("one one three and to the the the the")` bi trebalo da vrati `"one(2) three and to the(4)"`.

```js
assert.equal(compressString("one one three and to the the the the"), "one(2) three and to the(4)");
```

`compressString("route route route route route route tee tee tee tee tee tee")` bi trebalo da vrati `"route(6) tee(6)"`.

```js
assert.equal(compressString("route route route route route route tee tee tee tee tee tee"), "route(6) tee(6)");
```

# --seed--

## --seed-contents--

```js
function compressString(sentence) {

  return sentence;
}
```

# --solutions--

```js
function compressString(sentence) {
  const words = sentence.split(" ");
  const result = [];
  let count = 1;

  for (let i = 0; i < words.length; i++) {
    if (words[i] === words[i + 1]) {
      count++;
    } else {
      if (count > 1) {
        result.push(`${words[i]}(${count})`);
      } else {
        result.push(words[i]);
      }
      count = 1;
    }
  }

  return result.join(" ");
}
```