id: 69162d64f96574d9bb629efc
title: "Vežba 111: Pravac lopte"
challengeType: 28
dashedName: challenge-111
---

# --description--

Danasni izazov potiče od video igre Pong, lansirane 29. novembra 1972. godine.

Uzimajući u obzir matricu (niz podataka redova sa podacima) koja sadrži poziciju lopte (`2`), i prethodnu poziciju lopte (`1`), vratite matrične koordinate sledeće pozicije lopte.

- Lopta se uvek kreće po pravoj liniji.
- Pravac kretanja zavisi od toga kako je lopta prešla sa `1` na `2`.
- Granice matrice smatraju se zidovima. Ako loptica udari:
  - o gornji ili donji zid, ona skokom menja svoj vertikalni pravac.
  - o levi ili desni zid, ona skokom menja svoj horizontalni pravac.
  - u ćošak, ona skokom menja oba pravca.

# --hints--

`getNextLocation([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]])` bi trebalo da vrati `[2, 3]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]]), [2, 3]);
```

`getNextLocation([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]])` bi trebalo da vrati `[3, 0]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]]), [3, 0]);
```

`getNextLocation([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]])` bi trebalo da vrati `[1, 2]`.

```js
assert.deepEqual(getNextLocation([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]]), [1, 2]);
```

`getNextLocation([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]])` bi trebalo da vrati `[1, 1]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]]), [1, 1]);
```

`getNextLocation([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]])` bi trebalo da vrati `[2, 2]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]]), [2, 2]);
```

# --seed--

## --seed-contents--

```js
function getNextLocation(matrix) {

  return matrix;
}
```

# --solutions--

```js
function getNextLocation(matrix) {
  let prev = null;
  let curr = null;

  for (let row = 0; row < matrix.length; row++) {
    for (let col = 0; col < matrix[0].length; col++) {
      if (matrix[row][col] === 1) prev = [row, col];
      if (matrix[row][col] === 2) curr = [row, col];
    }
  }

  const [prevRow, prevCol] = prev;
  const [currRow, currCol] = curr;

  let dirX = currCol - prevCol;
  let dirY = currRow - prevRow;

  let nextRow = currRow + dirY;
  let nextCol = currCol + dirX;

  const maxRow = matrix.length - 1;
  const maxCol = matrix[0].length - 1;


  if (nextCol < 0 || nextCol > maxCol) dirX *= -1;
  if (nextRow < 0 || nextRow > maxRow) dirY *= -1;

  nextRow = currRow + dirY;
  nextCol = currCol + dirX;

  return [nextRow, nextCol];
}
```