id: 69162d64f96574d9bb629f00
title: "Vežba 115: Konverter elementa numerisane liste u Markdown-u"
challengeType: 28
dashedName: challenge-115
---

# --description--

Ako vam je dat niz znakova koji predstavlja element numerisane liste u Markdown-u, vratite odgovarajući HTML niz znakova.

Validan element numerisane liste u Markdown-u mora da:

- Počinje sa nula ili više razmaka, praćenih
- Brojem (1 ili više) i tačkom (`.`), praćenim
- Bar jednim razmakom, zatim
- Tekstom elementa liste.

Ako niz znakova ne odgovara tačno gornjoj strukturi, vratite `"Invalid format"`. U suprotnom, obavijte tekst elementa liste unutar `li` oznaka i vratite niz znakova.

Na primer, ako primite `"1. My item"`, vratite `"<li>My item</li>"`.

Napomena: Konola možda neće prikazivati HTML oznake u nizu znakova prilikom logovanja poruka. Proverite konzolu pretraživača da vidite zapise sa oznakama.

# --hints--

`convertListItem("1. My item")` treba vratiti `"<li>My item</li>"`.

```js
assert.equal(convertListItem("1. My item"), "<li>My item</li>");
```

`convertListItem(" 1.  Another item")` treba vratiti `"<li>Another item</li>"`.

```js
assert.equal(convertListItem(" 1.  Another item"), "<li>Another item</li>");
```

`convertListItem("1 . invalid item")` treba vratiti `"Invalid format"`.

```js
assert.equal(convertListItem("1 . invalid item"), "Invalid format");
```

`convertListItem("2. list item text")` treba vratiti `"<li>list item text</li>"`.

```js
assert.equal(convertListItem("2. list item text"), "<li>list item text</li>");
```

`convertListItem(". invalid again")` treba vratiti `"Invalid format"`.

```js
assert.equal(convertListItem(". invalid again"), "Invalid format");
```

`convertListItem("A. last invalid")` treba vratiti `"Invalid format"`.

```js
assert.equal(convertListItem("A. last invalid"), "Invalid format");
```

# --seed--

## --seed-contents--

```js
function convertListItem(markdown) {

  return markdown;
}
```

# --solutions--

```js
function convertListItem(markdown) {
  const match = markdown.match(/^\s*(\d+)\.\s+(.+)$/);
  if (!match) return "Invalid format";

  const text = match[2];
  return `<li>${text}</li>`;
}
```