id: 69162d64f96574d9bb629eff
title: "Vežba 114: Od camel ka snake"
challengeType: 28
dashedName: challenge-114
---

# --description--

Ako vam je dat niz znakova u camel case formatu, pretvorite ga u snake case koristeći sledeća pravila:

- Ulazni niz će sadržati samo slova (`A-Z` i `a-z`) i uvek će počinjati malim slovom.
- Svako veliko slovo u nizu camel case formatu označava početak nove reči.
- Prebacite sva slova u mala.
- Razdvojte reči donjom crtom (underscore) (`_`).

# --hints--

`toSnake("helloWorld")` bi trebalo da vrati `"hello_world"`.

```js
assert.equal(toSnake("helloWorld"), "hello_world");
```

`toSnake("myVariableName")` bi trebalo da vrati `"my_variable_name"`.

```js
assert.equal(toSnake("myVariableName"), "my_variable_name");
```

`toSnake("freecodecampDailyChallenges")` bi trebalo da vrati `"freecodecamp_daily_challenges"`.

```js
assert.equal(toSnake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges");
```

# --seed--

## --seed-contents--

```js
function toSnake(camel) {

  return camel;
}
```

# --solutions--

```js
function toSnake(camel) {
  return camel
    .replace(/([A-Z])/g, "_$1")
    .toLowerCase();
}
```