id: 691b559495c5cb5a37b9b488
title: "Vežba 128: Brojanje suglasnika"
challengeType: 28
dashedName: challenge-128
---

# --description--

Uzimajući u obzir niz slova i ciljani broj, utvrdite da li taj niz ima tačan broj suglasnika kao ciljani broj.

- Suglasnici su sva slova abecede osim "a", "e", "i", "o" i "u" u bilo kom slučaju slova.
- Ignorišite brojeve, interpunkciju, razmake i sva druga ne-slovena slova prilikom brojanja.

# --hints--

`hasConsonantCount("helloworld", 7)` bi trebalo da vrati `true`.

```js
assert.isTrue(hasConsonantCount("helloworld", 7));
```

`hasConsonantCount("eieio", 5)` bi trebalo da vrati `false`.

```js
assert.isFalse(hasConsonantCount("eieio", 5));
```

`hasConsonantCount("freeCodeCamp Rocks!", 11)` bi trebalo da vrati `true`.

```js
assert.isTrue(hasConsonantCount("freeCodeCamp Rocks!", 11));
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)` bi trebalo da vrati `false`.

```js
assert.isFalse(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)
);
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)` bi trebalo da vrati `true`.

```js
assert.isTrue(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)
);
```

# --seed--

## --seed-contents--

```js
function hasConsonantCount(text, target) {
  return text;
}
```

# --solutions--

```js
function hasConsonantCount(text, target) {
  const vowels = new Set(["a", "e", "i", "o", "u"]);
  let count = 0;

  for (const char of text.toLowerCase()) {
    if (/[a-z]/.test(char) && !vowels.has(char)) {
      count++;
    }
  }

  return count === target;
}
```