id: 691b559495c5cb5a37b9b489
title: "Vežba 129: Markdown blokuot parser"
challengeType: 28
dashedName: challenge-129
---

# --description--

Ako vam je dat niz znakova koji sadrži blokuot u Markdownu, vratite odgovarajući HTML niz.

Blokuot u Markdownu je bilo koji red koji:

- Počinje sa nula ili više razmaka
- Praćen znakiem veće od (>)(`>`)
- Zatim, jedan ili više razmaka
- I na kraju, tekst blokuota.

Vratite tekst blokuota upakovan u otvarajuće i zatvarajuće HTML oznake `blockquote`.

Na primer, ako primite `"> This is a quote"`, vratite `<blockquote >This is a quote</blockquote>`.

Napomena: Konola možda neće prikazati HTML oznake u nizu znakova tokom snimanja poruka. Proverite konzolu pretraživača da vidite snimke sa priloženim oznakama.

# --hints--

`parseBlockquote("> This is a quote")` bi trebalo da vrati `"<blockquote>This is a quote</blockquote>"`.

```js
assert.equal(parseBlockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>");
```

`parseBlockquote(" > This is also a quote")` bi trebalo da vrati `"<blockquote>This is also a quote</blockquote>"`.

```js
assert.equal(parseBlockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>");
```

`parseBlockquote("  >    So  Is  This")` bi trebalo da vrati `"<blockquote>So  Is  This</blockquote>"`.

```js
assert.equal(parseBlockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>");
```

# --seed--

## --seed-contents--

```js
function parseBlockquote(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBlockquote(markdown) {
  const blockquoteRegex = /^\s*>\s+(.+)$/;
  const match = markdown.match(blockquoteRegex);

  const text = match[1];
  return `<blockquote>${text}</blockquote>`;
}
```