id: 691b559495c5cb5a37b9b480
title: "Vežba 120: Uncije u Kilograme"
challengeType: 28
dashedName: challenge-120
---

# --description--

Kada vam je težina data u unci kao broj, vratite niz znakova `"(lbs) pounds equals (kgs) kilograms."`.

- Zamijenite `"(lbs)"` sa ulaznim brojem.
- Zamijenite `"(kgs)"` vrijednošću konvertovanom u kilograme, zaokruženom na dvije decimale i uvijek uključujući dvije decimalne mjesta.
- 1 uncia je jednaka 0.453592 kilograma.
- Ako je ulaz `1`, koristite `"pound"` umjesto `"pounds"`.
- Ako je konvertovana vrijednost `1`, koristite `"kilogram"` umjesto `"kilograms"`.

# --hints--

`convertToKgs(1)` bi trebalo vratiti `"1 pound equals 0.45 kilograms."`.

```js
assert.equal(convertToKgs(1), "1 pound equals 0.45 kilograms.");
```

`convertToKgs(0)` bi trebalo vratiti `"0 pounds equals 0.00 kilograms."`.

```js
assert.equal(convertToKgs(0), "0 pounds equals 0.00 kilograms.");
```

`convertToKgs(100)` bi trebalo vratiti `"100 pounds equals 45.36 kilograms."`.

```js
assert.equal(convertToKgs(100), "100 pounds equals 45.36 kilograms.");
```

`convertToKgs(2.5)` bi trebalo vratiti `"2.5 pounds equals 1.13 kilograms."`.

```js
assert.equal(convertToKgs(2.5), "2.5 pounds equals 1.13 kilograms.");
```

`convertToKgs(2.20462)` bi trebalo vratiti `"2.20462 pounds equals 1.00 kilogram."`.

```js
assert.equal(convertToKgs(2.20462), "2.20462 pounds equals 1.00 kilogram.");
```

# --seed--

## --seed-contents--

```js
function convertToKgs(lbs) {
  return lbs;
}
```

# --solutions--

```js
function convertToKgs(lbs) {
  const KG_PER_POUND = 0.453592;
  const kgs = (lbs * KG_PER_POUND).toFixed(2);

  const poundWord = lbs === 1 ? "pound" : "pounds";
  const kilogramWord = kgs === "1.00" ? "kilogram" : "kilograms";

  return `${lbs} ${poundWord} equals ${kgs} ${kilogramWord}.`;
}
```