id: 691b559495c5cb5a37b9b482
title: "Vežba 122: Parser podebljanog teksta u Markdownu"
challengeType: 28
dashedName: challenge-122
---

# --description--

Ako vam je dat niz znakova koji može sadržati podebljani tekst u Markdownu, vratite ekvivalentni HTML niz.

- Podebljani tekst u Markdownu je bilo koji tekst okružen dva zvjezdice (`**`) ili dva donje crte (`__`).
- Ne postoji razmak između teksta i zvjezdica ili donjih crta, ali razmaci mogu postojati unutar samog teksta.
- Zamijenite sve instance podebljanog teksta HTML tagovima `b` i vratite niz.

Na primjer, ako primite `"**This is bold**"`, vratite `"<b>This is bold</b>"`.

Napomena: Konzola možda neće prikazati HTML tagove u nizu prilikom logovanja poruka. Provjerite konzolu preglednika kako biste vidjeli zapise sa priloženim tagovima.

# --hints--

`parseBold("**This is bold**")` bi trebalo vratiti `"<b>This is bold</b>"`.

```js
assert.equal(parseBold("**This is bold**"), "<b>This is bold</b>");
```

`parseBold("__This is also bold__")` bi trebalo vratiti `"<b>This is also bold</b>"`.

```js
assert.equal(parseBold("__This is also bold__"), "<b>This is also bold</b>");
```

`parseBold("**This is not bold **")` bi trebalo vratiti `"**This is not bold **"`.

```js
assert.equal(parseBold("**This is not bold **"), "**This is not bold **");
```

`parseBold("__ This is also not bold__")` bi trebalo vratiti `"__ This is also not bold__"`.

```js
assert.equal(parseBold("__ This is also not bold__"), "__ This is also not bold__");
```

`parseBold("The **quick** brown fox __jumps__ over the **lazy** dog.")` bi trebalo vratiti `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."`.

```js
assert.equal(parseBold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.");
```

# --seed--

## --seed-contents--

```js
function parseBold(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBold(markdown) {
  markdown = markdown.replace(/\*\*(\S(?:.*?\S)?)\*\*/g, "<b>$1</b>");
  markdown = markdown.replace(/__(\S(?:.*?\S)?)__/g, "<b>$1</b>");

  return markdown;
}
```