id: 691b559495c5cb5a37b9b485
title: "Izazov 125: Igra života"
challengeType: 28
dashedName: challenge-125
---

# --description--

Ako vam je data matrica (redovi redova) koja predstavlja trenutno stanje u Igri života Konweja, vratite sledeće stanje matrice koristeći ove pravila:

- Svaka ćelija je `1` (živa) ili `0` (mrtva).
- Komšije ćelije su do osam susednih ćelija (vertikalno, horizontalno i dijagonalno).
- Ćelije blizu granica imaju manje od osam komšija.

Pravila ažuriranja svake ćelije:

- Bilo koja živa ćelija sa manje od dva živa komšija umire (nedostatak populacije).
- Bilo koja živa ćelija sa dva ili tri živa komšija nastavlja da živi.
- Bilo koja živa ćelija sa više od tri živa komšija umire (prekomerna populacija).
- Bilo koja mrtva ćelija sa tačno tri živa komšija oživljava se (rođenje).

Na primer, ako je dato:

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

vratite:

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

Svaka ćelija se ažurira na osnovu broja živih komšija. Na primer, `[0][0]` ostaje mrtva (2 živa komšija), `[0][1]` ostaje živa (2 živa komšija), `[0][2]` umire (3 živa komšija), i tako dalje.

# --hints--

`gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` bi trebalo da vrati `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`.

```js
assert.deepEqual(gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]]);
```

`gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` bi trebalo da vrati `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`.

```js
assert.deepEqual(gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]);
```

`gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` bi trebalo da vrati `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`.

```js
assert.deepEqual(gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]]);
```

`gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` bi trebalo da vrati `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`.

```js
assert.deepEqual(gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function gameOfLife(grid) {
  return grid;
}
```

# --solutions--

```js
function gameOfLife(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const countLiveNeighbors = (r, c) => {
    let count = 0;
    for (let i = r - 1; i <= r + 1; i++) {
      for (let j = c - 1; j <= c + 1; j++) {
        if (
          i >= 0 &&
          i < rows &&
          j >= 0 &&
          j < cols &&
          !(i === r && j === c) &&
          grid[i][j] === 1
        ) {
          count++;
        }
      }
    }
    return count;
  };

  const next = grid.map(row => [...row]);

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const liveNeighbors = countLiveNeighbors(r, c);

      if (grid[r][c] === 1) {
        if (liveNeighbors < 2 || liveNeighbors > 3) next[r][c] = 0;
      } else {
        if (liveNeighbors === 3) next[r][c] = 1;
      }
    }
  }

  return next;
}
```