id: 691b559495c5cb5a37b9b483
title: "Vežba 123: Generator Rimskih brojeva"
challengeType: 28
dashedName: challenge-123
---

# --description--

Ako dobijete potpuno число, vratite njegovu ekvivalentnu vrednost u rimskim brojevima.

Rimski brojevi koriste ove simbole:

| Simbol | Vrednost |
| ------ | ----- |
| I      | 1     |
| V      | 5     |
| X      | 10    |
| L      | 50    |
| C      | 100   |
| D      | 500   |
| M      | 1000  |

Rimski brojevi se pišu od najvećeg ka najmanjem, od levi ka desni koristeći sledeća pravila:

- Dodavanje se koristi kada simbol sledi drugi sa jednakom ili manjom vrednošću. Na primer, `18` se piše kao `XVIII` (10 + 5 + 1 + 1 + 1 = 18).
- Oduzimanje se koristi kada je manji simbol vidljiv pre većeg, predstavljajući 4 ili 9 u bilo kojoj poziciji. Na primer, 19 se piše kao `XIX` (10 + (10 - 1)).
- Nijedan simbol ne sme biti ponovljen više od tri puta uzastopno. Oduzimanje se koristi kada želite da napišete više od tri simbola uzastopno. Stoga, najveći broj koji možete napisati je 3999.

Evo još jednog primera: ako dobijete `1464`, vratite `"MCDLXIV"` (1000 + (500 - 100) + 50 + 10 + (5 - 1)).

# --hints--

`toRoman(18)` bi trebalo da vrati `"XVIII"`.

```js
assert.equal(toRoman(18), "XVIII");
```

`toRoman(19)` bi trebalo da vrati `"XIX"`.

```js
assert.equal(toRoman(19), "XIX");
```

`toRoman(1464)` bi trebalo da vrati `"MCDLXIV"`.

```js
assert.equal(toRoman(1464), "MCDLXIV");
```

`toRoman(2025)` bi trebalo da vrati `"MMXXV"`.

```js
assert.equal(toRoman(2025), "MMXXV");
```

`toRoman(3999)` bi trebalo da vrati `"MMMCMXCIX"`.

```js
assert.equal(toRoman(3999), "MMMCMXCIX");
```

# --seed--

## --seed-contents--

```js
function toRoman(num) {
  return num;
}
```

# --solutions--

```js
function toRoman(num) {
  const symbols = [
    ["M", 1000],
    ["CM", 900],
    ["D", 500],
    ["CD", 400],
    ["C", 100],
    ["XC", 90],
    ["L", 50],
    ["XL", 40],
    ["X", 10],
    ["IX", 9],
    ["V", 5],
    ["IV", 4],
    ["I", 1]
  ];

  let result = "";
  for (const [sym, val] of symbols) {
    while (num >= val) {
      result += sym;
      num -= val;
    }
  }
  return result;
}
```