id: 691b559495c5cb5a37b9b487
title: "Vežba 127: Provera brzine"
challengeType: 28
dashedName: challenge-127
---

# --description--

Na osnovu brzine koju prelazite u milja po satu (MPH) i ograničenja brzine u kilometri po satu (KPH), utvrdite da li prelazite brzinu i da li ćete dobiti upozorenje ili kaznu.

- 1 mila je jednako 1.60934 kilometra.
- Ako putujete brzinom koja ne prelazi niti je jednaka ograničenju brzine, vratite `"Not Speeding"`.
- Ako putujete brzinom između KPH 5 i ispod ograničenja brzine, vratite `"Warning"`.
- Ako putujete brzinom vekom od KPH 5 iznad ograničenja brzine, vratite `"Ticket"`.

# --hints--

`speedCheck(30, 70)` bi trebalo da vrati `"Not Speeding"`.

```js
assert.equal(speedCheck(30, 70), "Not Speeding");
```

`speedCheck(40, 60)` bi trebalo da vrati `"Warning"`.

```js
assert.equal(speedCheck(40, 60), "Warning");
```

`speedCheck(40, 65)` bi trebalo da vrati `"Not Speeding"`.

```js
assert.equal(speedCheck(40, 65), "Not Speeding");
```

`speedCheck(60, 90)` bi trebalo da vrati `"Ticket"`.

```js
assert.equal(speedCheck(60, 90), "Ticket");
```

`speedCheck(65, 100)` bi trebalo da vrati `"Warning"`.

```js
assert.equal(speedCheck(65, 100), "Warning");
```

`speedCheck(88, 40)` bi trebalo da vrati `"Ticket"`.

```js
assert.equal(speedCheck(88, 40), "Ticket");
```

# --seed--

## --seed-contents--

```js
function speedCheck(speedMph, speedLimitKph) {
  return speedMph;
}
```

# --solutions--

```js
function speedCheck(speedMph, speedLimitKph) {
  const speedKph = speedMph * 1.60934;

  if (speedKph <= speedLimitKph) {
    return "Not Speeding";
  }

  const over = speedKph - speedLimitKph;

  if (over <= 5) {
    return "Warning";
  }

  return "Ticket";
}
```