id: 691b559495c5cb5a37b9b481
title: "Vežba 121: Najčešće ponavljajući element"
challengeType: 28
dashedName: challenge-121
---

# --description--

Kada dobijete niz podataka elemenata, vratite element koji se najčešće ponavlja.

- Uvek će postojati samo jedan element koji se najčešće ponavlja.

# --hints--

`mostFrequent(["a", "b", "a", "c"])` treba da vrati `"a"`.

```js
assert.equal(mostFrequent(["a", "b", "a", "c"]), "a");
```

`mostFrequent([2, 3, 5, 2, 6, 3, 2, 7, 2, 9])` treba da vrati `2`.

```js
assert.equal(mostFrequent([2, 3, 5, 2, 6, 3, 2, 7, 2, 9]), 2);
```

`mostFrequent([true, false, "false", "true", false])` treba da vrati `false`.

```js
assert.isFalse(mostFrequent([true, false, "false", "true", false]));
```

`mostFrequent([40, 20, 70, 30, 10, 40, 10, 50, 40, 60])` treba da vrati `40`.

```js
assert.equal(mostFrequent([40, 20, 70, 30, 10, 40, 10, 50, 40, 60]), 40);
```

# --seed--

## --seed-contents--

```js
function mostFrequent(arr) {
  return arr;
}
```

# --solutions--

```js
function mostFrequent(arr) {
  const freq = {};
  let maxCount = 0;
  let mostElem = null;

  for (const el of arr) {
    freq[el] = (freq[el] || 0) + 1;
    if (freq[el] > maxCount) {
      maxCount = freq[el];
      mostElem = el;
    }
  }

  return mostElem;
}
```