id: 691b559495c5cb5a37b9b486
title: "Vežba 126: Kapitalizacija"
challengeType: 28
dashedName: challenge-126
---

# --description--

Kada dobijete niz znakova nazvan `title`, vratite novi niz formatiran u stilu velikih slova koristeći sledeća pravila:

- Pišite prvo slovo svake reči velikim slovom.
- Svi ostali slovi u svakoj reči učinite malim slovima.
- Reči se uvek razdvajaju samo jednim razmakom.

# --hints--

`titleCase("hello world")` bi trebalo da vrati `"Hello World"`.

```js
assert.equal(titleCase("hello world"), "Hello World");
```

`titleCase("the quick brown fox")` bi trebalo da vrati `"The Quick Brown Fox"`.

```js
assert.equal(titleCase("the quick brown fox"), "The Quick Brown Fox");
```

`titleCase("JAVASCRIPT AND PYTHON")` bi trebalo da vrati `"Javascript And Python"`.

```js
assert.equal(titleCase("JAVASCRIPT AND PYTHON"), "Javascript And Python");
```

`titleCase("AvOcAdO tOAst fOr brEAkfAst")` bi trebalo da vrati `"Avocado Toast For Breakfast"`.

```js
assert.equal(titleCase("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast");
```

# --seed--

## --seed-contents--

```js
function titleCase(title) {
  return title;
}
```

# --solutions--

```js
function titleCase(title) {
  return title
    .split(" ")
    .map(word => word.charAt(0).toUpperCase() + word.slice(1).toLowerCase())
    .join(" ");
}
```