id: 6925e2068081f40f549ced1d
title: "Vežba 139: Kamen, Papir, Nožare"
challengeType: 28
dashedName: challenge-139
---

# --description--

Dato je niz od dva slova, gde prvi predstavlja Igrača 1, a drugi Igrača 2, i odredite pobednika u igri Kamen, Papir, Nožare.

- String sa slovima unosa će uvek biti `"Rock"`, `"Paper"` ili `"Scissors"`.
- `"Rock"` pobjeđuje `"Scissors"`.
- `"Paper"` pobjeđuje `"Rock"`.
- `"Scissors"` pobjeđuje `"Paper"`.

Vratite:

- `"Player 1 wins"` ako Igrač 1 pobeđuje.
- `"Player 2 wins"` ako Igrač 2 pobeđuje.
- `"Tie"` ako oba igrača izaberu isti izbor.

# --hints--

`rockPaperScissors("Rock", "Rock")` should return `"Tie"`.

```js
assert.equal(rockPaperScissors("Rock", "Rock"), "Tie");
```

`rockPaperScissors("Rock", "Paper")` should return `"Player 2 wins"`.

```js
assert.equal(rockPaperScissors("Rock", "Paper"), "Player 2 wins");
```

`rockPaperScissors("Scissors", "Paper")` should return `"Player 1 wins"`.

```js
assert.equal(rockPaperScissors("Scissors", "Paper"), "Player 1 wins");
```

`rockPaperScissors("Rock", "Scissors")` should return `"Player 1 wins"`.

```js
assert.equal(rockPaperScissors("Rock", "Scissors"), "Player 1 wins");
```

`rockPaperScissors("Scissors", "Scissors")` should return `"Tie"`.

```js
assert.equal(rockPaperScissors("Scissors", "Scissors"), "Tie");
```

`rockPaperScissors("Scissors", "Rock")` should return `"Player 2 wins"`.

```js
assert.equal(rockPaperScissors("Scissors", "Rock"), "Player 2 wins");
```

# --seed--

## --seed-contents--

```js
function rockPaperScissors(player1, player2) {

  return player1;
}
```

# --solutions--

```js
function rockPaperScissors(player1, player2) {
  if (player1 === player2) return "Tie";

  if (
    (player1 === "Rock" && player2 === "Scissors") ||
    (player1 === "Paper" && player2 === "Rock") ||
    (player1 === "Scissors" && player2 === "Paper")
  ) {
    return "Player 1 wins";
  } else {
    return "Player 2 wins";
  }
}
```