id: 6925e2068081f40f549ced1b
title: "Vežba 137: Generator snježne čarolije"
challengeType: 28
dashedName: challenge-137
---

# --description--

Ako vam je dat niz znakova sa više redova koji koristi znake za novi red (`\n`) za predstavljanje razmaknih linija, vratite novi niz gde je svaki red ogledano i spojen na kraj originalnog reda.

- Ogledajte red tako što ćete obrnuti ceo niz znakova, uključujući razmake.

Na primer, ako primite `"* \n *\n* "`, što se ispisuje u konzoli kao:

```sh
* 
 *
* 
```

Vratite `"*  *\n ** \n*  *"`, što se ispisuje u konzoli kao:

```sh
*  *
 ** 
*  *
```

Obratite posebnu pažnju na razmake u datom nizu i onome koji vraćate. Uverite se da ne izbacujete nikakve razmake.

# --hints--

`generateSnowflake("* \n *\n* ")` bi trebalo vratiti `"*  *\n ** \n*  *"`.

```js
assert.equal(generateSnowflake("* \n *\n* "), "*  *\n ** \n*  *");
```

`generateSnowflake("X=~")` bi trebalo vratiti `"X=~~=X"`.

```js
assert.equal(generateSnowflake("X=~"), "X=~~=X");
```

`generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  ")` bi trebalo vratiti `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`.

```js
assert.equal(generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  "), " X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X ");
```

`generateSnowflake("*   *\n * * \n* * *\n * * \n*   *")` bi trebalo vratiti `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`.

```js
assert.equal(generateSnowflake("*   *\n * * \n* * *\n * * \n*   *"), "*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *");
```

`generateSnowflake("*  -\n * -\n*  -")` bi trebalo vratiti `"*  --  *\n * -- * \n*  --  *"`.

```js
assert.equal(generateSnowflake("*  -\n * -\n*  -"), "*  --  *\n * -- * \n*  --  *");
```

# --seed--

## --seed-contents--

```js
function generateSnowflake(crystals) {

  return crystals;
}
```

# --solutions--

```js
function generateSnowflake(crystals) {
  const lines = crystals.split("\n");

  const mirroredLines = lines.map(line => {
    const reversed = line.split("").reverse().join("");
    return line + reversed;
  });

  return mirroredLines.join("\n");
}
```