---
id: 6925e2068081f40f549ced1a
title: "Zoezi la 136: Mchambuzi wa picha wa Markdown"
challengeType: 28
dashedName: challenge-136
---

# --description--

Ukipata mfuatano wa herufi wa picha katika Markdown, rudisha mfuatano wa herufi sawa wa HTML.

Picha ya Markdown ina muundo ufuatao: `"![alt text](image_url)"`. Ambapo:

- `alt text` je opis slike (vrednost svojstva elementa za `alt`).
- `image_url` je URL izvor slike (vrednost svojstva elementa za `src`).

Rudisha mfuatano wa herufi wa lebo ya HTML `img` yenye `src` imewekwa kuwa URL ya picha na `alt` imewekwa kuwa maandishi ya alt.

Na primer, ako dobijete `"![Cute cat](cat.png)"` vratite `'<img src="cat.png" alt="Cute cat">'`;

- Osigurajte da su oznake, raspored atributa elementa, razmak i korišćenje navodnika u skladu sa gore.

Kumbuka: Konsoli huenda isionyeshe lebo za HTML katika mfuatano wa herufi wakati wa kurekodi ujumbe — angalia konsoli ya kivinjari kuona rekodi zenye lebo zikiwemo.

# --hints--

`parseImage("![Cute cat](cat.png)")` inapaswa kurudisha `'<img src="cat.png" alt="Cute cat">'`.

```js
assert.equal(parseImage("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">');
```

`parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` inapaswa kurudisha `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
assert.equal(parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">');
```

`parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` inapaswa kurudisha `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
assert.equal(parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">');
```

# --seed--

## --seed-contents--

```js
function parseImage(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseImage(markdown) {
  const regex = /!\[(.*?)\]\((.*?)\)/;
  const match = markdown.match(regex);
  if (!match) return "";
  const alt = match[1];
  const src = match[2];
  return `<img src="${src}" alt="${alt}">`;
}
```
