id: 6925e2068081f40f549ced1a
title: "Vežba 136: Parser Markdown slika"
challengeType: 28
dashedName: challenge-136
---

# --description--

Ako dobijete niz znakova za sliku u Markdownu, vratite ekvivalentni HTML niz znakova.

Markdown slika ima sledeći format: `"![alt text](image_url)"`. Gde:

- `alt text` je opis slike (vrednost `alt` atributa).
- `image_url` je izvorni URL slike (vrednost `src` atributa).

Vratite niz znakova HTML `<img>` taga sa `src` postavljenim na URL slike i `alt` postavljenim na alt tekst.

Na primer, ako dobijete `"![Cute cat](cat.png)"`, vratite `'<img src="cat.png" alt="Cute cat">'`;

- Uverite se da je tag, redosled atributa, razmak i upotreba navodnika identičan prethodno navedenom.

Napomena: Konzola možda neće prikazati HTML tagove u nizu znakova prilikom snimanja poruka — proverite konzolu pretraživača radi pregleda zapisa koji uključuju tagove.

# --hints--

`parseImage("![Cute cat](cat.png)")` bi trebalo da vrati `'<img src="cat.png" alt="Cute cat">'`.

```js
assert.equal(parseImage("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">');
```

`parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` bi trebalo da vrati `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
assert.equal(parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">');
```

`parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` bi trebalo da vrati `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
assert.equal(parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">');
```

# --seed--

## --seed-contents--

```js
function parseImage(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseImage(markdown) {
  const regex = /!\[(.*?)\]\((.*?)\)/;
  const match = markdown.match(regex);
  if (!match) return "";
  const alt = match[1];
  const src = match[2];
  return `<img src="${src}" alt="${alt}">`;
}
```