id: 6925e2068081f40f549ced1c
title: "Vežba 138: Zbir delitelja"
challengeType: 28
dashedName: challenge-138
---

# --description--

Kada vam je dat pozitivan ceo broj, vratite zbir svih njegovih delitelja.

- Delitelj je bilo koji ceo broj koji deli taj broj bez ostatka (ostatak je `0`).
- Brojte svaki delitelj samo jednom.

Na primer, ako dobijete `6`, vratite `12` jer su delitelji `6` brojevi `1`, `2`, `3` i `6`, a njihov zbir je `12`.

# --hints--

`sumDivisors(6)` bi trebalo da vrati `12`.

```js
assert.equal(sumDivisors(6), 12);
```

`sumDivisors(13)` bi trebalo da vrati `14`.

```js
assert.equal(sumDivisors(13), 14);
```

`sumDivisors(28)` bi trebalo da vrati `56`.

```js
assert.equal(sumDivisors(28), 56);
```

`sumDivisors(84)` bi trebalo da vrati `224`.

```js
assert.equal(sumDivisors(84), 224);
```

`sumDivisors(549)` bi trebalo da vrati `806`.

```js
assert.equal(sumDivisors(549), 806);
```

`sumDivisors(9348)` bi trebalo da vrati `23520`.

```js
assert.equal(sumDivisors(9348), 23520);
```

# --seed--

## --seed-contents--

```js
function sumDivisors(n) {

  return n;
}
```

# --solutions--

```js
function sumDivisors(n) {
  let sum = 0;
  for (let i = 1; i <= n; i++) {
    if (n % i === 0) {
      sum += i;
    }
  }
  return sum;
}
```