id: 69272dcf1c24b44fd79137c6
title: "Zadatak 143: Regex Parser za isticanje u Markdownu (Italics)"
challengeType: 28
dashedName: challenge-143
---

# --description--

Ako vam je dat niz znakova koji može sadržavati tekst u stilu italika u Markdownu, vratite odgovarajući HTML niz znakova.

- Italicki tekst u Markdownu je bilo koji tekst koji počinje i završava sa jednim zvjezdicom (`*`) ili jednom crtom (`_`).
- Nisu dozvoljeni razmaci između teksta i zvjezdice ili crte, ali razmaci mogu postojati unutar samog teksta.
- Prebacite sve instance italika u HTML tag `<i`> i vratite niz znakova.

Na primer, ako primite `"*This is italic*"`, vratite `"<i>This is italic</i>"`.

Napomena: Konola možda neće prikazivati HTML tagove unutar niza znakova tokom snimanja poruke. Proverite konzolu pretraživača da biste videli snimke sa tagovima uključeno.

# --hints--

`parseItalics("*This is italic*")` treba da vrati `"<i>This is italic</i>"`.

```js
assert.equal(parseItalics("*This is italic*"), "<i>This is italic</i>");
```

`parseItalics("_This is also italic_")` treba da vrati `"<i>This is also italic</i>"`.

```js
assert.equal(parseItalics("_This is also italic_"), "<i>This is also italic</i>");
```

`parseItalics("*This is not italic *")` treba da vrati `"*This is not italic *"`.

```js
assert.equal(parseItalics("*This is not italic *"), "*This is not italic *");
```

`parseItalics("_ This is also not italic_")` treba da vrati `"_ This is also not italic_"`.

```js
assert.equal(parseItalics("_ This is also not italic_"), "_ This is also not italic_");
```

`parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog.")` treba da vrati `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`.

```js
assert.equal(parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.");
```

# --seed--

## --seed-contents--

```js
function parseItalics(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseItalics(markdown) {
  return markdown.replace(/(\*|_)([^\s][^]*?[^\s])\1/g, '<i>$2</i>');
}
```