id: 69272dcf1c24b44fd79137c3
title: "Vežba 140: SCREAMING_SNAKE_CASE"
challengeType: 28
dashedName: challenge-140
---

# --description--

Dati je niz niz znakova koji predstavljaju ime varijable, vratite ime varijable konvertovano u SCREAMING_SNAKE_CASE.

Omogućenja imena varijabli će biti napisana u jednom od sljedećih formata:

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

U navedenim formatima, riječi su razdvojene donjom crtom (`_`), crticom (`-`), ili nova riječ počinje velikim slovom.

Da biste konvertovali u SCREAMING_SNAKE_CASE:

- Pretvorite sve slova u velika
- Razdvojte riječi donjom crtom (`_`)

# --hints--

`toScreamingSnakeCase("userEmail")` bi trebalo vratiti `"USER_EMAIL"`.

```js
assert.equal(toScreamingSnakeCase("userEmail"), "USER_EMAIL");
```

`toScreamingSnakeCase("UserPassword")` bi trebalo vratiti `"USER_PASSWORD"`.

```js
assert.equal(toScreamingSnakeCase("UserPassword"), "USER_PASSWORD");
```

`toScreamingSnakeCase("user_id")` bi trebalo vratiti `"USER_ID"`.

```js
assert.equal(toScreamingSnakeCase("user_id"), "USER_ID");
```

`toScreamingSnakeCase("user-address")` bi trebalo vratiti `"USER_ADDRESS"`.

```js
assert.equal(toScreamingSnakeCase("user-address"), "USER_ADDRESS");
```

`toScreamingSnakeCase("username")` bi trebalo vratiti `"USERNAME"`.

```js
assert.equal(toScreamingSnakeCase("username"), "USERNAME");
```

`toScreamingSnakeCase("my_variable_name")` bi trebalo vratiti `"MY_VARIABLE_NAME"`.

```js
assert.equal(toScreamingSnakeCase("my_variable_name"), "MY_VARIABLE_NAME");
```

# --seed--

## --seed-contents--

```js
function toScreamingSnakeCase(variableName) {

  return variableName;
}
```

# --solutions--

```js
function toScreamingSnakeCase(variableName) {
  let temp = variableName.replace(/[-_]+/g, ' ');
  temp = temp.replace(/([a-z0-9])([A-Z])/g, '$1 $2');
  const words = temp.trim().split(/\s+/);
  return words.join('_').toUpperCase();
}
```