id: 69272dcf1c24b44fd79137c5
title: "Zadatak 142: Sabiranje sekvence cifara"
challengeType: 28
dashedName: challenge-142
---

# --description--

Ako vam je dat niz znakova koji uključuje brojeve i druge znakove, vratite zbir svih brojeva prisutnih u tom nizu.

- Razmatrate brojeve koji slede jedan za drugim kao jednu celobrojnu cifru. Na primer, `"13"` se računa kao 13, a ne 1 + 3.
- Ignorišite sve znakove koji nisu brojevi.

# --hints--

`stringSum("3apples2bananas")` bi trebalo da vrati `5`.

```js
assert.equal(stringSum("3apples2bananas"), 5);
```

`stringSum("10cats5dogs2birds")` bi trebalo da vrati `17`.

```js
assert.equal(stringSum("10cats5dogs2birds"), 17);
```

`stringSum("125344")` bi trebalo da vrati `125344`.

```js
assert.equal(stringSum("125344"), 125344);
```

`stringSum("a1b20c300")` bi trebalo da vrati `321`.

```js
assert.equal(stringSum("a1b20c300"), 321);
```

`stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` bi trebalo da vrati `1653`.

```js
assert.equal(stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653);
```

# --seed--

## --seed-contents--

```js
function stringSum(str) {

  return str;
}
```

# --solutions--

```js
function stringSum(str) {
  const matches = str.match(/\d+/g); // find all consecutive digits
  return matches.reduce((sum, num) => sum + Number(num), 0);
}
```