---
id: 69272dcf1c24b44fd79137c6
title: Vežba 143: Markdown parser za naglašavanje/kurziv
challengeType: 28
dashedName: challenge-143
---

# --description--

Ikiwa umepewa mfuatano wa herufi ambao unaweza kuwa na maandishi yaliyo katika mtindo wa italiki katika Markdown, rudisha mfuatano wa herufi wa HTML unaolingana.

- Maandishi ya italiki katika Markdown ni maandishi yoyote yanayoanza na kumalizika na nyota moja (`*`) au mstari chini mmoja (`_`).
- Hakuna nafasi zinazoruhusiwa kati ya maandishi na nyota au mstari chini, lakini nafasi zinaweza kuwepo ndani ya maandishi yenyewe.
- Badilisha matukio yote ya italiki kuwa lebo za HTML `i` na rudisha mfuatano wa herufi.

Na primer, ako primite `"*This is italic*"`, vratite `"<i>This is italic</i>"`.

Napomena: Konsola možda neće prikazati tagove HTML unutar niza karaktera prilikom snimanja poruka. Proverite konzolu pretraživača da biste videli rekorde sa tim oznakama.

# --hints--

`parseItalics("*This is italic*")` inapaswa kurudisha `"<i>This is italic</i>"`.

```js
assert.equal(parseItalics("*This is italic*"), "<i>This is italic</i>");
```

`parseItalics("_This is also italic_")` inapaswa kurudisha `"<i>This is also italic</i>"`.

```js
assert.equal(parseItalics("_This is also italic_"), "<i>This is also italic</i>");
```

`parseItalics("*This is not italic *")` inapaswa kurudisha `"*This is not italic *"`.

```js
assert.equal(parseItalics("*This is not italic *"), "*This is not italic *");
```

`parseItalics("_ This is also not italic_")` inapaswa kurudisha `"_ This is also not italic_"`.

```js
assert.equal(parseItalics("_ This is also not italic_"), "_ This is also not italic_");
```

`parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog.")` inapaswa kurudisha `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`.

```js
assert.equal(parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.");
```

# --seed--

## --seed-contents--

```js
function parseItalics(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseItalics(markdown) {
  return markdown.replace(/(\*|_)([^\s][^]*?[^\s])\1/g, '<i>$2</i>');
}
```
