id: 69272dcf1c24b44fd79137c4
title: "Vežba 141: Dodavanje goriva"
challengeType: 28
dashedName: challenge-141
---

# --description--

Na osnovu trenutnog broja galona goriva u vašem avionu i potrebnog broja litara goriva za dolazak na odredište, utvrdite koliko dodatnih galona goriva treba dodati.

- 1 galon je jednak 3.78541 litri.
- Ako avion već ima dovoljno goriva, vratite `0`.
- Možete dodati samo cele galone.
- Ne uključujte decimalne brojeve u vraćenu brojku.

# --hints--

`fuelToAdd(0, 1)` treba vratiti `1`.

```js
assert.equal(fuelToAdd(0, 1), 1);
```

`fuelToAdd(5, 40)` treba vratiti `6`.

```js
assert.equal(fuelToAdd(5, 40), 6);
```

`fuelToAdd(10, 30)` treba vratiti `0`.

```js
assert.equal(fuelToAdd(10, 30), 0);
```

`fuelToAdd(896, 20500)` treba vratiti `4520`.

```js
assert.equal(fuelToAdd(896, 20500), 4520);
```

`fuelToAdd(1000, 50000)` treba vratiti `12209`.

```js
assert.equal(fuelToAdd(1000, 50000), 12209);
```

# --seed--

## --seed-contents--

```js
function fuelToAdd(currentGallons, requiredLiters) {

  return currentGallons;
}
```

# --solutions--

```js
function fuelToAdd(currentGallons, requiredLiters) {
  const litersPerGallon = 3.78541;
  const currentLiters = currentGallons * litersPerGallon;
  if (currentLiters >= requiredLiters) return 0;
  const litersNeeded = requiredLiters - currentLiters;
  return Math.ceil(litersNeeded / litersPerGallon);
}
```