id: 69306364df283fcaff2e1ada
title: "Vežba 149: vOwElcAsE"
challengeType: 28
dashedName: challenge-149
---

# --description--

Dato je niz znakova, vratite novi niz gde su sve samoglasnike konvertovane u velika slova, a svi ostali alfabetski znakovi konvertovani u mala.

- Samoglasnici su `"a"`, `"e"`, `"i"`, `"o"`, i `"u"` u bilo kom slučaju.
- Nealfabetanski znakovi se ne menjaju.

# --hints--

`vowelCase("vowelcase")` treba da vrati `"vOwElcAsE"`.

```js
assert.equal(vowelCase("vowelcase"), "vOwElcAsE");
```

`vowelCase("coding is fun")` treba da vrati `"cOdIng Is fUn"`.

```js
assert.equal(vowelCase("coding is fun"), "cOdIng Is fUn");
```

`vowelCase("HELLO, world!")` treba da vrati `"hEllO, wOrld!"`.

```js
assert.equal(vowelCase("HELLO, world!"), "hEllO, wOrld!");
```

`vowelCase("git cherry-pick")` treba da vrati `"gIt chErry-pIck"`.

```js
assert.equal(vowelCase("git cherry-pick"), "gIt chErry-pIck");
```

`vowelCase("HEAD~1")` treba da vrati `"hEAd~1"`.

```js
assert.equal(vowelCase("HEAD~1"), "hEAd~1");
```

# --seed--

## --seed-contents--

```js
function vowelCase(str) {

  return str;
}
```

# --solutions--

```js
function vowelCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (vowels.includes(char)) {
      result += char.toUpperCase();
    } else if (/[a-zA-Z]/.test(char)) {
      result += char.toLowerCase();
    } else {
      result += char;
    }
  }

  return result;
}
```