id: 69306364df283fcaff2e1ad6
title: "Vežba 145: N-ti Fibonacci broj"
challengeType: 28
dashedName: challenge-145
---

# --description--

Ako vam je dat celobrojni broj `n`, vratite broj na poziciji `n` u Fibonacci sekvenci.

Fibonacci sekvenca je niz brojeva gde je svaki broj zbir prethodnih dva. Prvih deset brojeva u sekvenci su `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

# --hints--

`nthFibonacci(4)` treba da vrati `2`.

```js
assert.equal(nthFibonacci(4), 2);
```

`nthFibonacci(10)` treba da vrati `34`.

```js
assert.equal(nthFibonacci(10), 34);
```

`nthFibonacci(15)` treba da vrati `377`.

```js
assert.equal(nthFibonacci(15), 377);
```

`nthFibonacci(40)` treba da vrati `63245986`.

```js
assert.equal(nthFibonacci(40), 63245986);
```

`nthFibonacci(75)` treba da vrati `1304969544928657`.

```js
assert.equal(nthFibonacci(75), 1304969544928657);
```

# --seed--

## --seed-contents--

```js
function nthFibonacci(n) {

  return n;
}
```

# --solutions--

```js
function nthFibonacci(n) {
  if (n === 1) return 0;
  if (n === 2) return 1;

  let a = 0;
  let b = 1;

  for (let i = 3; i <= n; i++) {
    const next = a + b;
    a = b;
    b = next;
  }

  return b;
}
```