id: 69306364df283fcaff2e1ad5
title: "Vežba 144: Niz rezolucije"
challengeType: 28
dashedName: challenge-144
---

# --description--

Ako dobijete niz podataka u redovima, gde svaki red predstavlja dan vaše aktivnosti vezane za rezoluciju i sadrži tri celobrojne broja: broj koraka koje ste prešli tog dana, minute provedene na ekranu tog dana i broj pročitanih stranica tog dana; utvrdite da li ste ispunili svoju rezoluciju.

- Prvi red predstavlja dan 1, drugi dan 2, i tako dalje.

Dan se smatra uspešnim ako su ostvareni sledeći uslovi:

- Prešli najmanje 10.000 koraka.
- Niste prešli 120 minuta korišćenja ekrana.
- Pročitali ste najmanje pet stranica.

Ako su svi navedeni dani uspešni, vratite `"Resolution on track: N day streak."` gde je `N` broj uspešnih dana.

Ako jedan ili više dana nisu uspešni, vratite `"Resolution failed on day X: N day streak."`. Gde je `X` broj prvog neuspešnog dana, a `N` je broj uspešnih dana pre prvog neuspešnog dana.

# --hints--

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]])` treba da vrati `"Resolution on track: 3 day streak."`

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]]), "Resolution on track: 3 day streak.");
```

`resolutionStreak([[10000, 120, 5], [10950, 121, 11]])` treba da vrati `"Resolution failed on day 2: 1 day streak."`

```js
assert.equal(resolutionStreak([[10000, 120, 5], [10950, 121, 11]]), "Resolution failed on day 2: 1 day streak.");
```

`resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]])` treba da vrati `"Resolution failed on day 3: 2 day streak."`

```js
assert.equal(resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]]), "Resolution failed on day 3: 2 day streak.");
```

`resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]])` treba da vrati `"Resolution failed on day 4: 3 day streak."`

```js
assert.equal(resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]]), "Resolution failed on day 4: 3 day streak.");
```

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]])` treba da vrati `"Resolution on track: 10 day streak."`

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]]), "Resolution on track: 10 day streak.");
```

# --seed--

## --seed-contents--

```js
function resolutionStreak(days) {

  return days;
}
```

# --solutions--

```js
function resolutionStreak(days) {
  let streak = 0;

  for (let i = 0; i < days.length; i++) {
    const [steps, screen, pages] = days[i];
    if (steps >= 10000 && screen <= 120 && pages >= 5) {
      streak++;
    } else {
      return `Resolution failed on day ${i + 1}: ${streak} day streak.`;
    }
  }

  return `Resolution on track: ${streak} day streak.`;
}
```