id: 69306364df283fcaff2e1ad8
title: "Vežba 147: Kalkulator za Godine skoka"
challengeType: 28
dashedName: challenge-147
---

# --description--

Ako dobijete kompletnu godinu broja, utvrdite da li je godina skokovna.

Godina je skokovna ako ispunjava sledeća pravila:

- Godina je deljiva sa 4, i
- Godina nije deljiva sa 100, osim ako
- Godina je deljiva sa 400.

# --hints--

`isLeapYear(2024)` bi trebalo da vrati `true`.

```js
assert.isTrue(isLeapYear(2024));
```

`isLeapYear(2023)` bi trebalo da vrati `false`.

```js
assert.isFalse(isLeapYear(2023));
```

`isLeapYear(2100)` bi trebalo da vrati `false`.

```js
assert.isFalse(isLeapYear(2100));
```

`isLeapYear(2000)` bi trebalo da vrati `true`.

```js
assert.isTrue(isLeapYear(2000));
```

`isLeapYear(1999)` bi trebalo da vrati `false`.

```js
assert.isFalse(isLeapYear(1999));
```

`isLeapYear(2040)` bi trebalo da vrati `true`.

```js
assert.isTrue(isLeapYear(2040));
```

`isLeapYear(2026)` bi trebalo da vrati `false`.

```js
assert.isFalse(isLeapYear(2026));
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

  return year;
}
```

# --solutions--

```js
function isLeapYear(year) {
    return ((year % 4 === 0 && year % 100 !== 0) || (year % 400 === 0))
}
```