id: 69306364df283fcaff2e1ad7
title: "Vežba 146: Sedište za osobu koja piše levom rukom u tabeli"
challengeType: 28
dashedName: challenge-146
---

# --description--

Dato je 4x2 tabela (redovi po redovima) koja predstavlja raspored sedišta za obroke, utvrdite koliko sedišta osoba koja piše levom rukom može zauzeti.

- Osoba koja piše levom rukom ne može da sedi na mestu gde bi osoba koja piše desnom rukom sedela na sedištu sa leve strane te osobe.

U datoj tabeli:

- `"R"` je sedište koje koristi osoba koja piše desnom rukom.
- `"L"` je sedište koje koristi osoba koja piše levom rukom.
- `"U"` je prazno sedište.
- Samo prazna sedišta su dostupna za zauzimanje.
- Sedišta u gornjem redu gledaju "dole", a sedišta u donjem redu gledaju "gore" (kao tabela), tako da zavisi da li je levo ili desno od pravca sedišta.
- Ugljena sedišta imaju samo jedno sedište pored sebe.

Na primer, u datoj tabeli:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

Gornje levu sedište ne može se zauzeti jer postoji osoba koja piše desnom rukom sa leve strane. Ostala dva prazna sedišta mogu se zauzeti jer nema osobe koja piše desnom rukom sa leve strane.

# --hints--

`findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` should return `2`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2);
```

`findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` should return `8`.

```js
assert.equal(findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8);
```

`findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` should return `0`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0);
```

`findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` should return `1`.

```js
assert.equal(findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1);
```

`findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` should return `5`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5);
```

# --seed--

## --seed-contents--

```js
function findLeftHandedSeats(table) {

  return table;
}
```

# --solutions--

```js
function findLeftHandedSeats(table) {
  let availableSeats = 0;

  const [topRow, bottomRow] = table;

  for (let i=0; i<topRow.length; i++) {
    if (topRow[i] === 'U') {
      if (i === 3) {
        availableSeats++;
      } else if (topRow[i+1] !== 'R') {
        availableSeats++;
      }
    }
  }

  for (let i=0; i<bottomRow.length; i++) {
    if (bottomRow[i] === 'U') {
      if (i === 0) {
        availableSeats++;
      } else if (bottomRow[i-1] !== 'R') {
        availableSeats++;
      }
    }
  }

  return availableSeats;
}
```