id: 69373793f5a867f769cde138
title: "Vežba 153: Tic-Tac-Toe"
challengeType: 28
dashedName: challenge-153
---

# --description--

Ako vam je data matrica 3x3 (matrica podataka redova podataka) koja predstavlja završenu igru Tic-Tac-Toe, odredite pobednika.

- Svaki element u datoj matrici je `"X"` ili `"O"`.

Igrač pobeđuje ako ima tri uzastopna karaktera - horizontalno, vertikalno ili dijagonalno.

Vratite:

- `"X wins"` ako imač X ima tri uzastopna karaktera.
- `"O wins"` ako imač O ima tri uzastopna karaktera.
- `"Draw"` ako nijedan igrač nema tri uzastopna karaktera.

# --hints--

`ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]])` bi trebalo da vrati `"X wins"`.

```js
assert.equal(ticTacToe([["X", "X", "X"], ["O", "O", "X"], ["O", "X", "O"]]), "X wins");
```

`ticTacToe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]])` bi trebalo da vrati `"X wins"`.

```js
assert.equal(ticTacToe([["O", "O", "X"], ["X", "O", "X"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]])` bi trebalo da vrati `"Draw"`.

```js
assert.equal(ticTacToe([["X", "O", "X"], ["O", "X", "O"], ["O", "X", "O"]]), "Draw");
```

`ticTacToe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]])` bi trebalo da vrati `"O wins"`.

```js
assert.equal(ticTacToe([["X", "X", "O"], ["X", "O", "O"], ["O", "O", "X"]]), "O wins");
```

`ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]])` bi trebalo da vrati `"X wins"`.

```js
assert.equal(ticTacToe([["X", "O", "O"], ["O", "X", "O"], ["O", "X", "X"]]), "X wins");
```

`ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]])` bi trebalo da vrati `"Draw"`.

```js
assert.equal(ticTacToe([["O", "X", "X"], ["X", "O", "O"], ["X", "O", "X"]]), "Draw");
```

# --seed--

## --seed-contents--

```js
function ticTacToe(board) {

  return board;
}
```

# --solutions--

```js
function ticTacToe(board) {
  const lines = [];

  for (let row of board) {
    lines.push(row);
  }

  for (let c = 0; c < 3; c++) {
    lines.push([board[0][c], board[1][c], board[2][c]]);
  }

  lines.push([board[0][0], board[1][1], board[2][2]]);
  lines.push([board[0][2], board[1][1], board[2][0]]);

  for (let line of lines) {
    if (line.every(cell => cell === "X")) return "X wins";
    if (line.every(cell => cell === "O")) return "O wins";
  }

  return "Draw";
}
```