id: 69373793f5a867f769cde135
title: "Vežba 150: Parsiranje neuređenog spiska u Markdownu"
challengeType: 28
dashedName: challenge-150
---

# --description--

Pretvoren niz znakova validnog neuređenog spiska u Markdownu, vratite odgovarajući HTML niz znakova.

Neuređeni spisak ima jedan ili više stavki spiska. Validna stavka spiska izgleda na sledećem redu i:

- Počinje sa crticom (`"-"`), praćeno
- Bar jednim razmakom, zatim
- Tekstom stavke spiska.

Spisak je dostavljen kao jedan niz znakova sa novim redovima odvojenima karakterom newline (`"\n"`). Ne uključujte karaktere novog reda u tekst stavke.

Omotajte svaku stavku spiska unutar HTML oznaka `li`, a ceo spisak stavki unutar oznaka `ul`.

Na primer, ako primite `"- Item A\n- Item B"`, vratite `"<ul><li>Item A</li><li>Item B</li></ul>"`.

Napomena: Konzola možda neće prikazati HTML oznake u nizu znakova tokom snimanja poruke. Proverite konzolu pretraživača da vidite zapise sa umetnutim oznakama.

# --hints--

`parseUnorderedList("- Item A\n- Item B")` treba da vrati `"<ul><li>Item A</li><li>Item B</li></ul>"`.

```js
assert.equal(parseUnorderedList("- Item A\n- Item B"), "<ul><li>Item A</li><li>Item B</li></ul>");
```

`parseUnorderedList("-  JavaScript\n-  Python")` treba da vrati `"<ul><li>JavaScript</li><li>Python</li></ul>"`.

```js
assert.equal(parseUnorderedList("-  JavaScript\n-  Python"), "<ul><li>JavaScript</li><li>Python</li></ul>");
```

`parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla")` treba da vrati `"<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>"`.

```js
assert.equal(parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla"), "<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>");
```

`parseUnorderedList("- A-1\n- A-2\n- B-1")` treba da vrati `"<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>"`.

```js
assert.equal(parseUnorderedList("- A-1\n- A-2\n- B-1"), "<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>");
```

# --seed--

## --seed-contents--

```js
function parseUnorderedList(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseUnorderedList(markdown) {
  const lines = markdown.split("\n");
  const listItems = [];
  
  for (let line of lines) {
    listItems.push(`<li>${line.slice(1).trim()}</li>`);
  }
  
  return `<ul>${listItems.join("")}</ul>`;
}
```