id: 69373793f5a867f769cde139
title: "Vežba 154: Izračunavanje par-a"
challengeType: 28
dashedName: challenge-154
---

# --description--

Uzimajući u obzir dva celobrojna broja, par za golf rupa i broj udaraca koje je igrač napravio na toj rupi, vratite ocenu igrača koristeći golf terminologiju.

Vratiti:

- `"Hole in one!"` ako je uzastopno napravljen jedan udarac.
- `"Eagle"` ako je napravljeno dva udarca ispod para.
- `"Birdie"` ako je napravljen jedan udarac ispod para.
- `"Par"` ako je napravljen isti broj udaraca kao par.
- `"Bogey"` ako je napravljen jedan udarac više od para.
- `"Double bogey"` ako je napravljeno dva udarca više od para.

# --hints--

`golfScore(3, 3)` treba da vrati `"Par"`.

```js
assert.equal(golfScore(3, 3), "Par");
```

`golfScore(4, 3)` treba da vrati `"Birdie"`.

```js
assert.equal(golfScore(4, 3), "Birdie");
```

`golfScore(3, 1)` treba da vrati `"Hole in one!"`.

```js
assert.equal(golfScore(3, 1), "Hole in one!");
```

`golfScore(5, 7)` treba da vrati `"Double bogey"`.

```js
assert.equal(golfScore(5, 7), "Double bogey");
```

`golfScore(4, 5)` treba da vrati `"Bogey"`.

```js
assert.equal(golfScore(4, 5), "Bogey");
```

`golfScore(5, 3)` treba da vrati `"Eagle"`.

```js
assert.equal(golfScore(5, 3), "Eagle");
```

# --seed--

## --seed-contents--

```js
function golfScore(par, strokes) {

  return par;
}
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) return "Hole in one!";

  const diff = strokes - par;

  if (diff === -2) return "Eagle";
  if (diff === -1) return "Birdie";
  if (diff === 0) return "Par";
  if (diff === 1) return "Bogey";
  if (diff === 2) return "Double bogey";
}
```