id: 69373793f5a867f769cde13a
title: "Vežba 155: Sadnja usjeva"
challengeType: 28
dashedName: challenge-155
---

# --description--

Uzimajući u obzir pun broj koji predstavlja veličinu vašeg polja, a `"acres"` ili `"hectares"` predstavljaju mjernu jedinicu veličine vašeg polja, zajedno sa vrstom usjeva, odredite koliko sadnica te vrste možete zasaditi na svom polju.

- 1 Aker je jednak 4046.86 kvadratnih metara.
- 1 Hektaar je jednak 10,000 kvadratnih metara.

Ovo je popis usjeva sa njihovim unosima i površinom koju zauzima jedna biljka:

|Usjev|Površina po biljci|
|-|-|
|`"corn"`|jedan kvadratni metar|
|`"wheat"`|0.1 kvadratnog metra|
|`"soybeans"`|0.5 kvadratnog metra|
|`"tomatoes"`|0.25 kvadratnog metra|
|`"lettuce"`|0.2 kvadratnog metra|

Vratite broj biljaka koje mogu stati na polje, zaokruženo prema najbližoj cijeloj biljci.

# --hints--

`getNumberOfPlants(1, "acres", "corn")` bi trebalo da vrati `4046`.

```js
assert.equal(getNumberOfPlants(1, "acres", "corn"), 4046);
```

`getNumberOfPlants(2, "hectares", "lettuce")` bi trebalo da vrati `100000`.

```js
assert.equal(getNumberOfPlants(2, "hectares", "lettuce"), 100000);
```

`getNumberOfPlants(20, "acres", "soybeans")` bi trebalo da vrati `161874`.

```js
assert.equal(getNumberOfPlants(20, "acres", "soybeans"), 161874);
```

`getNumberOfPlants(3.75, "hectares", "tomatoes")` bi trebalo da vrati `150000`.

```js
assert.equal(getNumberOfPlants(3.75, "hectares", "tomatoes"), 150000);
```

`getNumberOfPlants(16.75, "acres", "tomatoes")` bi trebalo da vrati `271139`.

```js
assert.equal(getNumberOfPlants(16.75, "acres", "tomatoes"), 271139);
```

# --seed--

## --seed-contents--

```js
function getNumberOfPlants(fieldSize, unit, crop) {

  return fieldSize;
}
```

# --solutions--

```js
function getNumberOfPlants(fieldSize, unit, crop) {
  const cropSpace = {
    corn: 1,
    wheat: 0.1,
    soybeans: 0.5,
    tomatoes: 0.25,
    lettuce: 0.2
  };

  let fieldInSqM = unit === "acres" ? fieldSize * 4046.86 : fieldSize * 10000;

  return Math.floor(fieldInSqM / cropSpace[crop.toLowerCase()]);
}
```