---
id: 69373793f5a867f769cde139
title: "Vežba 154: Par do Rupe"
challengeType: 28
dashedName: challenge-154
---

# --description--
Uzimajući u obzir tačno dve brojke, par za rupu gofu i broj crteža koje je igrač gofu uzeo iz te rupe, vratite rezultat igrača koristeći reči gofu.

Rudisha:

- `"Hole in one!"` ikiwa ilichukua mchoro mmoja.
- `"Eagle"` ikiwa ilichukua michoro miwili chini ya par.
- `"Birdie"` ikiwa ilichukua mchoro mmoja chini ya par.
- `"Par"` ikiwa ilichukua idadi sawa ya michoro na par.
- `"Bogey"` ikiwa ilichukua mchoro mmoja zaidi ya par.
- `"Double bogey"` ikiwa ilichukua michoro miwili zaidi ya par.

# --hints--

`golfScore(3, 3)` inapaswa kurudisha `"Par"`.

```js
assert.equal(golfScore(3, 3), "Par");
```

`golfScore(4, 3)` inapaswa kurudisha `"Birdie"`.

```js
assert.equal(golfScore(4, 3), "Birdie");
```

`golfScore(3, 1)` inapaswa kurudisha `"Hole in one!"`.

```js
assert.equal(golfScore(3, 1), "Hole in one!");
```

`golfScore(5, 7)` inapaswa kurudisha `"Double bogey"`.

```js
assert.equal(golfScore(5, 7), "Double bogey");
```

`golfScore(4, 5)` inapaswa kurudisha `"Bogey"`.

```js
assert.equal(golfScore(4, 5), "Bogey");
```

`golfScore(5, 3)` inapaswa kurudisha `"Eagle"`.

```js
assert.equal(golfScore(5, 3), "Eagle");
```

# --seed--

## --seed-contents--

```js
function golfScore(par, strokes) {

  return par;
}
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) return "Hole in one!";

  const diff = strokes - par;

  if (diff === -2) return "Eagle";
  if (diff === -1) return "Birdie";
  if (diff === 0) return "Par";
  if (diff === 1) return "Bogey";
  if (diff === 2) return "Double bogey";
}
```
