id: 6939b873185d8e00d453563d
title: "Vežba 158: Razmena elemenata niza"
challengeType: 28
dashedName: challenge-158
---

# --description--

Ako vam je dat niz podataka sa dve vrednosti, vratite niz sa tim vrednostima razmenjenim.

Primer, ako vam je dat `["A", "B"]`, vratite `["B", "A"]`.

# --hints--

`arraySwap(["A", "B"])` bi trebalo da vrati `["B", "A"]`.

```js
assert.deepEqual(arraySwap(["A", "B"]), ["B", "A"]);
```

`arraySwap([25, 20])` bi trebalo da vrati `[20, 25]`.

```js
assert.deepEqual(arraySwap([25, 20]), [20, 25]);
```

`arraySwap([true, false])` bi trebalo da vrati `[false, true]`.

```js
assert.deepEqual(arraySwap([true, false]), [false, true]);
```

`arraySwap(["1", 1])` bi trebalo da vrati `[1, "1"]`.

```js
assert.deepEqual(arraySwap(["1", 1]), [1, "1"]);
```

# --seed--

## --seed-contents--

```js
function arraySwap(arr) {

  return arr;
}
```

# --solutions--

```js
function arraySwap(arr) {

  return [arr[1], arr[0]];
}
```