id: 6939b873185d8e00d453563b
title: "Zadatak 156: Paarno ili neparodno?"
challengeType: 28
dashedName: challenge-156
---

# --description--

Ako vam je dat pozitivan ceo broj, vratite `"Odd"` ako je neparizan, i `"Even"` ako je parizan.

# --hints--

`oddOrEven(1)` bi trebalo da vrati `"Odd"`.

```js
assert.equal(oddOrEven(1), "Odd");
```

`oddOrEven(2)` bi trebalo da vrati `"Even"`.

```js
assert.equal(oddOrEven(2), "Even");
```

`oddOrEven(13)` bi trebalo da vrati `"Odd"`.

```js
assert.equal(oddOrEven(13), "Odd");
```

`oddOrEven(196)` bi trebalo da vrati `"Even"`.

```js
assert.equal(oddOrEven(196), "Even");
```

`oddOrEven(123456789)` bi trebalo da vrati `"Odd"`.

```js
assert.equal(oddOrEven(123456789), "Odd");
```

# --seed--

## --seed-contents--

```js
function oddOrEven(n) {

  return n;
}
```

# --solutions--

```js
function oddOrEven(n) {
  return n % 2 === 0 ? "Even" : "Odd";
}
```