id: 6939b873185d8e00d453563e
title: "Vežba 159: Hipotenuza celog broja"
challengeType: 28
dashedName: challenge-159
---

# --description--

Kada razmatrate dva pozitivna cela broja koja predstavljaju dužinu dve katete (dve kraće stranice) pravouglog trougla, utvrdite da li je hipotenuza cel broj.

Dužina hipotenuze se računa sabiranjem kvadrata dužina dve katete i zatim uzimanjem kvadratnog korena tog zbirnog rezultata ($a^2 + b^2 = c^2$).

# --hints--

`isIntegerHypotenuse(3, 4)` bi trebalo da vrati `true`.

```js
assert.isTrue(isIntegerHypotenuse(3, 4));
```

`isIntegerHypotenuse(2, 3)` bi trebalo da vrati `false`.

```js
assert.isFalse(isIntegerHypotenuse(2, 3));
```

`isIntegerHypotenuse(5, 12)` bi trebalo da vrati `true`.

```js
assert.isTrue(isIntegerHypotenuse(5, 12));
```

`isIntegerHypotenuse(10, 10)` bi trebalo da vrati `false`.

```js
assert.isFalse(isIntegerHypotenuse(10, 10));
```

`isIntegerHypotenuse(780, 1040)` bi trebalo da vrati `true`.

```js
assert.isTrue(isIntegerHypotenuse(780, 1040));
```

`isIntegerHypotenuse(250, 333)` bi trebalo da vrati `false`.

```js
assert.isFalse(isIntegerHypotenuse(250, 333));
```

# --seed--

## --seed-contents--

```js
function isIntegerHypotenuse(a, b) {

  return a;
}
```

# --solutions--

```js
function isIntegerHypotenuse(a, b) {
  const sum = a * a + b * b;
  const c = Math.floor(Math.sqrt(sum));
  return c * c === sum;
}
```