id: 6939b873185d8e00d453563c
title: "Vežba 157: Analiziranje Markdown linka"
challengeType: 28
dashedName: challenge-157
---

# --description--

Izvučite niz znakova linka u Markdownu i vratite ekvivalentan HTML niz znakova.

Markdown slika ima sledeću strukturu: `"[link_text](link_url)"`. Vratite string HTML oznake `a` sa `href` podešenim na `link_url` i `link_text` kao sadržajem te oznake.

Na primer, ako dobijete `"[freeCodeCamp](https://freecodecamp.org/)"`, vratite `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`;

Napomena: Konsola možda neće prikazati HTML oznake u nizu znakova prilikom snimanja poruke — proverite konzolu pretraživača da vidite snimke sa priloženim oznakama.

# --hints--

`parseLink("[freeCodeCamp](https://freecodecamp.org/)")` bi trebalo da vrati `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`.

```js
assert.equal(parseLink("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>');
```

`parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` bi trebalo da vrati `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`.

```js
assert.equal(parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>');
```

`parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` bi trebalo da vrati `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'`.

```js
assert.equal(parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>');
```

# --seed--

## --seed-contents--

```js
function parseLink(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseLink(markdown) {
  const match = markdown.match(/\[(.*?)\]\((.*?)\)/);

  const [, text, url] = match;
  return `<a href="${url}">${text}</a>`;
}
```