id: 694596b0585c11170ac7c7fa
title: "Vežba 162: Potrošnja energije"
challengeType: 28
dashedName: challenge-162
---

# --description--

Na osnovu broja kalorija sprosta tokom vežbanja i broja vati-sati potrošenih od vašeg elektronskog uređaja tokom tog vežbanja, utvrdite šta je potrošilo više energije.

Da biste uporedili, pretvorite obe vrednosti u džoule koristeći sledeće konverzije:

- 1 Kalorija je jednaka 4184 džoula.
- 1 vati-sat je jednak 3600 džuola.

Vratite:

- `"Workout"` ako su vežbanje potrošile više energije.
- `"Devices"` ako je uređaj potrošio više energije.
- `"Equal"` ako su sve potrošile jednak iznos energije.

# --hints--

`compareEnergy(250, 50)` should return `"Workout"`.

```js
assert.equal(compareEnergy(250, 50), "Workout");
```

`compareEnergy(100, 200)` should return `"Devices"`.

```js
assert.equal(compareEnergy(100, 200), "Devices");
```

`compareEnergy(450, 523)` should return `"Equal"`.

```js
assert.equal(compareEnergy(450, 523), "Equal");
```

`compareEnergy(300, 75)` should return `"Workout"`.

```js
assert.equal(compareEnergy(300, 75), "Workout");
```

`compareEnergy(200, 250)` should return `"Devices"`.

```js
assert.equal(compareEnergy(200, 250), "Devices");
```

`compareEnergy(900, 1046)` should return `"Equal"`.

```js
assert.equal(compareEnergy(900, 1046), "Equal");
```

# --seed--

## --seed-contents--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {

  return caloriesBurned;
}
```

# --solutions--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {
  const workoutEnergy = caloriesBurned * 4184;
  const deviceEnergy = wattHoursUsed * 3600;

  console.log(workoutEnergy);
  console.log(deviceEnergy);
  if (workoutEnergy > deviceEnergy) return "Workout";
  if (deviceEnergy > workoutEnergy) return "Devices";
  return "Equal";
}
```