id: 694596b0585c11170ac7c7fc
title: "Vežba 164: Parsiranje koda unutar linije u Markdownu"
challengeType: 28
dashedName: challenge-164
---

# --description--

Ako vam je dat niz znakova u formatu Markdown koji sadrži jedan ili više inline kodnih dijelova, vratite odgovarajući niz HTML znakova.

Inline kodni dijelovi u Markdownu koriste jedan backtick (``` ` ```) na početku i kraju teksta kodnog dijela.

Vratite zadati niz sa svim kodnim dijelovima zamijenjenim HTML tagom `<code/>`.

Na primjer, ako primite niz ```"Use `let` to declare the variable."```, vratite `"Koristite <code>let</code> za deklarisanje varijable."`.

Napomena: Konzola možda neće prikazati HTML tagove unutar niza prilikom snimanja poruke. Provjerite konzolu pretraživača da vidite snimke sa ugrađenim tagovima.

# --hints--

```parseInlineCode("Use `let` to declare the variable.")``` should return `"Use <code>let</code> to declare the variable."`.

```js
assert.equal(parseInlineCode("Tumia `let` kutangaza kigezo."), "Tumia <code>let</code> kutangaza kigezo.");
```

```parseInlineCode("Koristite `let` ili `const` za deklarisanje varijable.")``` should return `"Koristite <code>let</code> ili <code>const</code> za deklarisanje varijable."`.

```js
assert.equal(parseInlineCode("Use `let` or `const` to declare a variable."), "Use <code>let</code> or <code>const</code> to declare a variable.");
```

```parseInlineCode("Run `npm install` then `npm start`.")``` should return `"Run <code>npm install</code> then <code>npm start</code>."`.

```js
assert.equal(parseInlineCode("Pokrenite `npm install`, zatim `npm start`."), "Pokrenite <code>npm install</code>, zatim <code>npm start</code>.");
```

# --seed--

## --seed-contents--

```js
function parseInlineCode(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseInlineCode(markdown) {
  return markdown.replace(/`([^`]+)`/g, "<code>$1</code>");
}
```