id: 694596b0585c11170ac7c7fb
title: "Vežba 163: Mala slova za suglasnike"
challengeType: 28
dashedName: challenge-163
---

# --description--

Kada dobijete niz znakova koji predstavlja ime varijable, konvertujte ga u format sa malim slovima suglasnika koristeći sledeća pravila:

- Sve suglasnike pretvorite u VELIKA slova.
- Sve samoglasnike (`a`, `e`, `i`, `o`, `u` u bilo kom slučaju) pretvorite u mala slova.
- Razdvojnici riječi (najkraći crtici) (`-`) zamijenite sa podvučnicama (`_`).

# --hints--

`toConsonantCase("helloworld")` bi trebalo da vrati `"HeLLoWoRLD"`.

```js
assert.equal(toConsonantCase("helloworld"), "HeLLoWoRLD");
```

`toConsonantCase("HELLOWORLD")` bi trebalo da vrati `"HeLLoWoRLD"`.

```js
assert.equal(toConsonantCase("HELLOWORLD"), "HeLLoWoRLD");
```

`toConsonantCase("_hElLO-WOrlD-_")` bi trebalo da vrati `"_HeLLo_WoRLD_"`.

```js
assert.equal(toConsonantCase("_hElLO-WOrlD-_"), "_HeLLo_WoRLD_");
```

`toConsonantCase("_~-generic_~-variable_~-name_~-here-~_")` bi trebalo da vrati `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`.

```js
assert.equal(toConsonantCase("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_");
```

# --seed--

## --seed-contents--

```js
function toConsonantCase(str) {

  return str;
}
```

# --solutions--

```js
function toConsonantCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (char === "-") {
      result += "_";
    } else if (/[a-zA-Z]/.test(char)) {
      if (vowels.includes(char)) {
        result += char.toLowerCase();
      } else {
        result += char.toUpperCase();
      }
    } else {
      result += char;
    }
  }

  return result;
}
```