id: 696655d24b614176d4c9b78c
title: "Vežba 169: Mini FizzBuzz"
challengeType: 28
dashedName: challenge-169
---

# --description--

Kada dobijete celo число, vratite niz slova prema sledećim pravilima:

- Ako je broj deljiv sa 3, vratite `"Fizz"`.
- Ako je broj deljiv sa 5, vratite `"Buzz"`.
- Ako je broj deljiv i sa 3 i sa 5, vratite `"FizzBuzz"`.
- Inače, vratite taj broj kao niz slova.

# --hints--

`fizzBuzzMini(3)` treba da vrati `"Fizz"`.

```js
assert.equal(fizzBuzzMini(3), "Fizz");
```

`fizzBuzzMini(4)` treba da vrati `"4"`.

```js
assert.strictEqual(fizzBuzzMini(4), "4");
```

`fizzBuzzMini(35)` treba da vrati `"Buzz"`.

```js
assert.equal(fizzBuzzMini(35), "Buzz");
```

`fizzBuzzMini(75)` treba da vrati `"FizzBuzz"`.

```js
assert.equal(fizzBuzzMini(75), "FizzBuzz");
```

`fizzBuzzMini(98)` treba da vrati `"98"`.

```js
assert.strictEqual(fizzBuzzMini(98), "98");
```

# --seed--

## --seed-contents--

```js
function fizzBuzzMini(n) {

  return n;
}
```

# --solutions--

```js
function fizzBuzzMini(n) {
  if (n % 3 === 0 && n % 5 === 0) return "FizzBuzz";
  if (n % 3 === 0) return "Fizz";
  if (n % 5 === 0) return "Buzz";
  return String(n);
}
```