id: 696655d24b614176d4c9b78d
title: "Vežba 170: Čudan ili paran dan"
challengeType: 28
dashedName: challenge-170
---

# --description--

Kada dobijete vremenski pečat (broj milisekundi od početka Unix epoha), vratite:

- `"odd"` ako je broj dana meseca za taj vremenski pečat neparan.
- `"even"` ako je broj dana meseca za taj vremenski pečat paran.

Na primer, ukoliko dobijete `1769472000000`, za vremenski pečat 27. januara 2026. godine vratite `"odd"` jer je broj dana (`27`) neparan broj.

Napomena: Vremenski pečat je u milisekundama i morate koristiti datum po UTC vremenu, a ne lokalno vreme vaše zone.

# --hints--

`oddOrEvenDay(1769472000000)` treba da vrati `"odd"`.

```js
assert.equal(oddOrEvenDay(1769472000000), "odd");
```

`oddOrEvenDay(1769444440000)` treba da vrati `"even"`.

```js
assert.equal(oddOrEvenDay(1769444440000), "even");
```

`oddOrEvenDay(6739456780000)` treba da vrati `"odd"`.

```js
assert.equal(oddOrEvenDay(6739456780000), "odd");
```

`oddOrEvenDay(1)` treba da vrati `"odd"`.

```js
assert.equal(oddOrEvenDay(1), "odd");
```

`oddOrEvenDay(86400000)` treba da vrati `"even"`.

```js
assert.equal(oddOrEvenDay(86400000), "even");
```

# --seed--

## --seed-contents--

```js
function oddOrEvenDay(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function oddOrEvenDay(timestamp) {
  const date = new Date(timestamp);
  const day = date.getUTCDate();

  return day % 2 === 0 ? "even" : "odd";
}
```