id: 696655d24b614176d4c9b789
title: "Vežba 166: Provera heksadecimalnog niza karaktera"
challengeType: 28
dashedName: challenge-166
---

# --description--

Kada dobijete niz karaktera, proverite da li je validna CSS heksadecimalna boja. Validna CSS heksadecimalna boja mora:

- Početi sa `#`, i
- biti praćena 3 ili 6 heksadecimalnih karaktera.

Heksadecimalni karakteri su brojevi od `0` do `9` i slova `a` do `f` (nema razlike između velikih i malih slova).

# --hints--

`isValidHex("#123")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidHex("#123"));
```

`isValidHex("#123abc")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidHex("#123abc"));
```

`isValidHex("#ABCDEF")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidHex("#ABCDEF"));
```

`isValidHex("#0a1B2c")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidHex("#0a1B2c"));
```

`isValidHex("#12G")` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidHex("#12G"));
```

`isValidHex("#1234567")` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidHex("#1234567"));
```

`isValidHex("#12 3")` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidHex("#12 3"));
```

`isValidHex("fff")` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidHex("fff"));
```

# --seed--

## --seed-contents--

```js
function isValidHex(str) {

  return str;
}
```

# --solutions--

```js
function isValidHex(str) {
  if (str[0] !== "#") return false;

  const hex = str.slice(1);
  if (hex.length !== 3 && hex.length !== 6) return false;

  for (let char of hex) {
    if (!/[0-9a-fA-F]/.test(char)) {
      return false;
    }
  }

  return true;
}
```