id: 696655d24b614176d4c9b78b
title: "Vežba 168: Ponovno skalirana slika"
challengeType: 28
dashedName: challenge-168
---

# --description--

Dato je niz znakova koji predstavlja širinu i visinu slike, te faktor skaliranja za tu sliku, vratite skalirane širinu i visinu.

- Ulazni niz je u formatu `"ŠxV"`. Na primer, `"800x600"`.
- Skala je broj sa kojim treba pomnožiti širinu i visinu.

Vratite skalovane dimenzije u istom formatu `"ŠxV"`.

# --hints--

`scaleImage("800x600", 2)` bi trebalo da vrati `"1600x1200"`.

```js
assert.equal(scaleImage("800x600", 2), "1600x1200");
```

`scaleImage("100x100", 10)` bi trebalo da vrati `"1000x1000"`.

```js
assert.equal(scaleImage("100x100", 10), "1000x1000");
```

`scaleImage("1024x768", 0.5)` bi trebalo da vrati `"512x384"`.

```js
assert.equal(scaleImage("1024x768", 0.5), "512x384");
```

`scaleImage("300x200", 1.5)` bi trebalo da vrati `"450x300"`.

```js
assert.equal(scaleImage("300x200", 1.5), "450x300");
```

# --seed--

## --seed-contents--

```js
function scaleImage(size, scale) {

  return size;
}
```

# --solutions--

```js
function scaleImage(size, scale) {
  const [width, height] = size.split("x").map(Number);

  const newWidth = width * scale;
  const newHeight = height * scale;

  return `${newWidth}x${newHeight}`;
}
```