id: 69738771fb5a7b8b24cca2a2
title: "Vežba 176: Dan groundhoga"
challengeType: 28
dashedName: challenge-176
---

# --description--

Dan je dan groundhoga, kada groundhog predviđa vremensku prognozu na osnovu toga da li vidi svoju senku ili ne.

Kada dobijete vrednost koja predstavlja pojavu groundhoga, vratite tačnu predviđanju:

- Ako je zadata vrednost logički `true` (groundhog je video svoju senku), vratite `"Looks like we'll have six more weeks of winter."`.
- Ako je vrednost logički `false` (groundhog nije video svoju senku), vratite `"It's going to be an early spring."`.
- Ako je vrednost bilo šta drugo (groundhog nije viđen), vratite `"No prediction this year."`.

# --hints--

Funkcija `groundhogDayPrediction(true)` bi trebalo da vrati `"Looks like we'll have six more weeks of winter."`.

```js
assert.equal(groundhogDayPrediction(true), "Looks like we'll have six more weeks of winter.");
```

Funkcija `groundhogDayPrediction(false)` bi trebalo da vrati `"It's going to be an early spring."`.

```js
assert.equal(groundhogDayPrediction(false), "It's going to be an early spring.");
```

Funkcija `groundhogDayPrediction(null)` bi trebalo da vrati `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction(null), "No prediction this year.");
```

Funkcija `groundhogDayPrediction(" ")` bi trebalo da vrati `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction(" "), "No prediction this year.");
```

Funkcija `groundhogDayPrediction("true")` bi trebalo da vrati `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction("true"), "No prediction this year.");
```

# --seed--

## --seed-contents--

```js
function groundhogDayPrediction(appearance) {

  return appearance;
}
```

# --solutions--

```js
function groundhogDayPrediction(appearance) {
  if (appearance === true) return "Looks like we'll have six more weeks of winter.";
  if (appearance === false) return "It's going to be an early spring.";
  return "No prediction this year.";
}
```