id: 69738771fb5a7b8b24cca29f
title: "Zadatak 173: Legalni potez piona"
challengeType: 28
dashedName: challenge-173
---

# --description--

Iz pozicije jednog vašeg piona na šahovskoj tabli, vratite niz svih legalnih poteza koje se mogu napraviti u rastućem redosledu.

Standardna šahovska tabla je 8x8, a redovi su označeni od `A` do `H` (sa leve na desno) i kolone su označene od `1` do `8` (od dna ka vrhu). Izgleda ovako:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

Za ovaj zadatak:

- Vi ste igrač sa donje strane tablice.
- Pioni mogu se pomerati samo direktno "gore".
- Osim ako pion nije na početnom redu (red 2), može se pomeriti jedan ili dva koraka napred.

Na primer, ako dobijete `"D4"`, vratite `["D5"]`, što je jedina pozicija u koju vaš pion može da se pomakne.
Ako dobijete `"B2"`, vratite `["B3", "B4"]`, jer se nalazi na početnom redu i mora biti u rastućem redosledu.

# --hints--

`findPawnMoves("D4")` bi trebalo da vrati `["D5"]`.

```js
assert.deepEqual(findPawnMoves("D4"), ["D5"]);
```

`findPawnMoves("B2")` bi trebalo da vrati `["B3", "B4"]`.

```js
assert.deepEqual(findPawnMoves("B2"), ["B3", "B4"]);
```

`findPawnMoves("A7")` bi trebalo da vrati `["A8"]`.

```js
assert.deepEqual(findPawnMoves("A7"), ["A8"]);
```

`findPawnMoves("G2")` bi trebalo da vrati `["G3", "G4"]`.

```js
assert.deepEqual(findPawnMoves("G2"), ["G3", "G4"]);
```

`findPawnMoves("E3")` bi trebalo da vrati `["E4"]`.

```js
assert.deepEqual(findPawnMoves("E3"), ["E4"]);
```

# --seed--

## --seed-contents--

```js
function findPawnMoves(position) {

  return position;
}

```

# --solutions--

```js
function findPawnMoves(position) {
  const column = position[0].toUpperCase();
  const row = parseInt(position[1]);
  const moves = [];

  if (row >= 8) return moves;

  moves.push(`${column}${row + 1}`);

  if (row === 2) moves.push(`${column}${row + 2}`);

  return moves;
}
```