id: 69738771fb5a7b8b24cca2a3
title: "Vežba 177: Ogledalo slovačnog niza"
challengeType: 28
dashedName: challenge-177
---

# --description--

Dato vam je slovočni niz, vratite novi slovočni niz koji uključuje dati niz praćen kopijom tog niza okrenom na kraj niza.

# --hints--

`mirror("freeCodeCamp")` treba da vrati `"freeCodeCamppmaCedoCeerf"`.

```js
assert.equal(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf");
```

`mirror("RaceCar")` treba da vrati `"RaceCarraCecaR"`.

```js
assert.equal(mirror("RaceCar"), "RaceCarraCecaR");
```

`mirror("helloworld")` treba da vrati `"helloworlddlrowolleh"`.

```js
assert.equal(mirror("helloworld"), "helloworlddlrowolleh");
```

`mirror("The quick brown fox...")` treba da vrati `"The quick brown fox......xof nworb kciuq ehT"`.

```js
assert.equal(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT");
```

# --seed--

## --seed-contents--

```js
function mirror(str) {

  return str;
}
```

# --solutions--

```js
function mirror(str) {
  const reversed = str.split("").reverse().join("");
  return str + reversed;
}
```