id: 69738771fb5a7b8b24cca2a4
title: "Zadatak 178: Skraćivanje niza znakova"
challengeType: 28
dashedName: challenge-178
---

# --description--

Ako vam je dat niz znakova, vratite ga takav kakav jeste ako ima 20 ili manje znakova. Ako ima više od 20 znakova, skratite na prvih 17 znakova, a zatim dodajte `"..."` na kraj (da bi ukupan broj bio 20) i vratite rezultat.

# --hints--

`truncateText("Hello, world!")` treba da vrati `"Hello, world!"`.

```js
assert.equal(truncateText("Hello, world!"), "Hello, world!");
```

`truncateText("This string should get truncated.")` treba da vrati `"This string shoul..."`.

```js
assert.equal(truncateText("This string should get truncated."), "This string shoul...");
```

`truncateText("Exactly twenty chars")` treba da vrati `"Exactly twenty chars"`.

```js
assert.equal(truncateText("Exactly twenty chars"), "Exactly twenty chars");
```

`truncateText(".....................")` treba da vrati `"...................."`

```js
assert.equal(truncateText("....................."), "....................");
```

# --seed--

## --seed-contents--

```js
function truncateText(text) {

  return text;
}
```

# --solutions--

```js
function truncateText(text) {
  if (text.length <= 20) return text;
  return text.slice(0, 17) + "...";
}
```