id: 69738771fb5a7b8b24cca2a0
title: "Vežba 174: Pronalazač Zodijaka"
challengeType: 28
dashedName: challenge-174
---

# --description--

Dato je niz znakova dat u formatu `"YYYY-MM-DD"`; vratite zodijak za taj datum koristeći sledeću tabelu:

| Period datuma | Zodijak |
| - | - |
| 21. mart - 19. april | `"Aries"` |
| 20. april - 20. maj | `"Taurus"` |
| 21. maj - 20. jun | `"Gemini"` |
| 21. jun - 22. jul | `"Cancer"` |
| 23. jul - 22. avgust | `"Leo"` |
| 23. avgust - 22. septembar | `"Virgo"` |
| 23. septembar - 22. oktobar | `"Libra"` |
| 23. oktobar - 21. novembar | `"Scorpio"` |
| 22. novembar - 21. decembar | `"Sagittarius"` |
| 22. decembar - 19. januar | `"Capricorn"` |
| 20. januar - 18. februar | `"Aquarius"` |
| 19. februar - 20. mart | `"Pisces"` |

- Zodijak zavisi samo od meseca i dana, pa možete zanemariti godinu.

# --hints--

`getSign("2026-01-31")` bi trebalo da vrati `"Aquarius"`.

```js
assert.equal(getSign("2026-01-31"), "Aquarius");
```

`getSign("2001-06-10")` bi trebalo da vrati `"Gemini"`.

```js
assert.equal(getSign("2001-06-10"), "Gemini");
```

`getSign("1985-09-07")` bi trebalo da vrati `"Virgo"`.

```js
assert.equal(getSign("1985-09-07"), "Virgo");
```

`getSign("2023-03-19")` bi trebalo da vrati `"Pisces"`.

```js
assert.equal(getSign("2023-03-19"), "Pisces");
```

`getSign("2045-11-05")` bi trebalo da vrati `"Scorpio"`.

```js
assert.equal(getSign("2045-11-05"), "Scorpio");
```

`getSign("1985-12-06")` bi trebalo da vrati `"Sagittarius"`.

```js
assert.equal(getSign("1985-12-06"), "Sagittarius");
```

`getSign("2025-12-30")` bi trebalo da vrati `"Capricorn"`.

```js
assert.equal(getSign("2025-12-30"), "Capricorn");
```

`getSign("2018-10-08")` bi trebalo da vrati `"Libra"`.

```js
assert.equal(getSign("2018-10-08"), "Libra");
```

`getSign("1958-05-04")` bi trebalo da vrati `"Taurus"`.

```js
assert.equal(getSign("1958-05-04"), "Taurus");
```

# --seed--

## --seed-contents--

```js
function getSign(dateStr) {

  return dateStr;
}
```

# --solutions--

```js
function getSign(dateStr) {
  const [, monthStr, dayStr] = dateStr.split("-");
  const month = Number(monthStr);
  const day = Number(dayStr);

  if ((month === 3 && day >= 21) || (month === 4 && day <= 19)) return "Aries";
  if ((month === 4 && day >= 20) || (month === 5 && day <= 20)) return "Taurus";
  if ((month === 5 && day >= 21) || (month === 6 && day <= 20)) return "Gemini";
  if ((month === 6 && day >= 21) || (month === 7 && day <= 22)) return "Cancer";
  if ((month === 7 && day >= 23) || (month === 8 && day <= 22)) return "Leo";
  if ((month === 8 && day >= 23) || (month === 9 && day <= 22)) return "Virgo";
  if ((month === 9 && day >= 23) || (month === 10 && day <= 22)) return "Libra";
  if ((month === 10 && day >= 23) || (month === 11 && day <= 21)) return "Scorpio";
  if ((month === 11 && day >= 22) || (month === 12 && day <= 21)) return "Sagittarius";
  if ((month === 12 && day >= 22) || (month === 1 && day <= 19)) return "Capricorn";
  if ((month === 1 && day >= 20) || (month === 2 && day <= 18)) return "Aquarius";
  if ((month === 2 && day >= 19) || (month === 3 && day <= 20)) return "Pisces";
}
```