id: 69738771fb5a7b8b24cca2a5
title: "Vežba 179: Novac u džepu"
challengeType: 28
dashedName: challenge-179
---

# --description--

Dato je niz celih brojeva koji predstavljaju novčiće u vašem džepu, gde svaki ceo broj predstavlja vrednost novčića u centima. Vratite ukupan iznos u formatu `"$D.CC"`.

- 100 centi je jednako $1.
- U povratnoj vrednosti, uključite vodeće nule za iznos manji od dolara i uvek samo dve cifre za centi.

# --hints--

`countChange([25, 10, 5, 1])` treba da vrati `"$0.41"`.

```js
assert.equal(countChange([25, 10, 5, 1]), "$0.41");
```

`countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25])` treba da vrati `"$1.43"`.

```js
assert.equal(countChange([25, 10, 5, 1, 25, 10, 25, 1, 1, 10, 5, 25]), "$1.43");
```

`countChange([100, 25, 100, 1000, 5, 500, 2000, 25])` treba da vrati `"$37.55"`.

```js
assert.equal(countChange([100, 25, 100, 1000, 5, 500, 2000, 25]), "$37.55");
```

`countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10])` treba da vrati `"$0.70"`.

```js
assert.equal(countChange([10, 5, 1, 10, 1, 25, 1, 1, 5, 1, 10]), "$0.70");
```

`countChange([1])` treba da vrati `"$0.01"`.

```js
assert.equal(countChange([1]), "$0.01");
```

`countChange([25, 25, 25, 25])` treba da vrati `"$1.00"`.

```js
assert.equal(countChange([25, 25, 25, 25]), "$1.00");
```

# --seed--

## --seed-contents--

```js
function countChange(change) {

  return change;
}
```

# --solutions--

```js
function countChange(change) {
  const totalCents = change.reduce((sum, coin) => sum + coin, 0);

  const dollars = Math.floor(totalCents / 100);
  const cents = totalCents % 100;

  return `$${dollars}.${cents.toString().padStart(2, "0")}`;
}
```