id: 697a49e6ff50d756c9b6935e
title: "Vežba 181: Zimske igre 2026. Dan 2: Snoboarding"
challengeType: 28
dashedName: challenge-181
---

# --description--

Kada snoboarding sportista počinje sa rotacijom u stepenima, odredite kakva će mu biti konačna poza.

- "Poza snoboardista" je `"Regular"` ili `"Goofy"`.
- Rotacije manevra su često od 90 stepeni. Pozitivan pokazuje rotaciju po kazaljci na satu, a negativan pokazuje rotaciju suprotno kazaljci na satu.
- Poza prizemljavanja se menja svakih 180 stepeni rotacije.

Na primer, ako je dato `"Regular"` i `90`, vratite `"Regular"`. Ako je dato `"Regular"` i 180 stepeni, vratite `"Goofy"`.

# --hints--

`getLandingStance("Regular", 90)` bi trebalo da vrati `"Regular"`.

```js
assert.equal(getLandingStance("Regular", 90), "Regular");
```

`getLandingStance("Regular", 180)` bi trebalo da vrati `"Goofy"`.

```js
assert.equal(getLandingStance("Regular", 180), "Goofy");
```

`getLandingStance("Goofy", -270)` bi trebalo da vrati `"Regular"`.

```js
assert.equal(getLandingStance("Goofy", -270), "Regular");
```

`getLandingStance("Regular", 2340)` bi trebalo da vrati `"Goofy"`.

```js
assert.equal(getLandingStance("Regular", 2340), "Goofy");
```

`getLandingStance("Goofy", 2160)` bi trebalo da vrati `"Goofy"`.

```js
assert.equal(getLandingStance("Goofy", 2160), "Goofy");
```

`getLandingStance("Goofy", -540)` bi trebalo da vrati `"Regular"`.

```js
assert.equal(getLandingStance("Goofy", -540), "Regular");
```

# --seed--

## --seed-contents--

```js
function getLandingStance(startStance, rotation) {

  return startStance;
}
```

# --solutions--

```js
function getLandingStance(startStance, rotation) {
  const flips = Math.floor(Math.abs(rotation) / 180);
  if (flips % 2 === 0) {
    return startStance;
  } else {
    return startStance === "Regular" ? "Goofy" : "Regular";
  }
}
```