id: 697a49e9860d24853adef67e
title: "Vežba 193: Zimske sportove 2026 Dan 14: Skijanje niz planine"
challengeType: 28
dashedName: challenge-193
---

# --description--

Na osnovu dubine snega i strmine planine, utvrdite da li postoji rizik od snježne krupobitnosti.

- Vrednosti dubine snega su `"Shallow"`, `"Moderate"` ili `"Deep"`.
- Vrednosti strmine su `"Gentle"`, `"Steep"` ili `"Very Steep"`.

Vratite `"Safe"` (Sigurno) ili `"Risky"` (Rizično) na osnovu ove tabele:

| |`"Shallow"`|`"Moderate"`|`"Deep"`|
|-|-|-|-|
|`"Gentle"`|`"Safe"`|`"Safe"`|`"Safe"`|
|`"Steep"`|`"Safe"`|`"Risky"`|`"Risky"`|
|`"Very Steep"`|`"Safe"`| `"Risky"`|`"Risky"`|

# --hints--

`avalancheRisk("Shallow", "Gentle")` treba da vrati `"Safe"`.

```js
assert.equal(avalancheRisk("Shallow", "Gentle"), "Safe");
```

`avalancheRisk("Shallow", "Steep")` treba da vrati `"Safe"`.

```js
assert.equal(avalancheRisk("Shallow", "Steep"), "Safe");
```

`avalancheRisk("Shallow", "Very Steep")` treba da vrati `"Safe"`.

```js
assert.equal(avalancheRisk("Shallow", "Very Steep"), "Safe");
```

`avalancheRisk("Moderate", "Gentle")` treba da vrati `"Safe"`.

```js
assert.equal(avalancheRisk("Moderate", "Gentle"), "Safe");
```

`avalancheRisk("Moderate", "Steep")` treba da vrati `"Risky"`.

```js
assert.equal(avalancheRisk("Moderate", "Steep"), "Risky");
```

`avalancheRisk("Moderate", "Very Steep")` treba da vrati `"Risky"`.

```js
assert.equal(avalancheRisk("Moderate", "Very Steep"), "Risky");
```

`avalancheRisk("Deep", "Gentle")` treba da vrati `"Safe"`.

```js
assert.equal(avalancheRisk("Deep", "Gentle"), "Safe");
```

`avalancheRisk("Deep", "Steep")` treba da vrati `"Risky"`.

```js
assert.equal(avalancheRisk("Deep", "Steep"), "Risky");
```

`avalancheRisk("Deep", "Very Steep")` treba da vrati `"Risky"`.

```js
assert.equal(avalancheRisk("Deep", "Very Steep"), "Risky");
```

# --seed--

## --seed-contents--

```js
function avalancheRisk(snowDepth, slope) {

  return snowDepth;
}
```

# --solutions--

```js
function avalancheRisk(snowDepth, slope) {
  const riskTable = {
    Gentle: {
      Shallow: "Safe",
      Moderate: "Safe",
      Deep: "Safe",
    },
    Steep: {
      Shallow: "Safe",
      Moderate: "Risky",
      Deep: "Risky",
    },
    "Very Steep": {
      Shallow: "Safe",
      Moderate: "Risky",
      Deep: "Risky",
    },
  };

  return riskTable[slope][snowDepth];
}
```