id: 697a49e6ff50d756c9b69366
title: "Vežba 189: Zimske igre 2026 Dan 10: Skijanje niz planinu"
challengeType: 28
dashedName: challenge-189
---

# --description--

Padina skijane, horizontalna udaljenost i njen tip određuju nivo težine planine.

Da biste odredili nivo:

- Izračunajte strminu podelivanjem padina sa udaljenošću.
- Zatim izračunajte prilagođenu strminu u zavisnosti od tipa planine:
  - `"Downhill"`: pomnožite strminu sa 1.2
  - `"Slalom"`: pomnožite strminu sa 0.9
  - `"Giant Slalom"`: pomnožite strminu sa 1.0

Vratite:

- `"Green"` ako je prilagođena strmina manja ili jednaka 0.1
- `"Blue"` ako je prilagođena strmina veća od 0.1 i manja ili jednaka 0.25
- `"Black"` ako je prilagođena strmina veća od 0.25

# --hints--

`getHillRating(95, 900, "Slalom")` bi trebalo da vrati `"Green"`.

```js
assert.equal(getHillRating(95, 900, "Slalom"), "Green");
```

`getHillRating(620, 2800, "Downhill")` bi trebalo da vrati `"Black"`.

```js
assert.equal(getHillRating(620, 2800, "Downhill"), "Black");
```

`getHillRating(420, 1680, "Giant Slalom")` bi trebalo da vrati `"Blue"`.

```js
assert.equal(getHillRating(420, 1680, "Giant Slalom"), "Blue");
```

`getHillRating(250, 3000, "Downhill")` bi trebalo da vrati `"Green"`.

```js
assert.equal(getHillRating(250, 3000, "Downhill"), "Green");
```

`getHillRating(110, 900, "Slalom")` bi trebalo da vrati `"Blue"`.

```js
assert.equal(getHillRating(110, 900, "Slalom"), "Blue");
```

`getHillRating(380, 1500, "Giant Slalom")` bi trebalo da vrati `"Black"`.

```js
assert.equal(getHillRating(380, 1500, "Giant Slalom"), "Black");
```

# --seed--

## --seed-contents--

```js
function getHillRating(drop, distance, type) {

  return drop;
}
```

# --solutions--

```js
function getHillRating(drop, distance, type) {
  let steepness = drop / distance;

  if (type === "Downhill") {
    steepness *= 1.2;
  } else if (type === "Slalom") {
    steepness *= 0.9;
  } else if (type === "Giant Slalom") {
    steepness *= 1.0;
  }

  if (steepness <= 0.1) {
    return "Green";
  } else if (steepness <= 0.25) {
    return "Blue";
  } else {
    return "Black";
  }
}
```