id: 697a49e9860d24853adef67d
title: "Zoezi la 192: Igre zimskih Olimpijadi 2026, Dan 13: Nordički kombinovani"
challengeType: 28
dashedName: challenge-192
---

# --description--

Dato je niz popuna skokova rezultata sportista; izračunajte njihovo kašnjenje za start u trci nordijskog kombinovanja.

Sportista sa najvišim rezultatom skoka počinje prvi (kašnjenje od 0 sekundi). Svi ostali sportisti počinju kasnije, bazirano na razlici između njihovih rezultata skokova i najboljeg rezultata.

Da biste izračunali kašnjenje za svakog sportistu, oduzmite rezultat skoka sportiste od apsolutno najboljeg rezultata skoka, a zatim pomnožite rezultat sa 1.5. Zaokružite kašnjenje do najbližeg celog broja.

# --hints--

`calculateStartDelays([120, 110, 125])` bi trebalo da vrati `[8, 23, 0]`.

```js
assert.deepEqual(calculateStartDelays([120, 110, 125]), [8, 23, 0]);
```

`calculateStartDelays([118, 125, 122, 120])` bi trebalo da vrati `[11, 0, 5, 8]`.

```js
assert.deepEqual(calculateStartDelays([118, 125, 122, 120]), [11, 0, 5, 8]);
```

`calculateStartDelays([100, 105, 95, 110, 120, 115, 108])` bi trebalo da vrati `[30, 23, 38, 15, 0, 8, 18]`.

```js
assert.deepEqual(calculateStartDelays([100, 105, 95, 110, 120, 115, 108]), [30, 23, 38, 15, 0, 8, 18]);
```

`calculateStartDelays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124])` bi trebalo da vrati `[3, 11, 6, 18, 21, 15, 8, 26, 0, 12]`.

```js
assert.deepEqual(calculateStartDelays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124]), [3, 11, 6, 18, 21, 15, 8, 26, 0, 12]);
```

# --seed--

## --seed-contents--

```js
function calculateStartDelays(jumpScores) {

  return jumpScores
}
```

# --solutions--

```js
function calculateStartDelays(jumpScores) {
  const bestJump = Math.max(...jumpScores);
  return jumpScores.map(score => Math.ceil((bestJump - score) * 1.5));
}
```