id: 697a49e6ff50d756c9b6935d
title: "VeĹľba 180: Zimske Olimpijske igre 2026 Dan 1: Dana otvaranja"
challengeType: 28
dashedName: challenge-180
---

# --description--

Danas je poÄŤetak Zimskih Olimpijskih igara 2026. SledeÄ‡ih 17 dana donosiÄ‡e vam kod izazove inspirisane tim igrama.

Za prvi izazov, dat Ä‡e vam se kod dve-slovnog koda zemlje i potrebno je da vratite emoji zastave te zemlje.

Koristite ovu listu:

| DrĹľava | Kod | Zastava |
| - | -: | :--: |
| Albanija                    |            `"AL"` | `"đź‡¦đź‡±"` |
| Andora                    |            `"AD"` | `"đź‡¦đź‡©"` |
| Argentina                  |            `"AR"` | `"đź‡¦đź‡·"` |
| Armenija                    |            `"AM"` | `"đź‡¦đź‡˛"` |
| Australija                  |            `"AU"` | `"đź‡¦đź‡ş"` |
| Austrija                    |            `"AT"` | `"đź‡¦đź‡ą"` |
| AzerbejdĹľan                 |            `"AZ"` | `"đź‡¦đź‡ż"` |
| Belgija                    |            `"BE"` | `"đź‡§đź‡Ş"` |
| Benin                      |            `"BJ"` | `"đź‡§đź‡Ż"` |
| Bolivija                    |            `"BO"` | `"đź‡§đź‡´"` |
| Bosna i Hercegovina     |            `"BA"` | `"đź‡§đź‡¦"` |
| Brazil                     |            `"BR"` | `"đź‡§đź‡·"` |
| Bugarska                   |            `"BG"` | `"đź‡§đź‡¬"` |
| Kanada                     |            `"CA"` | `"đź‡¨đź‡¦"` |
| ÄŚile                       |            `"CL"` | `"đź‡¨đź‡±"` |
| Kina                      |            `"CN"` | `"đź‡¨đź‡ł"` |
| Kolumbija                   |            `"CO"` | `"đź‡¨đź‡´"` |
| Hrvatska                    |            `"HR"` | `"đź‡­đź‡·"` |
| Ciprus                     |            `"CY"` | `"đź‡¨đź‡ľ"` |
| ÄŚehska Republika             |            `"CZ"` | `"đź‡¨đź‡ż"` |
| Danska                    |            `"DK"` | `"đź‡©đź‡°"` |
| Ekvador                    |            `"EC"` | `"đź‡Şđź‡¨"` |
| Eritreja                    |            `"ER"` | `"đź‡Şđź‡·"` |
| Estonija                    |            `"EE"` | `"đź‡Şđź‡Ş"` |
| Finska                    |            `"FI"` | `"đź‡«đź‡®"` |
| Francuska                     |            `"FR"` | `"đź‡«đź‡·"` |
| Gruzija                    |            `"GE"` | `"đź‡¬đź‡Ş"` |
| NemaÄŤka                    |            `"DE"` | `"đź‡©đź‡Ş"` |
| Velika Britanija              |            `"GB"` | `"đź‡¬đź‡§"` |
| GrÄŤka                     |            `"GR"` | `"đź‡¬đź‡·"` |
| Gvineja-Bissau              |            `"GW"` | `"đź‡¬đź‡Ľ"` |
| Haiti                      |            `"HT"` | `"đź‡­đź‡ą"` |
| Hong Kong                  |            `"HK"` | `"đź‡­đź‡°"` |
| MaÄ‘arska                    |            `"HU"` | `"đź‡­đź‡ş"` |
| Island                    |            `"IS"` | `"đź‡®đź‡¸"` |
| Indija                      |            `"IN"` | `"đź‡®đź‡ł"` |
| Iran                       |            `"IR"` | `"đź‡®đź‡·"` |
| Irska                    |            `"IE"` | `"đź‡®đź‡Ş"` |
| Izrael                     |            `"IL"` | `"đź‡®đź‡±"` |
| Italija                      |            `"IT"` | `"đź‡®đź‡ą"` |
| Jamajka                    |            `"JM"` | `"đź‡Żđź‡˛"` |
| Japan                      |            `"JP"` | `"đź‡Żđź‡µ"` |
| Kazahstan                 |            `"KZ"` | `"đź‡°đź‡ż"` |
| Kenija                      |            `"KE"` | `"đź‡°đź‡Ş"` |
| Kosovo                     |            `"XK"` | `"đź‡˝đź‡°"` |
| Kirgistan                 |            `"KG"` | `"đź‡°đź‡¬"` |
| Letonija                    |            `"LV"` | `"đź‡±đź‡»"` |
| Liban                      |            `"LB"` | `"đź‡±đź‡§"` |
| LihtenĹˇtajn              |            `"LI"` | `"đź‡±đź‡®"` |
| Litvanija                  |            `"LT"` | `"đź‡±đź‡ą"` |
| Luksemburg                 |            `"LU"` | `"đź‡±đź‡ş"` |
| Madagaskar                 |            `"MG"` | `"đź‡˛đź‡¬"` |
| Malezija                   |            `"MY"` | `"đź‡˛đź‡ľ"` |
| Malta                      |            `"MT"` | `"đź‡˛đź‡ą"` |
| Meksiko                     |            `"MX"` | `"đź‡˛đź‡˝"` |
| Sudan                    |            `"MD"` | `"đź‡˛đź‡©"` |
| Monako                     |            `"MC"` | `"đź‡˛đź‡¨"` |
| Mongolsko                    |            `"MN"` | `"đź‡˛đź‡ł"` |
| Crna Gora                 |            `"ME"` | `"đź‡˛đź‡Ş"` |
| Maroko                    |            `"MA"` | `"đź‡˛đź‡¦"` |
| Holandija                |            `"NL"` | `"đź‡łđź‡±"` |
| Novogvinejska Zelanda                |            `"NZ"` | `"đź‡łđź‡ż"` |
| Nigerija                    |            `"NG"` | `"đź‡łđź‡¬"` |
| Severna Makedonija            |            `"MK"` | `"đź‡˛đź‡°"` |
| NorveĹˇka                     |            `"NO"` | `"đź‡łđź‡´"` |
| Pakistan                   |            `"PK"` | `"đź‡µđź‡°"` |
| Filipine                |            `"PH"` | `"đź‡µđź‡­"` |
| Poljska                     |            `"PL"` | `"đź‡µđź‡±"` |
| Portugal                   |            `"PT"` | `"đź‡µđź‡ą"` |
| Puerto Riko                |            `"PR"` | `"đź‡µđź‡·"` |
| Rumunija                    |            `"RO"` | `"đź‡·đź‡´"` |
| San Marino                 |            `"SM"` | `"đź‡¸đź‡˛"` |
| Savanska Arapska Emirati               |            `"SA"` | `"đź‡¸đź‡¦"` |
| Srbija                     |            `"RS"` | `"đź‡·đź‡¸"` |
| Singapura                  |            `"SG"` | `"đź‡¸đź‡¬"` |
| SlovaÄŤka                   |            `"SK"` | `"đź‡¸đź‡°"` |
| Slovenija                   |            `"SI"` | `"đź‡¸đź‡®"` |
| JuĹľna Afrika               |            `"ZA"` | `"đź‡żđź‡¦"` |
| JuĹľna Koreja                |            `"KR"` | `"đź‡°đź‡·"` |
| Ĺ panija                      |            `"ES"` | `"đź‡Şđź‡¸"` |
| Ĺ vajcarska                     |            `"SE"` | `"đź‡¸đź‡Ş"` |
| Ĺ vajcarska                    |            `"CH"` | `"đź‡¨đź‡­"` |
| Tajland                   |            `"TH"` | `"đź‡ąđź‡­"` |
| Trinidad i Tobago          |            `"TT"` | `"đź‡ąđź‡ą"` |
| Turska                     |            `"TR"` | `"đź‡ąđź‡·"` |
| Ukrajina                    |            `"UA"` | `"đź‡şđź‡¦"` |
| Ujedinjeni Arapski Emirati       |            `"AE"` | `"đź‡¦đź‡Ş"` |
| Sjedinjene DrĹľave              |            `"US"` | `"đź‡şđź‡¸"` |
| Uruguay                    |            `"UY"` | `"đź‡şđź‡ľ"` |
| Uzbehdristan                 |            `"UZ"` | `"đź‡şđź‡ż"` |
| Venecuela                  |            `"VE"` | `"đź‡»đź‡Ş"` |

# --hints--

`getFlag("AL")` treba da vrati `"đź‡¦đź‡±"`.

```js
assert.equal(getFlag("AL"), "đź‡¦đź‡±");
```

`getFlag("AD")` treba da vrati `"đź‡¦đź‡©"`.

```js
assert.equal(getFlag("AD"), "đź‡¦đź‡©");
```

`getFlag("AR")` treba da vrati `"đź‡¦đź‡·"`.

```js
assert.equal(getFlag("AR"), "đź‡¦đź‡·")
```

`getFlag("AM")` treba da vrati `"đź‡¦đź‡˛"`.

```js
assert.equal(getFlag("AM"), "đź‡¦đź‡˛")
```

`getFlag("AU")` treba da vrati `"đź‡¦đź‡ş"`.

```js
assert.equal(getFlag("AU"), "đź‡¦đź‡ş")
```

`getFlag("AT")` treba da vrati `"đź‡¦đź‡ą"`.

```js
assert.equal(getFlag("AT"), "đź‡¦đź‡ą")
```

`getFlag("AZ")` treba da vrati `"đź‡¦đź‡ż"`.

```js
assert.equal