id: 697a49e6ff50d756c9b69365
title: "Vežba 188: Zimični sportovi 2026, Dan 9: Skelet"
challengeType: 28
dashedName: challenge-188
---

# --description--

Dato je niz slova koji predstavlja zakrete na stazi za skeleton, odredite težinu te staze.

- Dat niz slova će sadržati samo sledeća slova:
  - `"L"` za levi zakret
  - `"R"` za desni zakret
  - `"S"` za ravni deo

- Svaka promena pravca dodaje 15 poena (`"L"` praćen sa `"R"` ili obrnuto).
- Sve ostale zakrete dodaju po 5 poena svaki (sva ostala slova `"L"` ili `"R"`).
- Ravni delovi ne dodaju poene.

Težina staze zavisi od ukupnog broja poena. Vratite:

- `"Easy"` ako je ukupan rezultat između 0 i 100
- `"Medium"` ako je ukupan rezultat između 101 i 200
- `"Hard"` ako je ukupan rezultat veći od 200

# --hints--

`getDifficulty("SLSLLSRRLSRLRL")` bi trebalo da vrati `"Easy"`.

```js
assert.equal(getDifficulty("SLSLLSRRLSRLRL"), "Easy");
```

`getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` bi trebalo da vrati `"Hard"`.

```js
assert.equal(getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard");
```

`getDifficulty("SRRRRLSLLRLRSSRLSRL")` bi trebalo da vrati `"Medium"`.

```js
assert.equal(getDifficulty("SRRRRLSLLRLRSSRLSRL"), "Medium");
```

`getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` bi trebalo da vrati `"Hard"`.

```js
assert.equal(getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard");
```

`getDifficulty("SLLSSLRLSLSLRSLSSLRL")` bi trebalo da vrati `"Medium"`.

```js
assert.equal(getDifficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium");
```

`getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` bi trebalo da vrati `"Easy"`.

```js
assert.equal(getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy");
```

# --seed--

## --seed-contents--

```js
function getDifficulty(track) {

  return track;
}
```

# --solutions--

```js
function getDifficulty(track) {
  let score = 0;

  for (let i = 0; i < track.length; i++) {
    const current = track[i];
    const previous = i > 0 ? track[i - 1] : null;

    if (current === 'S') {
      score += 0;
    } else if (current === 'L' || current === 'R') {
      if (previous && previous !== 'S' && previous !== current) {
        score += 15;
      } else if (previous !== 'S') {
        score += 5;
      } else {
        score += 5;
      }
    }
  }

  if (score <= 100) return 'Easy';
  if (score <= 200) return 'Medium';
  return 'Hard';
}
```