id: 697a49e6ff50d756c9b6935f
title: "Izazov 182: Zimske igre 2026, Dan 3: Bijaatlon"
challengeType: 28
dashedName: challenge-182
---

# --description--

Dato je niz cijelih brojeva, gdje svaka vrijednost predstavlja broj meta pogođenih u jednoj biatlon rundi; vratite ukupnu udaljenost kazne koju sportista mora da skija.

- Svaka runda ima 5 meta.
- Svaki propušteni met počinje krug kazne od 150 metara.

# --hints--

`calculatePenaltyDistance([4, 4])` bi trebalo da vrati `300`.

```js
assert.equal(calculatePenaltyDistance([4, 4]), 300);
```

`calculatePenaltyDistance([5, 5])` bi trebalo da vrati `0`.

```js
assert.equal(calculatePenaltyDistance([5, 5]), 0);
```

`calculatePenaltyDistance([4, 5, 3, 5])` bi trebalo da vrati `450`.

```js
assert.equal(calculatePenaltyDistance([4, 5, 3, 5]), 450);
```

`calculatePenaltyDistance([5, 4, 5, 5])` bi trebalo da vrati `150`.

```js
assert.equal(calculatePenaltyDistance([5, 4, 5, 5]), 150);
```

`calculatePenaltyDistance([4, 3, 0, 3])` bi trebalo da vrati `1500`.

```js
assert.equal(calculatePenaltyDistance([4, 3, 0, 3]), 1500);
```

# --seed--

## --seed-contents--

```js
function calculatePenaltyDistance(rounds) {

  return rounds;
}
```

# --solutions--

```js
function calculatePenaltyDistance(rounds) {
  let totalPenalty = 0;

  for (const hits of rounds) {
    const misses = 5 - hits;
    totalPenalty += misses * 150;
  }

  return totalPenalty;
}
```