id: 697a49e6ff50d756c9b69360
title: "Vežba 183: Zimske Igre 2026 Dan 4: Skoka s granja"
challengeType: 28
dashedName: challenge-183
---

# --description--

Uzimajući u obzir bodove udaljenosti, bodove stila, vrednost kompenzacije vetra i bonus K-bodova, izračunajte svoj rezultat skoka s granja i utvrdite da li ste osvojili medalju.

- Vaš rezultat se računa zbiranjem četiri navedena boda.

Ukupni rezultati drugih trenutnih skajača su:

```sh
165.5
172.0
158.0
180.0
169.5
175.0
162.0
170.0
```

- Ako je vaš rezultat najbolji, vratite `"Gold"`
- Ako je drugi po kvalitetu, vratite `"Silver"`
- Ako je treći po kvalitetu, vratite `"Bronze"`
- U suprotnom, vratite `"No Medal"`

# --hints--

`skiJumpMedal(125.0, 58.0, 0.0, 6.0)` treba da vrati `"Gold"`.

```js
assert.equal(skiJumpMedal(125.0, 58.0, 0.0, 6.0), "Gold");
```

`skiJumpMedal(119.0, 50.0, 1.0, 4.0)` treba da vrati `"Bronze"`.

```js
assert.equal(skiJumpMedal(119.0, 50.0, 1.0, 4.0), "Bronze");
```

`skiJumpMedal(122.0, 52.0, -1.0, 4.0)` treba da vrati `"Silver"`.

```js
assert.equal(skiJumpMedal(122.0, 52.0, -1.0, 4.0), "Silver");
```

`skiJumpMedal(118.0, 50.5, -1.5, 4.0)` treba da vrati `"No Medal"`.

```js
assert.equal(skiJumpMedal(118.0, 50.5, -1.5, 4.0), "No Medal");
```

`skiJumpMedal(124.0, 50.5, 2.0, 5.0)` treba da vrati `"Gold"`.

```js
assert.equal(skiJumpMedal(124.0, 50.5, 2.0, 5.0), "Gold");
```

`skiJumpMedal(119.0, 49.5, 0.0, 3.0)` treba da vrati `"No Medal"`.

```js
assert.equal(skiJumpMedal(119.0, 49.5, 0.0, 3.0), "No Medal");
```

# --seed--

## --seed-contents--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {

  return distancePoints;
}
```

# --solutions--

```js
function skiJumpMedal(distancePoints, stylePoints, windComp, kPointBonus) {
  const myScore = distancePoints + stylePoints + windComp + kPointBonus;
  const otherScores = [165.5, 172.0, 158.0, 180.0, 169.5, 175.0, 162.0, 170.0];
  const allScores = [...otherScores, myScore];
  allScores.sort((a, b) => b - a);
  const rank = allScores.indexOf(myScore) + 1;

  if (rank === 1) return "Gold";
  if (rank === 2) return "Silver";
  if (rank === 3) return "Bronze";
  return "No Medal";
}
```