id: 697a49e6ff50d756c9b69362
title: "Challenge 185: Winter Games 2026 Day 6: Figure Skating"
challengeType: 28
dashedName: challenge-185
---

# --description--

Uzimajući niz podataka ocena sudija i opcionih kazni, izračunajte konačnu ocenu za figure skating.

Prvi argument je niz podataka od 10 ocena sudija, svaka je broj od 0 do 10. Uklonite najviše i najniže ocene sudija i saberi preostalih 8 ocena kako biste dobili osnovnu ocenu.

Bilo koji dodatni argument prosleđen funkciji su kazne. Smanjite sve kazne od osnovne ocene da biste dobili konačnu ocenu.

# --hints--

`computeScore([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1)` bi trebalo da vrati `64`.

```js
assert.equal(computeScore([10, 8, 9, 6, 9, 8, 8, 9, 7, 7], 1), 64);
```

`computeScore([10, 10, 10, 10, 10, 10, 10, 10, 10, 10])` bi trebalo da vrati `80`.

```js
assert.equal(computeScore([10, 10, 10, 10, 10, 10, 10, 10, 10, 10]), 80);
```

`computeScore([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1)` bi trebalo da vrati `67`.

```js
assert.equal(computeScore([10, 8, 9, 10, 9, 8, 8, 9, 10, 7], 1, 2, 1), 67);
```

`computeScore([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0)` bi trebalo da vrati `67.5`.

```js
assert.equal(computeScore([8.0, 8.5, 9.0, 8.5, 9.0, 8.0, 9.0, 8.5, 9.0, 8.5], 0.5, 1.0), 67.5);
```

`computeScore([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5)` bi trebalo da vrati `59`.

```js
assert.equal(computeScore([6.0, 8.5, 7.0, 9.0, 7.5, 8.0, 6.5, 9.5, 7.0, 8.0], 1.5, 0.5, 0.5), 59);
```

# --seed--

## --seed-contents--

```js
function computeScore(judgeScores, ...penalties) {

  return judgeScores;
}
```

# --solutions--

```js
function computeScore(judgeScores, ...penalties) {
  const sortedScores = [...judgeScores].sort((a, b) => a - b);
  const baseScore = sortedScores.slice(1, 9).reduce((sum, s) => sum + s, 0);
  const totalPenalty = penalties.reduce((sum, p) => sum + p, 0);

  return baseScore - totalPenalty;
}
```