id: 697a49e9860d24853adef67c
title: "Vežba 191: Zimske igre 2026, Dan 12: Bobsled"
challengeType: 28
dashedName: challenge-191
---

# --description--

Dati je niz podataka koji predstavlja težinu sportista u bobsled timu i broj koji predstavlja težinu bobsleda; utvrdite da li je tim kvalifikovan za takmičenje.

- Dužina niza određuje veličinu tima: timovi od 1, 2 ili 4 osobe.
- Sve navedene vrednosti težine su u kilogramima (kg).

Bobsled (samo sled) mora imati minimalnu težinu od:

- 162 kg za tim od 1 osobe
- 170 kg za tim od 2 osobe
- 210 kg za tim od 4 osobe

Ukupna težina bobsleda (sportista zajedno sa sledom) ne sme preći:

- 247 kg za tim od 1 osobe
- 390 kg za tim od 2 osobe
- 630 kg za tim od 4 osobe

Vratite `"Eligible"` ako tim ispunjava sve uslove, ili `"Not Eligible"` ako tim ne ispuni jedan ili više ovih uslova.

# --hints--

`checkEligibility([78], 165)` bi trebalo da vrati `"Eligible"`.

```js
assert.equal(checkEligibility([78], 165), "Eligible");
```

`checkEligibility([80], 160)` bi trebalo da vrati `"Not Eligible"`.

```js
assert.equal(checkEligibility([80], 160), "Not Eligible");
```

`checkEligibility([80], 170)` bi trebalo da vrati `"Not Eligible"`.

```js
assert.equal(checkEligibility([80], 170), "Not Eligible");
```

`checkEligibility([85, 90], 170)` bi trebalo da vrati `"Eligible"`.

```js
assert.equal(checkEligibility([85, 90], 170), "Eligible");
```

`checkEligibility([85, 95], 168)` bi trebalo da vrati `"Not Eligible"`.

```js
assert.equal(checkEligibility([85, 95], 168), "Not Eligible");
```

`checkEligibility([112, 97], 185)` bi trebalo da vrati `"Not Eligible"`.

```js
assert.equal(checkEligibility([112, 97], 185), "Not Eligible");
```

`checkEligibility([110, 102, 90, 106], 222)` bi trebalo da vrati `"Eligible"`.

```js
assert.equal(checkEligibility([110, 102, 90, 106], 222), "Eligible");
```

`checkEligibility([106, 99, 90, 88], 205)` bi trebalo da vrati `"Not Eligible"`.

```js
assert.equal(checkEligibility([106, 99, 90, 88], 205), "Not Eligible");
```

`checkEligibility([106, 99, 103, 96], 227)` bi trebalo da vrati `"Not Eligible"`.

```js
assert.equal(checkEligibility([106, 99, 103, 96], 227), "Not Eligible");
```

# --seed--

## --seed-contents--

```js
function checkEligibility(athleteWeights, sledWeight) {

  return athleteWeights;
}
```

# --solutions--

```js
function checkEligibility(athleteWeights, sledWeight) {
  const teamSize = athleteWeights.length;

  const rules = {
    1: { minSled: 162, maxTotal: 247 },
    2: { minSled: 170, maxTotal: 390 },
    4: { minSled: 210, maxTotal: 630 }
  };

  const { minSled, maxTotal } = rules[teamSize];
  const athleteTotal = athleteWeights.reduce((sum, w) => sum + w, 0);
  const totalWeight = athleteTotal + sledWeight;

  if (sledWeight < minSled) {
    return "Not Eligible";
  }

  if (totalWeight > maxTotal) {
    return "Not Eligible";
  }

  return "Eligible";
}
```