id: 698a1a73ade5ac0e19180fa8
title: "Vežba 205: Brojanje savršenih kubova"
challengeType: 28
dashedName: challenge-205
---

# --description--

Datim dva celobrojka, odredite koliko savršenih kubova postoji u opsegu između i uključujući te dve broja.

- Niža broj nije garantovan kao prvi argument.
- Broj je savršen kub ako postoji ceo broj (`n`) takav da `n * n * n = number`. Na primer, 27 je savršen kub jer je `3 * 3 * 3 = 27`.

# --hints--

`countPerfectCubes(3, 30)` treba da vrati `2`.

```js
assert.equal(countPerfectCubes(3, 30), 2);
```

`countPerfectCubes(1, 30)` treba da vrati `3`.

```js
assert.equal(countPerfectCubes(1, 30), 3);
```

`countPerfectCubes(30, 0)` treba da vrati `4`.

```js
assert.equal(countPerfectCubes(30, 0), 4);
```

`countPerfectCubes(-64, 64)` treba da vrati `9`.

```js
assert.equal(countPerfectCubes(-64, 64), 9);
```

`countPerfectCubes(9214, -8127)` treba da vrati `41`.

```js
assert.equal(countPerfectCubes(9214, -8127), 41);
```

# --seed--

## --seed-contents--

```js
function countPerfectCubes(a, b) {

  return a;
}
```

# --solutions--

```js
function countPerfectCubes(a, b) {
  const start = Math.min(a, b);
  const end = Math.max(a, b);

  let count = 0;
  let n = 0;

  while (n ** 3 <= end) {
    if (n ** 3 >= start) count++;
    n++;
  }

  n = -1;
  while (n ** 3 >= start) {
    if (n ** 3 <= end) count++;
    n--;
  }

  return count;
}
```