id: 698a1a863194f1f4e63f6460
title: "Vežba 209: Veličina elementa"
challengeType: 28
dashedName: challenge-209
---

# --description--

Uzimajući u obzir veličinu prozora, širinu elementa u jedinicama viewport širine `"vw"` i visinu elementa u jedinicama viewport visine `"vh"`, odredite veličinu elementa u pikselima.

- Veličina prozora koja se pruža i veličina elementa koja se vraća su nizovi karaktera u formatu `"širina x visina"`, npr., `"1200 x 800"`.

- Jedinice `"vw"` predstavljaju procenat širine prozora. Na primjer, `"50vw"` je 50% širine prozora.
- Jedinice `"vh"` predstavljaju procenat visine prozora. Na primjer, `"50vh"` je 50% visine prozora.

# --hints--

`getElementSize("1200 x 800", "50vw", "50vh")` bi trebalo da vrati `"600 x 400"`.

```js
assert.equal(getElementSize("1200 x 800", "50vw", "50vh"), "600 x 400");
```

`getElementSize("320 x 480", "25vw", "50vh")` bi trebalo da vrati `"80 x 240"`.

```js
assert.equal(getElementSize("320 x 480", "25vw", "50vh"), "80 x 240");
```

`getElementSize("1000 x 500", "7vw", "3vh")` bi trebalo da vrati `"70 x 15"`.

```js
assert.equal(getElementSize("1000 x 500", "7vw", "3vh"), "70 x 15");
```

`getElementSize("1920 x 1080", "95vw", "100vh")` bi trebalo da vrati `"1824 x 1080"`.

```js
assert.equal(getElementSize("1920 x 1080", "95vw", "100vh"), "1824 x 1080");
```

`getElementSize("1200 x 800", "0vw", "0vh")` bi trebalo da vrati `"0 x 0"`.

```js
assert.equal(getElementSize("1200 x 800", "0vw", "0vh"), "0 x 0");
```

`getElementSize("1440 x 900", "100vw", "114vh")` bi trebalo da vrati `"1440 x 1026"`.

```js
assert.equal(getElementSize("1440 x 900", "100vw", "114vh"), "1440 x 1026");
```

# --seed--

## --seed-contents--

```js
function getElementSize(windowSize, elementVw, elementVh) {

  return windowSize;
}
```

# --solutions--

```js
function getElementSize(windowSize, elementVw, elementVh) {
  const [windowWidth, windowHeight] = windowSize
    .split(" x ")
    .map(Number);

  const vw = Number(elementVw.replace("vw", ""));
  const vh = Number(elementVh.replace("vh", ""));

  const widthPx = (vw / 100) * windowWidth;
  const heightPx = (vh / 100) * windowHeight;

  return `${widthPx} x ${heightPx}`;
}
```