id: 698a1a73ade5ac0e19180fa0
title: "Zadatak 197: Kompatibilnost krvnih grupa"
challengeType: 28
dashedName: challenge-197
---

# --description--

Uzimajući u obzir krvnu grupu donatora i krvnu grupu primača, utvrdite da li donator može dati krv primaču.

Svaka krvna grupa ima:

- Slovo: `"A"`, `"B"`, `"AB"`, ili `"O"`
- I Rh faktor: `"+"` ili `"-"`

Krvne grupe će biti jedno validno slovo praćeno Rh faktorom. Na primer, `"AB+"` i `"O-"` su validne krvne grupe.

Pravila za slova:

- `"O"` može dati krv bilo kojoj drugoj grupi slova.
- `"A"` može dati krv `"A"` i `"AB"`.
- `"B"` može dati krv `"B"` i `"AB"`.
- `"AB"` može dati krv samo `"AB"`.

Pravila za Rh:

- Negativan (`"-"`) može dati krv `"-"` i `"+"`.
- Pozitivan (`"+"`) može dati krv samo `"+"`.

Moraju biti ispunjena sva pravila slova i Rh faktora kako bi donator mogao dati krv primaču.

# --hints--

`canDonate("B+", "B+")` treba da vrati `true`.

```js
assert.isTrue(canDonate("B+", "B+"));
```

`canDonate("O-", "AB-")` treba da vrati `true`.

```js
assert.isTrue(canDonate("O-", "AB-"));
```

`canDonate("O+", "A-")` treba da vrati `false`.

```js
assert.isFalse(canDonate("O+", "A-"));
```

`canDonate("A+", "AB+")` treba da vrati `true`.

```js
assert.isTrue(canDonate("A+", "AB+"));
```

`canDonate("A-", "B-")` treba da vrati `false`.

```js
assert.isFalse(canDonate("A-", "B-"));
```

`canDonate("B-", "AB+")` treba da vrati `true`.

```js
assert.isTrue(canDonate("B-", "AB+"));
```

`canDonate("B-", "A+")` treba da vrati `false`.

```js
assert.isFalse(canDonate("B-", "A+"));
```

`canDonate("O-", "O+")` treba da vrati `true`.

```js
assert.isTrue(canDonate("O-", "O+"));
```

`canDonate("O+", "O-")` treba da vrati `false`.

```js
assert.isFalse(canDonate("O+", "O-"));
```

`canDonate("AB+", "AB-")` treba da vrati `false`.

```js
assert.isFalse(canDonate("AB+", "AB-"));
```

# --seed--

## --seed-contents--

```js
function canDonate(donor, recipient) {

  return donor;
}
```

# --solutions--

```js
function canDonate(donor, recipient) {
  const donorType = donor.slice(0, -1);
  const donorRh = donor.slice(-1);

  const recipientType = recipient.slice(0, -1);
  const recipientRh = recipient.slice(-1);

  const aboCompatibility = {
    "O": ["A", "B", "AB", "O"],
    "A": ["A", "AB"],
    "B": ["B", "AB"],
    "AB": ["AB"]
  };

  const aboMatch = aboCompatibility[donorType].includes(recipientType);
  const rhMatch =
    donorRh === "-" || (donorRh === "+" && recipientRh === "+");

  return aboMatch && rhMatch;
}
```