id: 698a1a863194f1f4e63f645f
title: "Vežba 208: Kretanje po stazi"
challengeType: 28
dashedName: challenge-208
---

# --description--

Dato je niz niz znakova koji predstavlja mapu vaše staze; vratite string znakova koji prikazuje korake potrebne za dolazak do cilja.

Omjer znakova će imati sljedeće vrijednosti:

- `"C"`: Vaša trenutna lokacija
- `"G"`: Ваш cilj
- `"T"`: Prolazni dijelovi staze
- `"-"`: Neprozlazni dijelovi mape

Vratite string znakova koji prikazuje korake potrebne za praćenje puta od vaše trenutne lokacije do vašeg cilja, gdje je:

- `"R"` korak kretanja udesno
- `"D"` korak kretanja nadole
- `"L"` korak kretanja ulijevo
- `"U"` korak kretanja nagore

- Uvijek će postojati jedna staza koja povezuje bez grana, od vaše trenutne lokacije do vašeg cilja.
- Svaki dio staze ima dva susjedna prolaziva dijela koji ga dodiruju.

Na primjer, ako je dato:

```js
[
  "-CT--",
  "--T--",
  "--TT-",
  "---T-",
  "---G-"
]
```

Vratite `"RDDRDD"`.

# --hints--

`navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"])` bi trebao vratiti `"RDDRDD"`.

```js
assert.equal(navigateTrail(["-CT--", "--T--", "--TT-", "---T-", "---G-"]), "RDDRDD");
```

`navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"])` bi trebao vratiti `"RRUUURR"`.

```js
assert.equal(navigateTrail(["-----", "--TTG", "--T--", "--T--", "CTT--"]), "RRUUURR");
```

`navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"])` bi trebao vratiti `"DLDDRRRRD"`.

```js
assert.equal(navigateTrail(["-C----", "TT----", "T-----", "TTTTT-", "----G-"]), "DLDDRRRRD");
```

`navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"])` bi trebao vratiti `"RRRDDL"`.

```js
assert.equal(navigateTrail(["--------", "-CTTT---", "----T---", "---GT---", "--------"]), "RRRDDL");
```

`navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"])` bi trebao vratiti `"ULLLUUURRRRRRDDDR"`.

```js
assert.equal(navigateTrail(["TTTTTTT-", "T-----T-", "T-----T-", "TTTT--TG", "---C----"]), "ULLLUUURRRRRRDDDR");
```

# --seed--

## --seed-contents--

```js
function navigateTrail(map) {

  return map;
}
```

# --solutions--

```js
function navigateTrail(map) {
  const rows = map.length;
  const cols = map[0].length;

  const directions = [
    [0, 1, "R"],
    [1, 0, "D"],
    [0, -1, "L"],
    [-1, 0, "U"],
  ];

  let row, col;
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (map[r][c] === "C") {
        row = r;
        col = c;
      }
    }
  }

  let result = "";
  let prev = null;

  while (map[row][col] !== "G") {
    for (let [dr, dc, move] of directions) {
      const newRow = row + dr;
      const newCol = col + dc;

      if (
        newRow >= 0 &&
        newRow < rows &&
        newCol >= 0 &&
        newCol < cols &&
        (map[newRow][newCol] === "T" ||
          map[newRow][newCol] === "G") &&
        (!prev || newRow !== prev[0] || newCol !== prev[1])
      ) {
        result += move;
        prev = [row, col];
        row = newRow;
        col = newCol;
        break;
      }
    }
  }

  return result;
}
```