id: 698a1a73ade5ac0e19180fa9
title: "Vežba 206: Vrednosti karata"
challengeType: 28
dashedName: challenge-206
---

# --description--

Kada dobijete niz podataka igračkih karata, vratite novi niz podataka sa numeričkom vrednošću svake karte.

Vrednosti karata:

- As (`"A"`) ima vrednost 1.
- Karate sa brojevima (`"2"` - `"10"`) imaju vrednost jednaku broju koji se nalazi ispred njih: 2 - 10, po redu.
- Karate sa likovima: Džek (`"J"`), Kraljica (`"Q"`), i Kralj (`"K"`) svaki imaju vrednost 10.

Vrste karata:

- Svaka karta ima vrstu: Pik (`"S"`), Špicevi (`"C"`), Karo (`"D"`), ili Srca (`"H"`).

Struktura karte:

- Svaka karta predstavljena je kao niz slova: `"valueSuit"`. Na primer: `"AS"` je As od Pika, `"10H"` je Deset od Srca, a `"QC"` je Kraljica od Špiceva.

# --hints--

`cardValues(["3H", "4D", "5S"])` treba da vrati `[3, 4, 5]`.

```js
assert.deepEqual(cardValues(["3H", "4D", "5S"]), [3, 4, 5]);
```

`cardValues(["AS", "10S", "10H", "6D", "7D"])` treba da vrati `[1, 10, 10, 6, 7]`.

```js
assert.deepEqual(cardValues(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7]);
```

`cardValues(["8D", "QS", "2H", "JC", "9C"])` treba da vrati `[8, 10, 2, 10, 9]`.

```js
assert.deepEqual(cardValues(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9]);
```

`cardValues(["AS", "KS"])` treba da vrati `[1, 10]`.

```js
assert.deepEqual(cardValues(["AS", "KS"]), [1, 10]);
```

`cardValues(["10H", "JH", "QH", "KH", "AH"])` treba da vrati `[10, 10, 10, 10, 1]`.

```js
assert.deepEqual(cardValues(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1]);
```

# --seed--

## --seed-contents--

```js
function cardValues(cards) {

  return cards;
}
```

# --solutions--

```js
function cardValues(cards) {
  const values = [];

  for (let card of cards) {
    let valueStr = card.slice(0, card.length - 1);

    let value;
    if (valueStr === "A") {
      value = 1;
    } else if (["J", "Q", "K"].includes(valueStr)) {
      value = 10;
    } else {
      value = parseInt(valueStr, 10);
    }

    values.push(value);
  }

  return values;
}
```