id: 698a1a73ade5ac0e19180fa1
title: "Vežba 198: Računanje radnih dana"
challengeType: 28
dashedName: challenge-198
---

# --description--

Ako dobijete početni i krajnji datum, vratite broj radnih dana između njih.

- Datumi su u formatu `"YYYY-MM-DD"`.
- Radni dani su dane sedmice (Ponedeljak do Petak).
- Vikendi nisu radni dani (Subota i Nedelja).
- Uključite oba datuma, početni i krajnji, kada brojiте.

# --hints--

`countBusinessDays("2026-02-24", "2026-02-26")` bi trebalo da vrati `3`.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-26"), 3);
```

`countBusinessDays("2026-02-24", "2026-02-28")` bi trebalo da vrati `4`.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-28"), 4);
```

`countBusinessDays("2026-02-21", "2026-03-01")` bi trebalo da vrati `5`.

```js
assert.equal(countBusinessDays("2026-02-21", "2026-03-01"), 5);
```

`countBusinessDays("2026-03-08", "2026-03-17")` bi trebalo da vrati `7`.

```js
assert.equal(countBusinessDays("2026-03-08", "2026-03-17"), 7);
```

`countBusinessDays("2026-02-24", "2027-02-24")` bi trebalo da vrati `262`.

```js
assert.equal(countBusinessDays("2026-02-24", "2027-02-24"), 262);
```

# --seed--

## --seed-contents--

```js
function countBusinessDays(start, end) {

  return start;
}
```

# --solutions--

```js
function countBusinessDays(start, end) {
  const startDate = new Date(start + "T00:00:00Z");
  const endDate = new Date(end + "T00:00:00Z");

  let count = 0;
  let current = new Date(startDate);

  while (current <= endDate) {
    const day = current.getUTCDay();
    if (day !== 0 && day !== 6) count++;
    current.setUTCDate(current.getUTCDate() + 1);
  }

  return count;
}
```