id: 698a1a73ade5ac0e19180fa7
title: "Vežba 204: Sabiranje slova"
challengeType: 28
dashedName: challenge-204
---

# --description--

Kada vam je dat niz znakova, vratite zbir njihovih slova.

- Slovi su A-Z u velikim ili malim slovima.
- Vrednosti slova su: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26.
- Velika i mala slova imaju istu vrednost.
- Ignorišite sve ne-slovene znakove.

# --hints--

`sumLetters("Hello")` bi trebalo da vrati `52`.

```js
assert.equal(sumLetters("Hello"), 52);
```

`sumLetters("freeCodeCamp")` bi trebalo da vrati `94`.

```js
assert.equal(sumLetters("freeCodeCamp"), 94);
```

`sumLetters("The quick brown fox jumps over the lazy dog.")` bi trebalo da vrati `473`.

```js
assert.equal(sumLetters("The quick brown fox jumps over the lazy dog."), 473);
```

`sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci.")` bi trebalo da vrati `1681`.

```js
assert.equal(sumLetters("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Aenean ex nisl, pretium eu varius blandit, facilisis quis eros. Vestibulum ante ipsum primis in faucibus orci."), 1681);
```

`sumLetters("</404>")` bi trebalo da vrati `0`.

```js
assert.equal(sumLetters("</404>"), 0);
```

# --seed--

## --seed-contents--

```js
function sumLetters(str) {

  return str;
}
```

# --solutions--

```js
function sumLetters(str) {
  let total = 0;
  for (let char of str) {
    let upper = char.toUpperCase();
    if (upper >= "A" && upper <= "Z") {
      total += upper.charCodeAt(0) - "A".charCodeAt(0) + 1;
    }
  }
  return total;
}
```