id: 698a1a73ade5ac0e19180fa6
title: "Vežba 203: Niz podataka bez podniza"
challengeType: 28
dashedName: challenge-203
---

# --description--

Ako vam je dat niz podataka, proverite da li je to ravan niz.

- Niz je ravan ako nijedan od njegovih elemenata nije niz sam po sebi.

# --hints--

`isFlat([1, 2, 3, 4])` treba da vrati `true`.

```js
assert.isTrue(isFlat([1, 2, 3, 4]));
```

`isFlat([1, [2, 3], 4])` treba da vrati `false`.

```js
assert.isFalse(isFlat([1, [2, 3], 4]));
```

`isFlat([1, 0, false, true, "a", "b"])` treba da vrati `true`.

```js
assert.isTrue(isFlat([1, 0, false, true, "a", "b"]));
```

`isFlat(["a", [0], "b", true])` treba da vrati `false`.

```js
assert.isFalse(isFlat(["a", [0], "b", true]));
```

`isFlat([1, [2, [3, [4, [5]]]], 6])` treba da vrati `false`.

```js
assert.isFalse(isFlat([1, [2, [3, [4, [5]]]], 6]));
```

# --seed--

## --seed-contents--

```js
function isFlat(arr) {

  return arr;
}
```

# --solutions--

```js
function isFlat(arr) {
  for (let el of arr) {
    if (Array.isArray(el)) return false;
  }
  return true;
}
```