id: 698a1a863194f1f4e63f645e
title: "Vežba 207: Najmanji razmak"
challengeType: 28
dashedName: challenge-207
---

# --description--

Dato vam je niz slova; vratite podniz niz između parova jednakih slova koji imaju najmanji razmak slova između njih (najmanji razmak).

- Uvek će postojati barem jedan par jednakih slova.
- Podniz niz koji se vraća ne sme sadržavati jednaka slova.
- Ako postoji dva ili više razmaka iste dužine, vratite slova iz prvog razmaka.

Na primer, ako je dato `"ABCDAC"`, vratiti je `"DA"` jer:

- Samo su "A" i "C" ponovljena slova u nizu.
- Broj slova između dva slova "A" je 3, a između slova "C" je 2.
- Stoga se vraća niz između dva slova "C".

# --hints--

`smallestGap("ABCDAC")` bi trebalo da vrati `"DA"`.

```js
assert.equal(smallestGap("ABCDAC"), "DA");
```

`smallestGap("racecar")` bi trebalo da vrati `"e"`.

```js
assert.equal(smallestGap("racecar"), "e");
```

`smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` bi trebalo da vrati `"#q6e&rkf(p"`.

```js
assert.equal(smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p");
```

`smallestGap("Hello World")` bi trebalo da vrati `""`.

```js
assert.equal(smallestGap("Hello World"), "");
```

`smallestGap("The quick brown fox jumps over the lazy dog.")` bi trebalo da vrati `"fox"`.

```js
assert.equal(smallestGap("The quick brown fox jumps over the lazy dog."), "fox");
```

# --seed--

## --seed-contents--

```js
function smallestGap(str) {

  return str;
}
```

# --solutions--

```js
function smallestGap(str) {
  let minGap = Infinity;
  let result = "";

  for (let i = 0; i < str.length; i++) {
    for (let j = i + 1; j < str.length; j++) {
      if (str[i] === str[j]) {
        const gap = j - i - 1;

        if (gap < minGap) {
          minGap = gap;
          result = str.slice(i + 1, j);
        }

        break;
      }
    }
  }

  return result;
}
```