id: 6994cff2290543b3aec9f50f
title: "Vežba 210: Provera validnosti HSL"
challengeType: 28
dashedName: challenge-210
---

# --description--

Dato je niz znakova; utvrdite da li predstavlja validnu vrednost CSS boje `hsl()`.

- Validna HSL vrednost mora biti u formatu `"hsl(h, s%, l%)"`, gde:
  - `h` (hue) mora biti broj između 0 i 360 (uključujući).
  - `s` (saturation) mora biti procenat između 0% i 100%.
  - `l` (lightness) mora biti procenat između 0% i 100%.

- Dopušteni su samo razmaci:
  - Nakon otvarajuće zagrade
  - Pre i/ili posle zareza
  - Pre i/ili posle zatvarajuće zagrade

- Opciono, vrednost može da se završava tačkom-zarezom (";").

Na primer, `"hsl(240, 50%, 50%)"` je validna HSL vrednost.

# --hints--

`isValidHSL("hsl(240, 50%, 50%)")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidHSL("hsl(240, 50%, 50%)"));
```

`isValidHSL("hsl( 200 , 50% , 75% )")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidHSL("hsl( 200 , 50% , 75% )"));
```

`isValidHSL("hsl(99,60%,80%);")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidHSL("hsl(99,60%,80%);"));
```

`isValidHSL("hsl(0, 0%, 0%) ;")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidHSL("hsl(0, 0%, 0%) ;"));
```

`isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;"));
```

`isValidHSL("hsl(361, 50%, 80%)")` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidHSL("hsl(361, 50%, 80%)"));
```

`isValidHSL("hsl(300, 101%, 70%)")` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidHSL("hsl(300, 101%, 70%)"));
```

`isValidHSL("hsl(200, 55%, 75)")` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidHSL("hsl(200, 55%, 75)"));
```

`isValidHSL("hsl (80, 20%, 10%)")` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidHSL("hsl (80, 20%, 10%)"));
```

# --seed--

## --seed-contents--

```js
function isValidHSL(hsl) {

  return hsl;
}
```

# --solutions--

```js
function isValidHSL(hsl) {
  const regex = /^hsl\(\s*(\d{1,3})\s*,\s*(\d{1,3})%\s*,\s*(\d{1,3})%\s*\)\s*;?$/;

  const match = hsl.match(regex);
  if (!match) return false;

  const hue = Number(match[1]);
  const sat = Number(match[2]);
  const light = Number(match[3]);

  if (hue < 0 || hue > 360) return false;
  if (sat < 0 || sat > 100) return false;
  if (light < 0 || light > 100) return false;

  return true;
}
```