id: 699c8e045ee7cb94ed2322d6
title: "Vežba 215: Kalkulator naknade za parking"
challengeType: 28
dashedName: challenge-215
---

# --description--

Kada vam se date dve sekvence od dva karaktera koje predstavljaju vreme kada ste parkirali i vreme kada ste napustili vozilo, izračunajte naknadu za parking.

- Date sekvence karaktera će biti u formatu `"HH:MM"` koristeći 24-satni sistem. Na primer, `"14:00"` je 2 popodne.
- Prva sekvenca karaktera će biti vreme kada ste parkirali, a druga vreme kada ste napustili vozilo.
- Ako je vreme napuštanja ranije nego vreme dolaska, to znači da je parking trajao do posle ponoći i u sledećem danu.

Pravila naknade:

- Svaki sat parkiranja košta 3$.
- Nepuni sati se zaokružuju do sledećeg punog sata.
- Ako je parking trajao preko noći (nakon ponoći), naplaže se dodatna noćna naknada od 10$.
- Postoji minimalna naknada od 5$ (primjenjuje se samo ako je ukupan iznos manji od 5$).

Vratite ukupnu cenu u formatu `"$cost"`, npr. `"$5"`.

# --hints--

`calculateParkingFee("09:00", "11:00")` bi trebalo da vrati `"$6"`.

```js
assert.equal(calculateParkingFee("09:00", "11:00"), "$6");
```

`calculateParkingFee("10:00", "10:30")` bi trebalo da vrati `"$5"`.

```js
assert.equal(calculateParkingFee("10:00", "10:30"), "$5");
```

`calculateParkingFee("08:10", "10:45")` bi trebalo da vrati `"$9"`.

```js
assert.equal(calculateParkingFee("08:10", "10:45"), "$9");
```

`calculateParkingFee("14:40", "23:10")` bi trebalo da vrati `"$27"`.

```js
assert.equal(calculateParkingFee("14:40", "23:10"), "$27");
```

`calculateParkingFee("18:15", "01:30")` bi trebalo da vrati `"$34"`.

```js
assert.equal(calculateParkingFee("18:15", "01:30"), "$34");
```

`calculateParkingFee("11:11", "11:10")` bi trebalo da vrati `"$82"`.

```js
assert.equal(calculateParkingFee("11:11", "11:10"), "$82");
```

# --seed--

## --seed-contents--

```js
function calculateParkingFee(parkTime, pickupTime) {

  return parkTime;
}
```

# --solutions--

```js
function calculateParkingFee(parkTime, pickupTime) {
  function toMinutes(time) {
    const [hours, minutes] = time.split(":").map(Number);
    return hours * 60 + minutes;
  }

  const entryMinutes = toMinutes(parkTime);
  const exitMinutes = toMinutes(pickupTime);

  let totalMinutes;
  let overnight = false;

  if (exitMinutes < entryMinutes) {
    overnight = true;
    totalMinutes = (24 * 60 - entryMinutes) + exitMinutes;
  } else {
    totalMinutes = exitMinutes - entryMinutes;
  }

  const hours = Math.ceil(totalMinutes / 60);
  let cost = hours * 3;

  if (overnight) {
    cost += 10;
  }

  if (cost < 5) {
    cost = 5;
  }

  return `$${cost}`;
}
```