id: 699c8e045ee7cb94ed2322d8
title: "Vežba 217: Uhvaćeni šahovski komadi"
challengeType: 28
dashedName: challenge-217
---

# --description--

Dato je popisa znakova koji predstavljaju šahovski komade preostale na tabli; izračunajte vrijednost komada koje je vaš protivnik uhvatio.

U šahu, počinjete sa 16 komada:

| Komad | Skraćenica | Broj | Vrijednost |
| :--- | :--- | :--- | :--- |
| Pešak (Pawn) | `"P"` | 8 | 1 |
| Tore (Rook) | `"R"` | 2 | 5 |
| Konj (Knight) | `"N"` | 2 | 3 |
| Slon (Bishop) | `"B"` | 2 | 3 |
| Kraljica (Queen) | `"Q"` | 1 | 9 |
| Kralj (King) | `"K"` | 1 | 0 |

- Pruženi niz podataka će sadržati samo skraćenice navedene iznad.
- Bilo koji od 16 komada koji nisu u priloženom nizu su uhvaćeni.
- Vratite ukupnu vrijednost svih uhvaćenih komada, osim ako...
- Ako je Kralj uhvaćen, vratiti `"Checkmate"`.

# --hints--

`getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"])` treba da vrati `8`.

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"]), 8);
```

`getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"])` treba da vrati `26`.

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"]), 26);
```

`getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"])` treba da vrati `16`.

```js
assert.equal(getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"]), 16);
```

`getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"])` treba da vrati `4`.

```js
assert.equal(getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"]), 4);
```

`getCapturedValue(["P", "K"])` treba da vrati `38`.

```js
assert.equal(getCapturedValue(["P", "K"]), 38);
```

`getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"])` treba da vrati `0`.

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"]), 0);
```

`getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"])` treba da vrati `"Checkmate"`.

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"]), "Checkmate");
```

# --seed--

## --seed-contents--

```js
function getCapturedValue(pieces) {

  return pieces;
}
```

# --solutions--

```js
function getCapturedValue(pieces) {
  if (!pieces.includes("K")) return "Checkmate";

  const values = { P:1, R:5, N:3, B:3, Q:9, K:0 };
  let remaining = "PPPPPPPPRRNNBBQK";

  for (const p of pieces) {
    remaining = remaining.replace(p, "");
  }

  return [...remaining].reduce((sum, p) => sum + values[p], 0);
}
```