id: 699c8e045ee7cb94ed2322d5
title: "Vežba 214: Proveri lanac domina"
challengeType: 28
dashedName: challenge-214
---

# --description--

Ako vam je dat niz podataka domina, utvrdite da li formira validan lanac.

- Svaki element u nizu podataka predstavlja dominom i biće niz od dve brojeve od 1 do 6 (uključujući).
- Da bi lanac bio validan, drugi broj svake domine mora odgovarati prvom broju sledeće domine.
- Prvi broj prve domine i poslednji broj zadnje domine ne moraju se poklapati sa ničim.

# --hints--

`isValidDominoChain([[1, 3], [3, 6], [6, 5]])` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidDominoChain([[1, 3], [3, 6], [6, 5]]));
```

`isValidDominoChain([[6, 2], [3, 4], [4, 1]])` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidDominoChain([[6, 2], [3, 4], [4, 1]]));
```

`isValidDominoChain([[2, 5], [5, 6], [5, 1]])` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidDominoChain([[2, 5], [5, 6], [5, 1]]));
```

`isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]])` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]]));
```

`isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]])` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]]));
```

# --seed--

## --seed-contents--

```js
function isValidDominoChain(dominoes) {

  return dominoes;
}
```

# --solutions--

```js
function isValidDominoChain(dominoes) {
  for (let i = 0; i < dominoes.length - 1; i++) {
    if (dominoes[i][1] !== dominoes[i+1][0]) return false
  }

  return true;
}
```