id: 699c8e045ee7cb94ed2322db
title: "Vežba 220: Najveći broj"
challengeType: 28
dashedName: challenge-220
---

# --description--

Uzevši u obzir niz brojeva razdvojen različitim separatorima, vratite najveći broj.

- Pruženi niz će sadržati samo brojeve i separatore.
- Separatori mogu biti zarez (","), uzvičnik ("!"), znak pitanja ("?"), dvotačka (":") ili tačka-zarez(";").

# --hints--

`largestNumber("1,2")` bi trebalo da vrati `2`.

```js
assert.equal(largestNumber("1,2"), 2);
```

`largestNumber("4;15:60,26?52!0")` bi trebalo da vrati `60`.

```js
assert.equal(largestNumber("4;15:60,26?52!0"), 60);
```

`largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011")` bi trebalo da vrati `-402`.

```js
assert.equal(largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011"), -402);
```

`largestNumber("12;-50,99.9,49.1!-10.1?88?16")` bi trebalo da vrati `99.9`.

```js
assert.equal(largestNumber("12;-50,99.9,49.1!-10.1?88?16"), 99.9);
```

# --seed--

## --seed-contents--

```js
function largestNumber(str) {

  return str;
}
```

# --solutions--

```js
function largestNumber(str) {
  const numbers = str.split(/[,!?:;]/).map(Number);
  return Math.max(...numbers);
}
```