id: 699c8e045ee7cb94ed2322dd
title: "Vežba 222: Senke ekvinokija"
challengeType: 28
dashedName: challenge-222
---

# --description--

Danas je dan ekvinokcija, kada se sunce nalazi tačno na središnjoj liniji Zemlje i najviše je podne s šest sati. Kada dobijete određeno vreme, odredite senku koju proizvodi stub visine 4 stope.

- Vreme će biti niz slova u formatu 24 sata `"HH:MM"` (primer: `"15:00"` je 3 popodne).
- Biće vam dato vreme samo sa intervalima od 30 minuta.

Pravila:

- Sunce izlazi u 6 ujutro tačno "east", i zalazi u 6 uveče tačno "west".
- Senka je uvek naspram sunca.
- Dužina senke (u stopama) je broj sati od podne s šest sati, podignut na treći stepen.
- Nema senke pre nego što sunce izađe (pre 6 ujutro), posle nego što sunce zađe (nakon 6 uveče) ili u 12 sati.

Vratite:

- Ako postoji senka, vratite `"(dužina)ft (smer)"`. Primer: `"8ft west"`.
- Inače, vratite `"No shadow"`.

Na primer, ako dobijete `"10:00"`, vratite `"8ft west"` jer je 10 ujutro 2 sata od podne s šest sati, pa je $2^3$ = 8 stopa, a senka je na zapad jer je sunce istočno u 10 ujutro.

# --hints--

`getShadow("10:00")` treba da vrati `"8ft west"`.

```js
assert.equal(getShadow("10:00"), "8ft west");
```

`getShadow("15:00")` treba da vrati `"27ft east"`.

```js
assert.equal(getShadow("15:00"), "27ft east");
```

`getShadow("12:00")` treba da vrati `"No shadow"`.

```js
assert.equal(getShadow("12:00"), "No shadow");
```

`getShadow("17:30")` treba da vrati `"166.375ft east"`.

```js
assert.equal(getShadow("17:30"), "166.375ft east");
```

`getShadow("05:00")` treba da vrati `"No shadow"`.

```js
assert.equal(getShadow("05:00"), "No shadow");
```

`getShadow("06:00")` treba da vrati `"216ft west"`.

```js
assert.equal(getShadow("06:00"), "216ft west");
```

`getShadow("18:00")` treba da vrati `"No shadow"`.

```js
assert.equal(getShadow("18:00"), "No shadow");
```

`getShadow("07:30")` treba da vrati `"91.125ft west"`.

```js
assert.equal(getShadow("07:30"), "91.125ft west");
```

`getShadow("00:00")` treba da vrati `"No shadow"`.

```js
assert.equal(getShadow("00:00"), "No shadow");
```

# --seed--

## --seed-contents--

```js
function getShadow(time) {

  return time;
}
```

# --solutions--

```js
function getShadow(time) {
  const [hourStr, minuteStr] = time.split(":");
  const hours = Number(hourStr);
  const minutes = Number(minuteStr);
  const timeInHours = hours + minutes / 60;

  if (timeInHours < 6 || timeInHours >= 18 || timeInHours === 12) {
    return "No shadow";
  }

  const hoursFromNoon = Math.abs(12 - timeInHours);
  const length = Math.pow(hoursFromNoon, 3);
  const direction = timeInHours < 12 ? "west" : "east";

  return `${length}ft ${direction}`;
}
```