id: 699c8e045ee7cb94ed2322d9
title: "Zadatak 218: Savrševo deljenje"
challengeType: 28
dashedName: challenge-218
---

# --description--

Kada vam se daju dva cela broja, utvrdite da li je prvi broj savršeno deljiv sa drugim.

# --hints--

`isEvenlyDivisible(4, 2)` treba da vrati `true`.

```js
assert.isTrue(isEvenlyDivisible(4, 2));
```

`isEvenlyDivisible(7, 3)` treba da vrati `false`.

```js
assert.isFalse(isEvenlyDivisible(7, 3));
```

`isEvenlyDivisible(5, 10)` treba da vrati `false`.

```js
assert.isFalse(isEvenlyDivisible(5, 10));
```

`isEvenlyDivisible(48, 6)` treba da vrati `true`.

```js
assert.isTrue(isEvenlyDivisible(48, 6));
```

`isEvenlyDivisible(3186, 9)` treba da vrati `true`.

```js
assert.isTrue(isEvenlyDivisible(3186, 9));
```

`isEvenlyDivisible(4192, 11)` treba da vrati `false`.

```js
assert.isFalse(isEvenlyDivisible(4192, 11));
```

# --seed--

## --seed-contents--

```js
function isEvenlyDivisible(a, b) {

  return a;
}
```

# --solutions--

```js
function isEvenlyDivisible(a, b) {
  return a % b === 0;
}
```