id: 699c8e045ee7cb94ed2322d4
title: "Vežba 213: Konverter dužine reči"
challengeType: 28
dashedName: challenge-213
---

# --description--

Kada dobijete niz reči, vratite novi niz gde je svaka reč zamenjena svojom dužinom.

- Reči u datom nizu ostaju sa jednim razmakom između njih
- Sačuvajte te razmake u vraćenom nizu.

Na primer, ako dobijete `"hello world"`, vratite `"5 5"`.

# --hints--

`convertWords("hello world")` treba da vrati `"5 5"`.

```js
assert.equal(convertWords("hello world"), "5 5");
```

`convertWords("Thanks and happy coding")` treba da vrati `"6 3 5 6"`.

```js
assert.equal(convertWords("Thanks and happy coding"), "6 3 5 6");
```

`convertWords("The quick brown fox jumps over the lazy dog")` treba da vrati `"3 5 5 3 5 4 3 4 3"`.

```js
assert.equal(convertWords("The quick brown fox jumps over the lazy dog"), "3 5 5 3 5 4 3 4 3");
```

`convertWords("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl")` treba da vrati `"5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4"`.

```js
assert.equal(convertWords("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl"), "5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4");
```

# --seed--

## --seed-contents--

```js
function convertWords(str) {

  return str;
}
```

# --solutions--

```js
function convertWords(str) {
  return str
    .split(" ")
    .map(word => word.length)
    .join(" ");
}
```