id: 699c8e045ee7cb94ed2322da
title: "Vežba 219: Spomenici godišnjica venčanja"
challengeType: 28
dashedName: challenge-219
---

# --description--

Ako vam je dat ceo broj koji predstavlja broj godina koje su muž i žena zajedno u braku, vratite im najnoviji spomenik godišnjice prema ovoj tabeli:

| Godine braka | Spomenik  |
| ------------- | ------------ |
| 1             | `"Paper"`    |
| 5             | `"Wood"`     |
| 10            | `"Tin"`      |
| 25            | `"Silver"`   |
| 40            | `"Ruby"`     |
| 50            | `"Gold"`     |
| 60            | `"Diamond"`  |
| 70            | `"Platinum"` |

- Ako nisu dostigli prvi spomenik, vratite `"Newlyweds"`.

# --hints--

`getMilestone(0)` treba da vrati `"Newlyweds"`.

```js
assert.equal(getMilestone(0), "Newlyweds");
```

`getMilestone(1)` treba da vrati `"Paper"`.

```js
assert.equal(getMilestone(1), "Paper");
```

`getMilestone(8)` treba da vrati `"Wood"`.

```js
assert.equal(getMilestone(8), "Wood");
```

`getMilestone(10)` treba da vrati `"Tin"`.

```js
assert.equal(getMilestone(10), "Tin");
```

`getMilestone(26)` treba da vrati `"Silver"`.

```js
assert.equal(getMilestone(26), "Silver");
```

`getMilestone(45)` treba da vrati `"Ruby"`.

```js
assert.equal(getMilestone(45), "Ruby");
```

`getMilestone(50)` treba da vrati `"Gold"`.

```js
assert.equal(getMilestone(50), "Gold");
```

`getMilestone(64)` treba da vrati `"Diamond"`.

```js
assert.equal(getMilestone(64), "Diamond");
```

`getMilestone(71)` treba da vrati `"Platinum"`.

```js
assert.equal(getMilestone(71), "Platinum");
```

# --seed--

## --seed-contents--

```js
function getMilestone(years) {

  return years;
}
```

# --solutions--

```js
function getMilestone(years) {
  if (years < 1) return "Newlyweds";

  const milestones = [
    [1, "Paper"],
    [5, "Wood"],
    [10, "Tin"],
    [25, "Silver"],
    [40, "Ruby"],
    [50, "Gold"],
    [60, "Diamond"],
    [70, "Platinum"]
  ];

  for (let i = milestones.length - 1; i >= 0; i--) {
    if (years >= milestones[i][0]) {
      return milestones[i][1];
    }
  }
}
```