id: 69a890af247de743333bd4d0
title: "Vežba 227: Vreme odmora"
challengeType: 28
dashedName: challenge-227
---

# --description--

Uzimajući u obzir dva vremenska oznakivanja, prvo koje predstavlja vreme kada je korisnik završio test, a drugo koje predstavlja trenutno vreme, utvrdite da li korisnik može ponovo da položi test.

- Oba vremenska oznakivanja će biti data u formatu: `"YYYY-MM-DDTHH:MM:SS"`, na primer `"2026-03-25T14:00:00"`. Zapamtite da je vreme 24-satni sistem.
- Korisnik mora sačekati najmanje 48 sati pre ponovnog polaganja testa.

# --hints--

`canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` treba da vrati `true`.

```js
assert.isTrue(canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00"));
```

`canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` treba da vrati `false`.

```js
assert.isFalse(canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00"));
```

`canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` treba da vrati `true`.

```js
assert.isTrue(canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00"));
```

`canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59")` treba da vrati `false`.

```js
assert.isFalse(canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59"));
```

# --seed--

## --seed-contents--

```js
function canRetake(finishTime, currentTime) {

  return finishTime;
}
```

# --solutions--

```js
function canRetake(finishTime, currentTime) {
  function parse(ts) {
    const parts = ts.split(/[-T:]/).map(Number);
    const [y, m, d, h, min, s] = parts;
    return Date.UTC(y, m - 1, d, h, min, s);
  }

  const finishMs = parse(finishTime);
  const currentMs = parse(currentTime);

  const cooldown = 48 * 60 * 60 * 1000;

  return currentMs - finishMs >= cooldown;
}
```