id: 69a890af247de743333bd4d1
title: "Vežba 228: Filmska noć"
challengeType: 28
dashedName: challenge-228
---

# --description--

Ako dobijete niz slova koji pokazuje dan u sedmici, drugi niz slova koji pokazuje vreme prikazivanja filma i ukupan broj ulaznica, vratite ukupnu cenu ulaznica za film za to prikazivanje.

Dan koji je dat će biti jedan od:

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

Vreme prikazivanja će biti dato u formatu `"H:MMam"` ili `"H:MMpm"`. Na primer `"10:00am"` ili `"10:00pm"`.

Vratite ukupnu cenu u formatu `"$D.CC"` koristeći ova pravila:

- Vikendje (Petak - Nedelja): $12.00 po ulaznici.
- Radni dani (Ponedeljak - Četvrtak): $10.00 po ulaznici.
- Popodnevno prikazivanje (pre 11 ujutru): smanjenje od $2.00 po ulaznici (osim u utorak).
- Sreda: sve ulaznice su po $5.00.

# --hints--

`getMovieNightCost("Saturday", "10:00pm", 1)` bi trebalo da vrati `"$12.00"`.

```js
assert.equal(getMovieNightCost("Saturday", "10:00pm", 1), "$12.00");
```

`getMovieNightCost("Sunday", "10:00am", 1)` bi trebalo da vrati `"$10.00"`.

```js
assert.equal(getMovieNightCost("Sunday", "10:00am", 1), "$10.00");
```

`getMovieNightCost("Tuesday", "7:20pm", 2)` bi trebalo da vrati `"$10.00"`.

```js
assert.equal(getMovieNightCost("Tuesday", "7:20pm", 2), "$10.00");
```

`getMovieNightCost("Wednesday", "5:40pm", 3)` bi trebalo da vrati `"$30.00"`.

```js
assert.equal(getMovieNightCost("Wednesday", "5:40pm", 3), "$30.00");
```

`getMovieNightCost("Monday", "11:50am", 4)` bi trebalo da vrati `"$32.00"`.

```js
assert.equal(getMovieNightCost("Monday", "11:50am", 4), "$32.00");
```

`getMovieNightCost("Friday", "4:30pm", 5)` bi trebalo da vrati `"$50.00"`.

```js
assert.equal(getMovieNightCost("Friday", "4:30pm", 5), "$50.00");
```

`getMovieNightCost("Tuesday", "11:30am", 1)` bi trebalo da vrati `"$5.00"`.

```js
assert.equal(getMovieNightCost("Tuesday", "11:30am", 1), "$5.00");
```

# --seed--

## --seed-contents--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {

  return day;
}
```

# --solutions--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {
  let price;

  if (day === "Tuesday") {
    price = 5;
  } else {
    if (["Friday", "Saturday", "Sunday"].includes(day)) {
      price = 12;
    } else {
      price = 10;
    }

    const hour = parseInt(showtime.split(":")[0]);
    const period = showtime.slice(-2);

    let hour24 = hour;

    if (period === "pm" && hour !== 12) hour24 += 12;
    if (period === "am" && hour === 12) hour24 = 0;

    if (hour24 < 17) {
      price -= 2;
    }
  }

  const total = price * numberOfTickets;

  return `$${total.toFixed(2)}`;
}
```