id: 69a890af247de743333bd4cc
title: "Zadatak 223: QR Dekoder"
challengeType: 28
dashedName: challenge-223
---

# --description--

Ako vam je data matrica 6x6 (redovi podataka sekvence karaktera) koja predstavlja QR kod, vratite niz binarnih podataka unutar tog koda.

- QR kod može biti dat u bilo kom rotacionom uglu od 90 stepeni.
- Kod koji je u ispravnom orijentaciji ima grupu `1` veličine 2x2 (markeri orijentacije) u donjem levom, gornjem levom i gornjem desnom uglu.
- Tri markerja orijentacije od 2x2 nisu deo binarnih podataka.
- Binarni podaci se čitaju od leve ka desno, odozgo nadole (kao knjiga) kada je QR kod u ispravnoj orijentaciji.
- Kod će imati samo jednu validnu orijentaciju.

Na primer, ako dobijete:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "110000",
  "110001"
]
```

ili ako dobijete isti kod u različitoj orijentaciji:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "000011",
  "100011"
]
```

Vratite `"000000000000000000000001"`, sve binarne podatke osim tri markerja orijentacije od 2x2.

# --hints--

`decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"])` bi trebalo da vrati `"000000000000000000000001"`.

```js
assert.equal(decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"]), "000000000000000000000001");
```

`decodeQr(["100011", "000011", "000000", "000000", "110011", "111011"])` bi trebalo da vrati `"000000000000000000000001"`.

```js
assert.equal(decodeQr(["100011", "000011", "000000", "000000", "110011", "111011"]), "000000000000000000000001");
```

`decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"])` bi trebalo da vrati `"001101000011000000110100"`.

```js
assert.equal(decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"]), "001101000011000000110100");
```

`decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"])` bi trebalo da vrati `"010001000100010101010110"`.

```js
assert.equal(decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"]), "010001000100010101010110");
```

`decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"])` bi trebalo da vrati `"010000100100100101001110"`.

```js
assert.equal(decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"]), "010000100100100101001110");
```

# --seed--

## --seed-contents--

```js
function decodeQr(qrCode) {

  return qrCode;
}
```

# --solutions--

```js
function isOriented(code) {
  const markerIndices = [
    [0, 0],
    [0, 1],
    [0, 4],
    [0, 5],
    [1, 0],
    [1, 1],
    [1, 4],
    [1, 5],
    [4, 0],
    [4, 1],
    [5, 0],
    [5, 1]
  ];

  for(const [row, col] of markerIndices) {
    if (code[row][col] !== '1') return false;
  }

  return true;
}

function rotate(code) {
  const newCode = ['', '', '', '', '', ''];
  for (let i = 0; i < 6; i++) {
    for (let j = 5; j >= 0; j--) {
      newCode[i] += code[j][i];
    }
  }
  return newCode;
}

function getData(code) {
  let data = '';
  data += code[0].substring(2, 4)
  data += code[1].substring(2, 4)
  data += code[2]
  data += code[3]
  data += code[4].substring(2)
  data += code[5].substring(2)
  return data;
}

function decodeQr(qrCode) {
  let code = [...qrCode];

  while(!isOriented(code)) {
    code = rotate(code);
  }

  return getData(code);
}
```