id: 69a890af247de743333bd4ce
title: "Vežba 225: Nema ponavljajućih uzastopnih slova"
challengeType: 28
dashedName: challenge-225
---

# --description--

Za dati niz slova, utvrdite da li nema ponavljajućih slova.

- Niz slova nema ponavljanja ako ne sadrži isto slovo dva ili više puta uzastopno.

# --hints--

`hasNoRepeats("hi world")` bi trebalo da vrati `true`.

```js
assert.isTrue(hasNoRepeats("hi world"));
```

`hasNoRepeats("hello world")` bi trebalo da vrati `false`.

```js
assert.isFalse(hasNoRepeats("hello world"));
```

`hasNoRepeats("abcdefghijklmnopqrstuvwxyz")` bi trebalo da vrati `true`.

```js
assert.isTrue(hasNoRepeats("abcdefghijklmnopqrstuvwxyz"));
```

`hasNoRepeats("freeCodeCamp")` bi trebalo da vratiti `false`.

```js
assert.isFalse(hasNoRepeats("freeCodeCamp"));
```

`hasNoRepeats("The quick brown fox jumped over the lazy dog.")` bi trebalo da vrati `true`.

```js
assert.isTrue(hasNoRepeats("The quick brown fox jumped over the lazy dog."));
```

`hasNoRepeats("Mississippi")` bi trebalo da vrati `false`.

```js
assert.isFalse(hasNoRepeats("Mississippi"));
```

# --seed--

## --seed-contents--

```js
function hasNoRepeats(str) {

  return str;
}
```

# --solutions--

```js
function hasNoRepeats(str) {
  for (let i = 1; i < str.length; i++) {
    if (str[i] === str[i - 1]) {
      return false;
    }
  }
  return true;
}
```