id: 69a890af247de743333bd4cd
title: "Vežba 224: Detektor stepena prženja kafe"
challengeType: 28
dashedName: challenge-224
---

# --description--

Dato je niz znakova koji predstavljaju zrna korišćena za pravljenje šolje kafe; odredite stepen prženja te šolje.

- Dat niz će sadržati sledeće znake, svaki od kojih predstavlja određenu vrstu zrna:

  - Apostrof (`'`) je svetlo praženo zrno sa vrednošću od 1 poen za svako pojavljivanje.
  - Znak crte (`-`) je srednje praženo zrno sa vrednošću od 2 poena za svako pojavljivanje.
  - Tačka (`.`) je tamno praženo zrno sa vrednošću od 3 poena za svako pojavljivanje.

- Stepen prženja se određuje prosekom svih zrna.

Vratite:

- `"Light"` ako je prosečna vrednost ispod 1.75.
- `"Medium"` ako je prosečna vrednost između 1.75 i 2.5.
- `"Dark"` ako je prosečna vrednost veća od 2.5.

# --hints--

`detectRoast("''-''''''-'-''--''''")` should return `"Light"`.

```js
assert.equal(detectRoast("''-''''''-'-''--''''"), "Light");
```

`detectRoast(".'-''-''..'''.-.-''-")` should return `"Medium"`.

```js
assert.equal(detectRoast(".'-''-''..'''.-.-''-"), "Medium");
```

`detectRoast("--.''--'-''.--..-.--")` should return `"Medium"`.

```js
assert.equal(detectRoast("--.''--'-''.--..-.--"), "Medium");
```

`detectRoast("-...'-......-..-...-")` should return `"Dark"`.

```js
assert.equal(detectRoast("-...'-......-..-...-"), "Dark");
```

`detectRoast(".--.-..-......----.'")` should return `"Medium"`.

```js
assert.equal(detectRoast(".--.-..-......----.'"), "Medium");
```

`detectRoast("..-..-..-..-....-.-.")` should return `"Dark"`.

```js
assert.equal(detectRoast("..-..-..-..-....-.-."), "Dark");
```

`detectRoast("-'-''''''..-'.''-'.'")` should return `"Light"`.

```js
assert.equal(detectRoast("-'-''''''..-'.''-'.'"), "Light");
```

# --seed--

## --seed-contents--

```js
function detectRoast(beans) {

  return beans;
}
```

# --solutions--

```js
function detectRoast(beans) {
  let total = 0;

  for (const bean of beans) {
    if (bean === "'") total += 1;
    else if (bean === "-") total += 2;
    else if (bean === ".") total += 3;
  }

  const avg = total / beans.length;

  if (avg < 1.75) return "Light";
  if (avg <= 2.5) return "Medium";
  return "Dark";
}
```