id: 69a890af247de743333bd4d2
title: "Vežba 229: Skraćivanje niza karaktera 2"
challengeType: 28
dashedName: challenge-229
---

# --description--

Ako vam je dat niz karaktera, vratite skraćeni novi niz kako bi ukupna širina karaktera ne prešla 50 jedinica.

Svaki karakter ima posebnu širinu:

| Karakter | Širina |
| - | - |
| `"ilI"` | 1 |
| `"fjrt"` | 2 |
| `"abcdeghkmnopqrstuvwxyzJL"` | 3 |
| `"ABCDEFGHKMNOPQRSTUVWXYZ"` | 4 |

Tabela iznad uključuje sva velika i mala slova. Pored toga:

- Razmak (`" "`) ima širinu od 2
- Tačka (`"."`) ima širinu od 1

- Ako je dat niz 50 jedinica ili manje, vratite ga takav kakav jeste; inače
- Skratite niz i dodajte tri tačke na kraj (`"..."`) tako da ukupna širina, uključujući tri tačke, bude što bliže 50 jedinica bez prelaska.

# --hints--

`truncateText("The quick brown fox")` bi trebalo da vrati `"The quick brown f..."`.

```js
assert.equal(truncateText("The quick brown fox"), "The quick brown f...");
```

`truncateText("The silky smooth sloth")` bi trebalo da vrati `"The silky smooth s..."`.

```js
assert.equal(truncateText("The silky smooth sloth"), "The silky smooth s...");
```

`truncateText("THE LOUD BRIGHT BIRD")` bi trebalo da vrati `"THE LOUD BRIG..."`.

```js
assert.equal(truncateText("THE LOUD BRIGHT BIRD"), "THE LOUD BRIG...");
```

`truncateText("The fast striped zebra")` bi trebalo da vrati `"The fast striped z..."`.

```js
assert.equal(truncateText("The fast striped zebra"), "The fast striped z...");
```

`truncateText("The big black bear")` bi trebalo da vrati `"The big black bear"`.

```js
assert.equal(truncateText("The big black bear"), "The big black bear");
```

# --seed--

## --seed-contents--

```js
function truncateText(str) {

  return str;
}
```

# --solutions--

```js
function truncateText(str) {
  const MAX_WIDTH = 50;
  const ELLIPSIS = "...";
  const ELLIPSIS_WIDTH = 3;
  const TRUNCATE_LIMIT = MAX_WIDTH - ELLIPSIS_WIDTH;

  const charGroups = {
    'ilI.': 1,
    'fjrt ': 2,
    'abcdeghkmnopqrstuvwxyzJL': 3,
    'ABCDEFGHKMNOPQRSTUVWXYZ': 4
  };

  function getCharWidth(char) {
    for (const key in charGroups) {
      if (key.includes(char)) return charGroups[key];
    }
    return 3;
  }

  function stringWidth(str) {
    let totalWidth = 0;
    for (const char of str) {
      totalWidth += getCharWidth(char);
    }
    return totalWidth;
  }

  if (stringWidth(str) <= MAX_WIDTH) return str;

  let result = "";
  let totalWidth = 0;

  for (const char of str) {
    const charWidth = getCharWidth(char);

    if (totalWidth + charWidth > TRUNCATE_LIMIT) break;

    result += char;
    totalWidth += charWidth;
  }

  return result + ELLIPSIS;
}
```