id: 69b1028d6e265413d0198a2a
title: "Vežba 231: Proverite ISBN-10"
challengeType: 28
dashedName: challenge-231
---

# --description--

Datoj niz znakova, utvrdite da li je validan ISBN-10.

ISBN-10 sadrži razmakke (hyphen) (`"-"`) i još 10 znakova. Nakon uklanjanja razmakaka (`"-"`):

- Prvih 9 znakova moraju biti brojevi, i
- Posljednji znak može biti broj ili slovo "X", koje predstavlja broj 10.

Za validaciju:

- Pomnožite svaki broj (ili vrijednost) sa svojom pozicijom (pomnožite prvi broj sa 1, drugi sa 2, itd.).
- Saberi sve rezultate.
- Ako je zbir djeljiv sa 11, onda je validan.

# --hints--

`isValidIsbn10("0-306-40615-2")` should return `true`.

```js
assert.isTrue(isValidIsbn10("0-306-40615-2"));
```

`isValidIsbn10("0-306-40615-1")` should return `false`.

```js
assert.isFalse(isValidIsbn10("0-306-40615-1"));
```

`isValidIsbn10("0-8044-2957-X")` should return `true`.

```js
assert.isTrue(isValidIsbn10("0-8044-2957-X"));
```

`isValidIsbn10("X-306-40615-2")` should return `false`.

```js
assert.isFalse(isValidIsbn10("X-306-40615-2"));
```

`isValidIsbn10("0-6822-2589-4")` should return `true`.

```js
assert.isTrue(isValidIsbn10("0-6822-2589-4"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn10(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn10(str) {
  const isbn = str.replace(/-/g, "");

  if (isbn.length !== 10) return false;

  let sum = 0;

  for (let i = 0; i < 9; i++) {
    const digit = Number(isbn[i]);
    if (Number.isNaN(digit)) return false;

    sum += digit * (i + 1);
  }

  const last = isbn[9];

  if (last === "X") {
    sum += 10 * 10;
  } else {
    const digit = Number(last);
    if (Number.isNaN(digit)) return false;

    sum += digit * 10;
  }

  return sum % 11 === 0;
}
```