---
id: 69b1028d6e265413d0198a2e
title: "Vežba 235: Fibonacci sa velikim slovima"
challengeType: 28
dashedName: challenge-235
---

# --description--

Kada dobijete niz slova, vratite novi niz gde je svako slovo napisano velikim slovom ako je njegov indeks Fibonacci broj, i malim slovom u suprotnom slučaju.

Fibonacci niz je sekvenca brojeva gde je svaki broj zbir prethodnih dva. Prvih deset brojeva u nizu su `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

- Prvo slovo je na indeksu `0`.
- Ako je indeks znaka nije slovo, ostavi ga nepromenjenog.

# --hints--

`capitalizeFibonacci("hello world")` bi trebalo da vrati `"HELLo woRld"`.

```js
assert.equal(capitalizeFibonacci("hello world"), "HELLo woRld");
```

`capitalizeFibonacci("HELLO WORLD")` bi trebalo da vrati `"HELLo woRld"`.

```js
assert.equal(capitalizeFibonacci("HELLO WORLD"), "HELLo woRld");
```

`capitalizeFibonacci("hello, world!")` bi trebalo da vrati `"HELLo, wOrld!"`.

```js
assert.equal(capitalizeFibonacci("hello, world!"), "HELLo, wOrld!");
```

`capitalizeFibonacci("The quick brown fox jumped over the lazy dog.")` bi trebalo da vrati `"THE qUicK broWn fox jUmped over thE lazy dog."`.

```js
assert.equal(capitalizeFibonacci("The quick brown fox jumped over the lazy dog."), "THE qUicK broWn fox jUmped over thE lazy dog.");
```

`capitalizeFibonacci("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Proin pulvinar ex nibh, vel ullamcorper ligula egestas quis. Integer tincidunt fringilla accumsan. Integer et metus placerat, gravida felis at, pellentesque nisl.")` bi trebalo da vrati `"LOREm ipSum dOlor sit amet, consecTetur adipiscing elit. proin pulvinar ex nibh, vel ullaMcorper ligula egestas quis. integer tincidunt fringillA accumsan. integer et metus placerat, gravida felis at, pellentesque nisl."`.

```js
assert.equal(capitalizeFibonacci("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Proin pulvinar ex nibh, vel ullamcorper ligula egestas quis. Integer tincidunt fringilla accumsan. Integer et metus placerat, gravida felis at, pellentesque nisl."), "LOREm ipSum dOlor sit amet, consecTetur adipiscing elit. proin pulvinar ex nibh, vel ullaMcorper ligula egestas quis. integer tincidunt fringillA accumsan. integer et metus placerat, gravida felis at, pellentesque nisl.");
```

# --seed--

## --seed-contents--

```js
function capitalizeFibonacci(str) {

  return str;
}
```

# --solutions--

```js
function capitalizeFibonacci(str) {
  const fibs = new Set([0, 1]);
  let a = 0, b = 1;
  while (b < str.length) {
    [a, b] = [b, a + b];
    fibs.add(b);
  }

  return str
    .split("")
    .map((char, i) => {
      if (!/[a-zA-Z]/.test(char)) return char;
      return fibs.has(i) ? char.toUpperCase() : char.toLowerCase();
    })
    .join("");
}
```