id: 69b1028d6e265413d0198a2c
title: "Vežba 233: Alarm za buđenje"
challengeType: 28
dashedName: challenge-233
---

# --description--

Dati su vam niz znakova koji predstavlja vreme kada ste postavili alarm i niz znakova koji predstavlja vreme kada ste se probudili. Odredite da li ste ustali rano, na vreme ili kasno.

- Oba vremena će biti data u formatu 24-satnog hrambata `"HH:MM"`.

Vratite:

- `"early"` ako ste ustali pre alarma.
- `"on time"` ako ste ustali tačno na vreme alarma, ili unutar vremenskog prozora od 10 minuta kašnjenja nakon alarma.
- `"late"` ako ste ustali više od 10 minuta nakon alarma.

Oba vremena su u istom danu.

# --hints--

`alarmCheck("07:00", "06:45")` inapaswa kurudisha `"early"`.

```js
assert.equal(alarmCheck("07:00", "06:45"), "early");
```

`alarmCheck("06:30", "06:30")` inapaswa kurudisha `"on time"`.

```js
assert.equal(alarmCheck("06:30", "06:30"), "on time");
```

`alarmCheck("08:10", "08:15")` inapaswa kurudisha `"on time"`.

```js
assert.equal(alarmCheck("08:10", "08:15"), "on time");
```

`alarmCheck("09:30", "09:45")` inapaswa kurudisha `"late"`.

```js
assert.equal(alarmCheck("09:30", "09:45"), "late");
```

`alarmCheck("08:15", "08:25")` inapaswa kurudisha `"on time"`.

```js
assert.equal(alarmCheck("08:15", "08:25"), "on time");
```

`alarmCheck("05:45", "05:56")` inapaswa kurudisha `"late"`.

```js
assert.equal(alarmCheck("05:45", "05:56"), "late");
```

`alarmCheck("04:30", "04:00")` inapaswa kurudisha `"early"`.

```js
assert.equal(alarmCheck("04:30", "04:00"), "early");
```

# --seed--

## --seed-contents--

```js
function alarmCheck(alarmTime, wakeTime) {

  return alarmTime;
}
```

# --solutions--

```js
function alarmCheck(alarmTime, wakeTime) {
  const [alarmH, alarmM] = alarmTime.split(":").map(Number);
  const [wakeH, wakeM] = wakeTime.split(":").map(Number);

  const alarmMinutes = alarmH * 60 + alarmM;
  const wakeMinutes = wakeH * 60 + wakeM;
  const diff = wakeMinutes - alarmMinutes;

  if (diff < 0) return "early";
  if (diff <= 10) return "on time";
  return "late";
}
```