id: 69b58ce40693f140c84c8559
title: "Vežba 240: Palindromska slova"
challengeType: 28
dashedName: challenge-240
---

# --description--

Dato je niz znakova, utvrdite da li je palindrom i vratite srednji karakter (ako je dužina neparna) ili dva srednja karaktera (ako je dužina parna).

- Palindrom je niz znakova koji se čita isto napred i nazad.
- Ako nije palindrom, vratite `"none"`.

# --hints--

`palindromeLocator("racecar")` treba da vrati `"e"`.

```js
assert.equal(palindromeLocator("racecar"), "e");
```

`palindromeLocator("level")` treba da vrati `"v"`.

```js
assert.equal(palindromeLocator("level"), "v");
```

`palindromeLocator("freecodecamp")` treba da vrati `"none"`.

```js
assert.equal(palindromeLocator("freecodecamp"), "none");
```

`palindromeLocator("noon")` treba da vrati `"oo"`.

```js
assert.equal(palindromeLocator("noon"), "oo");
```

`palindromeLocator("11100111")` treba da vrati `"00"`.

```js
assert.equal(palindromeLocator("11100111"), "00");
```

# --seed--

## --seed-contents--

```js
function palindromeLocator(str) {

  return str;
}
```

# --solutions--

```js
function palindromeLocator(str) {
  if (str !== str.split("").reverse().join("")) return "none";

  const mid = Math.floor(str.length / 2);
  return str.length % 2 === 1 ? str[mid] : str[mid - 1] + str[mid];
}
```