id: 69b559d2903b9e4afe9075f9
title: "Vežba 239: Koji je danas dan?"
challengeType: 28
dashedName: challenge-239
---

# --description--

Kada dobijete Unix vremenski stamp od milisekundi, vratite dan u sedmici.

Validni dani za vraćanje su:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

Obavezno uzmite u obzir vremonske zone.

# --hints--

`getDayOfWeek(1775492249000)` bi trebalo da vrati `"Monday"`.

```js
assert.equal(getDayOfWeek(1775492249000), "Monday");
```

`getDayOfWeek(1766246400000)` bi trebalo da vrati `"Saturday"`.

```js
assert.equal(getDayOfWeek(1766246400000), "Saturday");
```

`getDayOfWeek(33791256000000)` bi trebalo da vrati `"Tuesday"`.

```js
assert.equal(getDayOfWeek(33791256000000), "Tuesday");
```

`getDayOfWeek(1773576000000)` bi trebalo da vrati `"Sunday"`.

```js
assert.equal(getDayOfWeek(1773576000000), "Sunday");
```

`getDayOfWeek(0)` bi trebalo da vrati `"Thursday"`.

```js
assert.equal(getDayOfWeek(0), "Thursday");
```

# --seed--

## --seed-contents--

```js
function getDayOfWeek(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function getDayOfWeek(timestamp) {
  return new Date(timestamp).toLocaleDateString("en-US", { weekday: "long", timeZone: "UTC" });
}
```