id: 69b559d2903b9e4afe9075f8
title: "Vežba 238: Izbegavanje rotacije broja"
challengeType: 28
dashedName: challenge-238
---

# --description--

Ako je dat pozitivan ceo broj, proverite da li taj broj ili bilo koja rotacija tog broja može biti podeljiv sa brojem njegovih cifara.

Rotacija znači pomeranje prve cifre na kraj. Na primer, nakon jedne rotacije, 123 postaje 231.

- Prvo proverite rotaciju `0` (dat broj).
- Dani brojevi neće sadržati nulu.
- Vratite broj prve rotacije ako postoji, ili `"none"` ako ne postoji.

# --hints--

`getRotation(123)` bi trebalo da vrati `0`.

```js
assert.equal(getRotation(123), 0);
```

`getRotation(13579)` bi trebalo da vrati `3`.

```js
assert.equal(getRotation(13579), 3);
```

`getRotation(24681)` bi trebalo da vrati `"none"`.

```js
assert.equal(getRotation(24681), "none");
```

`getRotation(84138789345)` bi trebalo da vrati `6`.

```js
assert.equal(getRotation(84138789345), 6);
```

# --seed--

## --seed-contents--

```js
function getRotation(n) {

  return n;
}
```

# --solutions--

```js
function getRotation(n) {
  const str = String(n);
  const digitCount = str.length;
  let current = str;

  for (let i = 0; i < digitCount; i++) {
    if (Number(current) % digitCount === 0) return i;
    current = current.slice(1) + current[0];
  }

  return "none";
}
```