id: 69b559d2903b9e4afe9075f7
title: "Vežba 237: Provera matematičke jednačine"
challengeType: 28
dashedName: challenge-237
---

# --description--

Dat je niz znakova koji predstavlja matematički izraz; utvrdite da li je on tačan.

- Leva strana može sadržavati do tri celobrojnih broja i operatore `+`, `-`, `*`, i `/`.
- Matematički izraz će biti dat u formatu: `"broj operator broj = broj"` (sa dva ili tri broja na levoj strani). Primeri su: `"2 + 2 = 4"` ili `"2 + 3 - 1 = 4"`.
- Desna strana će uvek biti jedan celobrojni broj.

Pratite standardni red operacija: množenje i deljenje se procenjuju pre sabiranja i oduzimanja, od leve ka desnoj.

# --hints--

`isValidEquation("2 + 2 = 4")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidEquation("2 + 2 = 4"));
```

`isValidEquation("2 + 3 - 1 = 4")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidEquation("2 + 3 - 1 = 4"));
```

`isValidEquation("8 / 2 = 4")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidEquation("8 / 2 = 4"));
```

`isValidEquation("10 * 5 = 50")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidEquation("10 * 5 = 50"));
```

`isValidEquation("2 - 2 = 0")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidEquation("2 - 2 = 0"));
```

`isValidEquation("2 + 9 / 3 = 5")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidEquation("2 + 9 / 3 = 5"));
```

`isValidEquation("20 - 2 * 3 = 14")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidEquation("20 - 2 * 3 = 14"));
```

`isValidEquation("2 + 5 = 6")` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidEquation("2 + 5 = 6"));
```

`isValidEquation("10 - 2 * 3 = 24")` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidEquation("10 - 2 * 3 = 24"));
```

`isValidEquation("3 + 9 / 3 = 4")` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidEquation("3 + 9 / 3 = 4"));
```

# --seed--

## --seed-contents--

```js
function isValidEquation(equation) {

  return equation;
}
```

# --solutions--

```js
function isValidEquation(equation) {
  const [left, right] = equation.split(" = ");
  const tokens = left.split(" ");

  for (let i = 1; i < tokens.length - 1; i += 2) {
    const operator = tokens[i];
    if (operator === "*" || operator === "/") {
      const leftOperand = Number(tokens[i - 1]);
      const rightOperand = Number(tokens[i + 1]);
      const computed = operator === "*" ? leftOperand * rightOperand : leftOperand / rightOperand;
      tokens.splice(i - 1, 3, String(computed));
      i -= 2;
    }
  }

  let result = Number(tokens[0]);
  for (let i = 1; i < tokens.length - 1; i += 2) {
    const operator = tokens[i];
    const operand = Number(tokens[i + 1]);
    result = operator === "+" ? result + operand : result - operand;
  }

  return result === Number(right);
}
```