id: 69b5b2be76ec8135a7fbe974
title: "Zadatak 245: Spirala matrice"
challengeType: 28
dashedName: challenge-245
---

# --description--

Ako vam je data 2D matrica, vratite ravni niz podataka sa svim vrednostima u smeru po godini.

Niz podataka treba da sadrži vrednosti od gornje leve tačke, zatim se kreće desno duž gornjeg reda, pa dole duž desne kolone, zatim levo duž donjeg reda, i na kraju gore duž levog redisa. Ponovite ovaj proces ka unutra za sve preostale slojeve.

Na primer, ako dobijete:

```json
[
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9]
]
```

Vratite `[1, 2, 3, 6, 9, 8, 7, 4, 5]`.

# --hints--

`spiralMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]])` bi trebalo da vrati `[1, 2, 3, 6, 9, 8, 7, 4, 5]`.

```js
assert.deepEqual(spiralMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]]), [1, 2, 3, 6, 9, 8, 7, 4, 5]);
```

`spiralMatrix([["a", "b", "c", "d"], ["l", "m", "n", "e"], ["k", "p", "o", "f"], ["j", "i", "h", "g"]])` bi trebalo da vrati `["a", "b", "c", "d", "e", "f", "g", "h", "i", "j", "k", "l", "m", "n", "o", "p"]`.

```js
assert.deepEqual(spiralMatrix([["a", "b", "c", "d"], ["l", "m", "n", "e"], ["k", "p", "o", "f"], ["j", "i", "h", "g"]]), ["a", "b", "c", "d", "e", "f", "g", "h", "i", "j", "k", "l", "m", "n", "o", "p"]);
```

`spiralMatrix([[true, false, false], [false, true, true], [false, true, false], [true, true, false]])` bi trebalo da vrati `[true, false, false, true, false, false, true, true, false, false, true, true]`.

```js
assert.deepEqual(spiralMatrix([[true, false, false], [false, true, true], [false, true, false], [true, true, false]]), [true, false, false, true, false, false, true, true, false, false, true, true]);
```

`spiralMatrix([[25, 24, 23, 22, 21], [10, 9, 8, 7, 20], [11, 2, 1, 6, 19], [12, 3, 4, 5, 18], [13, 14, 15, 16, 17]])` bi trebalo da vrati `[25, 24, 23, 22, 21, 20, 19, 18, 17, 16, 15, 14, 13, 12, 11, 10, 9, 8, 7, 6, 5, 4, 3, 2, 1]`.

```js
assert.deepEqual(spiralMatrix([[25, 24, 23, 22, 21], [10, 9, 8, 7, 20], [11, 2, 1, 6, 19], [12, 3, 4, 5, 18], [13, 14, 15, 16, 17]]), [25, 24, 23, 22, 21, 20, 19, 18, 17, 16, 15, 14, 13, 12, 11, 10, 9, 8, 7, 6, 5, 4, 3, 2, 1]);
```

# --seed--

## --seed-contents--

```js
function spiralMatrix(matrix) {

    return matrix
}
```

# --solutions--

```js
function spiralMatrix(matrix) {
  const result = [];
  let top = 0, bottom = matrix.length - 1;
  let left = 0, right = matrix[0].length - 1;

  while (top <= bottom && left <= right) {
    for (let i = left; i <= right; i++) result.push(matrix[top][i]);
    top++;

    for (let i = top; i <= bottom; i++) result.push(matrix[i][right]);
    right--;

    if (top <= bottom) {
      for (let i = right; i >= left; i--) result.push(matrix[bottom][i]);
      bottom--;
    }

    if (left <= right) {
      for (let i = bottom; i >= top; i--) result.push(matrix[i][left]);
      left++;
    }
  }

  return result;
}
```