id: 69b58ce40693f140c84c855b
title: "Vežba 242: Sledeći Bingo broj"
challengeType: 28
dashedName: challenge-242
---

# --description--

Kada dobijete Bingo broj, vratite sledeći po redu Bingo broj.

Bingo broj je slovo praćeno brojem u njegovom opsegu prema ovoj tabeli:

| Slovo | Opseg brojeva |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

Na primer, ako dobijete `"B10"`, vratite `"B11"`, sledeći Bingo broj. Ako dobijete poslednji Bingo broj, vratite `"B1"`.

# --hints--

`getNextBingoNumber("B10")` bi trebalo da vrati `"B11"`.

```js
assert.equal(getNextBingoNumber("B10"), "B11");
```

`getNextBingoNumber("N33")` bi trebalo da vrati `"N34"`.

```js
assert.equal(getNextBingoNumber("N33"), "N34");
```

`getNextBingoNumber("I30")` bi trebalo da vrati `"N31"`.

```js
assert.equal(getNextBingoNumber("I30"), "N31");
```

`getNextBingoNumber("G60")` bi trebalo da vrati `"O61"`.

```js
assert.equal(getNextBingoNumber("G60"), "O61");
```

`getNextBingoNumber("O75")` bi trebalo da vrati `"B1"`.

```js
assert.equal(getNextBingoNumber("O75"), "B1");
```

# --seed--

## --seed-contents--

```js
function getNextBingoNumber(n) {

  return n;
}
```

# --solutions--

```js
function getNextBingoNumber(n) {
  const num = parseInt(n.slice(1));
  const next = num === 75 ? 1 : num + 1;

  if (next <= 15) return "B" + next;
  if (next <= 30) return "I" + next;
  if (next <= 45) return "N" + next;
  if (next <= 60) return "G" + next;
  return "O" + next;
}
```