id: 69bc6cb30c1d112a2e110a06
title: "Vežba 249: Matematika sekvence slova"
challengeType: 28
dashedName: challenge-249
---

# --description--

Ako vam je dat niz brojeva i drugih znakova, izvršite proračun na brojevima na osnovu broja nebrojnih znakova između njih.

- Ako je broj znakova koji razdvajaju dva broja parno, koristite sabiranje.
- Ako je neparno, koristite oduzimanje.
- Uzastopni brojevi čine jedan broj.
- Operacije se izvode od levi ka desni.
- Ignorišite vodeće i završne nebrojne znakove.

Na primer, ako dobijete `"3ab10c8"`, vratite `5`. Sabirate 3 i 10 da biste dobili 13 jer između njih ima parno broja znakova. Zatim oduzmite 8 od 13 jer između rezultata i 8 postoji neparno broja znakova.

# --hints--

`doMath("3ab10c8")` bi trebalo da vrati `5`.

```js
assert.equal(doMath("3ab10c8"), 5);
```

`doMath("6MINUS4")` bi trebalo da vrati `2`.

```js
assert.equal(doMath("6MINUS4"), 2);
```

`doMath("9plus3")` bi trebalo da vrati `12`.

```js
assert.equal(doMath("9plus3"), 12);
```

`doMath("5fkwo#10i#%.<>15P=@20!#B/25")` bi trebalo da vrati `15`.

```js
assert.equal(doMath("5fkwo#10i#%.<>15P=@20!#B/25"), 15);
```

`doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt")` bi trebalo da vrati `67`.

```js
assert.equal(doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt"), 67);
```

# --seed--

## --seed-contents--

```js
function doMath(str) {

  return str;
}
```

# --solutions--

```js
function doMath(str) {
  const tokens = str.match(/(\d+)|(\D+)/g);
  let result = null;
  let pendingOp = null;

  for (const token of tokens) {
    if (/^\d+$/.test(token)) {
      if (result === null) {
        result = parseInt(token);
      } else {
        result = pendingOp === 'add' ? result + parseInt(token) : result - parseInt(token);
      }
    } else {
      pendingOp = token.length % 2 === 0 ? 'add' : 'subtract';
    }
  }

  return result;
}
```