id: 69bc6cb30c1d112a2e110a08
title: "Zadatak 251: Pronađite zbirke podataka niza"
challengeType: 28
dashedName: challenge-251
---

# --description--

Kada se podaju niz brojeva i ciljna brojka, vratite prvi podmnožak sa dve ili više brojeve čija je suma jednaka ciljnoj brojci.

- "Prvi" podmnožak je onaj čiji elementi imaju najmanju moguće veličinu, uzimajući u obzir najraniji element prvo.
- Svaki broj u nizu podataka može se koristiti samo jednom.
- Ako ne postoji validan podmnožak, vratite `"Sum not found"`.

Vratite odgovarajuće brojeve kao niz podataka u redosledu kako su vidljivi u originalnom nizu podataka.

# --hints--

`findSum([1, 3, 5, 7], 6)` bi trebalo da vrati `[1, 5]`.

```js
assert.deepEqual(findSum([1, 3, 5, 7], 6), [1, 5]);
```

`findSum([1, 2, 3, 4, 5], 5)` bi trebalo da vrati `[1, 4]`.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 5), [1, 4]);
```

`findSum([1, 2, 3, 4, 5], 6)` bi trebalo da vrati `[1, 2, 3]`.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5], 6), [1, 2, 3]);
```

`findSum([-1, -2, 3, 4], 1)` bi trebalo da vrati `[-1, -2, 4]`.

```js
assert.deepEqual(findSum([-1, -2, 3, 4], 1), [-1, -2, 4]);
```

`findSum([3, 1, 4, 1, 5, 9, 2, 6], 10)` bi trebalo da vrati `[3, 1, 4, 2]`.

```js
assert.deepEqual(findSum([3, 1, 4, 1, 5, 9, 2, 6], 10), [3, 1, 4, 2]);
```

`findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20)` bi trebalo da vrati `[1, 2, 3, 5, 9]`.

```js
assert.deepEqual(findSum([1, 2, 3, 4, 5, 6, 7, 8, 9], 20), [1, 2, 3, 5, 9]);
```

`findSum([7, 9, 4, 2, 5], 10)` bi trebalo da vrati `"Sum not found"`.

```js
assert.equal(findSum([7, 9, 4, 2, 5], 10), "Sum not found");
```

# --seed--

## --seed-contents--

```js
function findSum(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findSum(arr, target) {
  function search(start, remaining, current) {
    if (remaining === 0 && current.length >= 2) return current;
    if (start >= arr.length) return null;

    for (let i = start; i < arr.length; i++) {
      const result = search(i + 1, remaining - arr[i], [...current, arr[i]]);
      if (result) return result;
    }

    return null;
  }

  return search(0, target, []) ?? "Sum not found";
}
```