id: 69bc6cb30c1d112a2e110a04
title: "Vežba 247: Poslednja slova"
challengeType: 28
dashedName: challenge-247
---

# --description--

Dati niz slova, vratite slovo koje se pojavljuje najkasnije u alfabetu.

- Ako više od dva slova izgledaju kao poslednja slova u alfabetu, vratite prvo pronađeno.
- Ignorišite sve nealfabetičke znakove.

# --hints--

`getLastLetter("world")` treba da vrati `"w"`.

```js
assert.equal(getLastLetter("world"), "w");
```

`getLastLetter("Hello World")` treba da vrati `"W"`.

```js
assert.equal(getLastLetter("Hello World"), "W");
```

`getLastLetter("The quick brown fox jumped over the lazy dog.")` treba da vrati `"z"`.

```js
assert.equal(getLastLetter("The quick brown fox jumped over the lazy dog."), "z");
```

`getLastLetter("HeLl0")` treba da vrati `"L"`.

```js
assert.equal(getLastLetter("HeLl0"), "L");
```

`getLastLetter("!#$ er@R asd fT.,> 2t0e9")` treba da vrati `"T"`.

```js
assert.equal(getLastLetter("!#$ er@R asd fT.,> 2t0e9"), "T");
```

# --seed--

## --seed-contents--

```js
function getLastLetter(str) {

  return str;
}
```

# --solutions--

```js
function getLastLetter(str) {
  const letters = str.split('').filter(char => /[a-zA-Z]/.test(char));
  return letters.reduce((last, char) => char.toLowerCase() > last.toLowerCase() ? char : last);
}
```