id: 69bc6cb30c1d112a2e110a09
title: "Vežba 252: Jedinstveni penjač stepeništa"
challengeType: 28
dashedName: challenge-252
---

# --description--

Uzimajući u obzir broj stepeništa, vratite koliko različitih načina osoba može popeti uz korake od 1 ili 2 u jednom trenutku.

# --hints--

`getUniqueClimbs(4)` bi trebalo da vrati `5`.

```js
assert.equal(getUniqueClimbs(4), 5);
```

`getUniqueClimbs(5)` bi trebalo da vrati `8`.

```js
assert.equal(getUniqueClimbs(5), 8);
```

`getUniqueClimbs(10)` bi trebalo da vrati `89`.

```js
assert.equal(getUniqueClimbs(10), 89);
```

`getUniqueClimbs(18)` bi trebalo da vrati `4181`.

```js
assert.equal(getUniqueClimbs(18), 4181);
```

`getUniqueClimbs(29)` bi trebalo da vrati `832040`.

```js
assert.equal(getUniqueClimbs(29), 832040);
```

`getUniqueClimbs(50)` bi trebalo da vrati `20365011074`.

```js
assert.equal(getUniqueClimbs(50), 20365011074);
```

# --seed--

## --seed-contents--

```js
function getUniqueClimbs(steps) {

  return steps;
}
```

# --solutions--

```js
function getUniqueClimbs(steps) {
  if (steps <= 0) return 0;
  if (steps === 1) return 1;
  if (steps === 2) return 2;

  let prev2 = 1, prev1 = 2;
  for (let i = 3; i <= steps; i++) {
    [prev2, prev1] = [prev1, prev2 + prev1];
  }
  return prev1;
}
```