id: 69c5f3d787b1725d5f00c8b8
title: "Vežba 254: Reči sa neparnim brojem slova"
challengeType: 28
dashedName: challenge-254
---

# --description--

Ako dobijete niz reči, vratite reči koje imaju neparno trajanje.

- Reči u datom nizu će biti razdvojene samo jednim razmakom.
- Vratite reči koje su razdvojene samo jednim razmakom.

# --hints--

`getOddWords("This is a super good test")` bi trebalo da vrati `"a super"`.

```js
assert.equal(getOddWords("This is a super good test"), "a super");
```

`getOddWords("one two three four")` bi trebalo da vrati `"one two three"`.

```js
assert.equal(getOddWords("one two three four"), "one two three");
```

`getOddWords("banana split sundae with rainbow sprinkles on top")` bi trebalo da vrati `"split rainbow sprinkles top"`.

```js
assert.equal(getOddWords("banana split sundae with rainbow sprinkles on top"), "split rainbow sprinkles top");
```

`getOddWords("The quick brown fox jumped over the lazy river")` bi trebalo da vrati `"The quick brown fox the river"`.

```js
assert.equal(getOddWords("The quick brown fox jumped over the lazy river"), "The quick brown fox the river");
```

# --seed--

## --seed-contents--

```js
function getOddWords(str) {

  return str;
}
```

# --solutions--

```js
function getOddWords(str) {
  return str
    .split(" ")
    .filter(word => word.length % 2 === 1)
    .join(" ");
}
```