id: 69c5f3d787b1725d5f00c8ba
title: "Vežba 256: Najkraći vremenski pravac"
challengeType: 28
dashedName: challenge-256
---

# --description--

Uzimajući u obzir dva vremena, utvrdite da li je brže stići od prvog do drugog krećući se napred ili nazad.

- Vremena su data u formatu 24 sata (`"HH:MM"`)
- Sat je cikličan (23:59 postaje 00:00 kada idete napred, a 00:00 postaje 23:59 kada idete nazad)

Vratite:

- `"forward"` ako je kretanje napred kraće
- `"backward"` ako je kretanje nazad kraće
- `"equal"` ako oba pravca traju isto vreme

# --hints--

`getDirection("10:00", "12:00")` treba da vrati `"forward"`.

```js
assert.equal(getDirection("10:00", "12:00"), "forward");
```

`getDirection("11:00", "05:00")` treba da vrati `"backward"`.

```js
assert.equal(getDirection("11:00", "05:00"), "backward");
```

`getDirection("00:00", "12:00")` treba da vrati `"equal"`.

```js
assert.equal(getDirection("00:00", "12:00"), "equal");
```

`getDirection("15:45", "01:10")` treba da vrati `"forward"`.

```js
assert.equal(getDirection("15:45", "01:10"), "forward");
```

`getDirection("03:30", "19:50")` treba da vrati `"backward"`.

```js
assert.equal(getDirection("03:30", "19:50"), "backward");
```

`getDirection("06:30", "18:30")` treba da vrati `"equal"`.

```js
assert.equal(getDirection("06:30", "18:30"), "equal");
```

# --seed--

## --seed-contents--

```js
function getDirection(time1, time2) {

  return time1;
}
```

# --solutions--

```js
function getDirection(time1, time2) {
  function toMinutes(t) {
    const [h, m] = t.split(":").map(Number);
    return h * 60 + m;
  }

  const start = toMinutes(time1);
  const end = toMinutes(time2);

  const forward = (end - start + 1440) % 1440;
  const backward = (start - end + 1440) % 1440;

  if (forward < backward) return "forward";
  if (backward < forward) return "backward";
  return "equal";
}
```