id: 69c5f3d787b1725d5f00c8bd
title: "Zadatak 259: Eksplozija FizzBuzz-a"
challengeType: 28
dashedName: challenge-259
---

# --description--

Kada dobijete celo brojeve, vratite broj koraka potrebnih za transformisanje stringa `"fizzbuzz"` u sekvencu slova koja sadrži najmanje određeni broj `"z"` koristeći sledeća pravila:

- Počinjete sa stringom `"fizzbuzz"`.
- Svaki korak, koristite standardna FizzBuzz pravila na osnovu pozicije slova u sekvenci (prvo slovo `"f"` je na poziciji 1).
  - Ako se pozicija deli sa 3, zamenite to slovo sa `"fizz"`
  - Ako se deli sa 5, zamenite to slovo sa `"buzz"`
  - Ako se deli i sa 3 i sa 5, zamenite to slovo sa `"fizzbuzz"`

Dakle, nakon koraka 1, `"fizzbuzz"` postaje `"fifizzzbuzzfizzzz"`, koji sadrži devet `"z"`.

# --hints--

`explodeFizzbuzz(9)` treba da vrati `1`.

```js
assert.equal(explodeFizzbuzz(9), 1);
```

`explodeFizzbuzz(15)` treba da vrati `2`.

```js
assert.equal(explodeFizzbuzz(15), 2);
```

`explodeFizzbuzz(51)` treba da vrati `3`.

```js
assert.equal(explodeFizzbuzz(51), 3);
```

`explodeFizzbuzz(52)` treba da vrati `4`.

```js
assert.equal(explodeFizzbuzz(52), 4);
```

`explodeFizzbuzz(359)` treba da vrati `5`.

```js
assert.equal(explodeFizzbuzz(359), 5);
```

`explodeFizzbuzz(789)` treba da vrati `6`.

```js
assert.equal(explodeFizzbuzz(789), 6);
```

`explodeFizzbuzz(54482)` treba da vrati `11`.

```js
assert.equal(explodeFizzbuzz(54482), 11);
```

`explodeFizzbuzz(1000000)` treba da vrati `14`.

```js
assert.equal(explodeFizzbuzz(1000000), 14);
```

# --seed--

## --seed-contents--

```js
function explodeFizzbuzz(targetZCount) {

  return targetZCount;
}
```

# --solutions--

```js
function explodeFizzbuzz(targetZCount) {
  let str = "fizzbuzz";
  let steps = 0;

  while ((str.match(/z/g) || []).length < targetZCount) {
    let next = "";
    for (let i = 0; i < str.length; i++) {
      const pos = i + 1;
      if (pos % 15 === 0) next += "fizzbuzz";
      else if (pos % 3 === 0) next += "fizz";
      else if (pos % 5 === 0) next += "buzz";
      else next += str[i];
    }
    str = next;
    steps++;
  }

  return steps;
}
```