id: 69c6ff713a52713463aa7927
title: "Vežba 261: Reči brojeva"
challengeType: 28
dashedName: challenge-261
---

# --description--

Ako dobijete kompletan broj od 0 do 99, vratite njegov naziv na engleskom jeziku.

- `0` vraća `"zero"`.
- Brojevi 1-19 imaju svoje jedinstvene nazive (`"one"`, `"two"`, ..., `"ten"`, `"eleven"`, ..., `"eighteen"`, `"nineteen"`).
- Višegodišnji brojevi (multiples of 10) od 20 do 90 imaju svoje nazive (`"twenty"`, `"thirty"`, ..., `"eighty"`, `"ninety"`).
- Brojevi 21-99 koji nisu višestruki sa 10 pišu se kao dva reči spojena crticom. Primeri su `"forty-two"` i `"fifty-three"`.

# --hints--

`getNumberWords(0)` bi trebalo da vrati `"zero"`.

```js
assert.equal(getNumberWords(0), "zero");
```

`getNumberWords(10)` bi trebalo da vrati `"ten"`.

```js
assert.equal(getNumberWords(10), "ten");
```

`getNumberWords(19)` bi trebalo da vrati `"nineteen"`.

```js
assert.equal(getNumberWords(19), "nineteen");
```

`getNumberWords(30)` bi trebalo da vrati `"thirty"`.

```js
assert.equal(getNumberWords(30), "thirty");
```

`getNumberWords(53)` bi trebalo da vrati `"fifty-three"`.

```js
assert.equal(getNumberWords(53), "fifty-three");
```

`getNumberWords(7)` bi trebalo da vrati `"seven"`.

```js
assert.equal(getNumberWords(7), "seven");
```

`getNumberWords(12)` bi trebalo da vrati `"twelve"`.

```js
assert.equal(getNumberWords(12), "twelve");
```

`getNumberWords(60)` bi trebalo da vrati `"sixty"`.

```js
assert.equal(getNumberWords(60), "sixty");
```

`getNumberWords(67)` bi trebalo da vrati `"sixty-seven"`.

```js
assert.equal(getNumberWords(67), "sixty-seven");
```

`getNumberWords(98)` bi trebalo da vrati `"ninety-eight"`.

```js
assert.equal(getNumberWords(98), "ninety-eight");
```

# --seed--

## --seed-contents--

```js
function getNumberWords(n) {

  return n;
}
```

# --solutions--

```js
function getNumberWords(n) {
  const ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
                 "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
                 "seventeen","eighteen","nineteen"];
  const tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"];

  if (n < 20) return ones[n];
  if (n % 10 === 0) return tens[Math.floor(n / 10)];
  return tens[Math.floor(n / 10)] + "-" + ones[n % 10];
}
```