id: 69c6ff713a52713463aa7929
title: "Izazov 263: Binarni vokabular"
challengeType: 28
dashedName: challenge-263
---

# --description--

Ako vam se da slovo, odredite da li njegov binarni prikaz od 8 bita može biti pronađen u sledećoj mreži, horizontalno ili vertikalno u bilo kom smeru:

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

Na primer, slovo `"A"` ima binarni prikaz `01000001`, koji se vidi u prvom redu od leve ka desne.

# --hints--

`isInCrossword("I")` bi trebalo da vrati `true`.

```js
assert.isTrue(isInCrossword("I"));
```

`isInCrossword("D")` bi trebalo da vrati `true`.

```js
assert.isTrue(isInCrossword("D"));
```

`isInCrossword("0")` bi trebalo da vrati `true`.

```js
assert.isTrue(isInCrossword("0"));
```

`isInCrossword("u")` bi trebalo da vrati `true`.

```js
assert.isTrue(isInCrossword("u"));
```

`isInCrossword("Y")` bi trebalo da vrati `false`.

```js
assert.isFalse(isInCrossword("Y"));
```

`isInCrossword("p")` bi trebalo da vrati `false`.

```js
assert.isFalse(isInCrossword("p"));
```

`isInCrossword("1")` bi trebalo da vrati `false`.

```js
assert.isFalse(isInCrossword("1"));
```

`isInCrossword("Q")` bi trebalo da vrati `false`.

```js
assert.isFalse(isInCrossword("Q"));
```

# --seed--

## --seed-contents--

```js
function isInCrossword(char) {

  return char;
}
```

# --solutions--

```js
function isInCrossword(char) {
  const grid = [
    "01000001",
    "01101111",
    "01000100",
    "01100101",
    "01010010",
    "01010100",
    "01101000",
    "10101110",
  ];

  const target = char.charCodeAt(0).toString(2).padStart(8, '0');
  const cols = Array.from({length: 8}, (_, c) => grid.map(r => r[c]).join(''));

  return [...grid, ...cols].some(seq => seq.includes(target) || seq.split('').reverse().join('').includes(target));
}
```