id: 69c6ff713a52713463aa7928
title: "Vežba 262: Analizator pretrigarnog upita URL-a"
challengeType: 28
dashedName: challenge-262
---

# --description--

Ako URL sadrži niz znakova pretrigarnog upita, analizirajte taj niz znakova u objekat (ili rječnik) parova ključ-vrijednost.

- Niz znakova pretrigarnog upita počinje nakon `"?"`,
- svaki parametar je razdvojen sa `"&"`,
- svaki par ključ/vrijednost je razdvojen sa `"="`

Na primjer, za `"https://example.com/search?name=Alice&age=30"`, vratite:

```json
{
  "name": "Alice",
  "age": "30"
}
```

Sve vrijednosti treba da se vrate kao niz znakova (string).

# --hints--

`parseUrlQuery("https://example.com/search?name=Alice&age=30")` bi trebalo da vrati `{"name": "Alice", "age": "30"}`

```js
assert.deepEqual(parseUrlQuery("https://example.com/search?name=Alice&age=30"), {"name": "Alice", "age": "30"});
```

`parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python")` bi trebalo da vrati `{"skill": "programming", "language": "python"}`

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python"), {"skill": "programming", "language": "python"});
```

`parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2")` bi trebalo da vrati `{"category": "books", "sort": "asc", "page": "2"}`

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2"), {"category": "books", "sort": "asc", "page": "2"});
```

`parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now")` bi trebalo da vrati `{"redirect": "freecodecamp.org/learn", "when": "now"}`

```js
assert.deepEqual(parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now"), {"redirect": "freecodecamp.org/learn", "when": "now"});
```

# --seed--

## --seed-contents--

```js
function parseUrlQuery(url) {

  return url;
}
```

# --solutions--

```js
function parseUrlQuery(url) {
  const queryString = url.split('?')[1];
  return Object.fromEntries(
    queryString.split('&').map(param => param.split('='))
  );
}
```