id: 69c6ff713a52713463aa7926
title: "Vežba 260: Bodovanje reči"
challengeType: 28
dashedName: challenge-260
---

# --description--

Kada dobijete reč, vratite njene bodove koristeći standardnu tabelu vrednosti slova:

| Slovo | Vrednost |
| --- | --- |
|   A |   1 |
|   B |   2 |
| ... | ... |
|   Z |  26 |

- Velika i mala slova imaju istu vrednost.

# --hints--

`getWordScore("hi")` bi trebalo da vrati `17`.

```js
assert.equal(getWordScore("hi"), 17);
```

`getWordScore("hello")` bi trebalo da vrati `52`.

```js
assert.equal(getWordScore("hello"), 52);
```

`getWordScore("hippopotamus")` bi trebalo da vrati `169`.

```js
assert.equal(getWordScore("hippopotamus"), 169);
```

`getWordScore("freeCodeCamp")` bi trebalo da vrati `94`.

```js
assert.equal(getWordScore("freeCodeCamp"), 94);
```

# --seed--

## --seed-contents--

```js
function getWordScore(word) {

  return word;
}
```

# --solutions--

```js
function getWordScore(word) {
  return word.toLowerCase().split('').reduce((sum, char) => {
    return sum + (char.charCodeAt(0) - 96);
  }, 0);
}
```