id: 69cfca90e8a0a6d4d6871c51
title: "Vežba 267: Konverter Parsek"
challengeType: 28
dashedName: challenge-267
---

# --description--

U dalekoj galaksiji, parsek se koristi za merenje vremena i udaljenosti. Kada dobijete celoban broj parsek, vratite njegov ekvivalent u vreme ili udaljenost.

- Ako je celobrojni broj neparan, predstavlja vreme. Ako je paran, predstavlja udaljenost.

Koristite ove konverzije:

| Parsecs | Vreme/Udaljenost |
| - | - |
| 1 | 2 sata |
| 2 | 6 svetlosnih godina |

Vratite konvertovanu vrednost kao celobrojni broj.

# --hints--

`convertParsecs(1)` treba da vrati `2`.

```js
assert.equal(convertParsecs(1), 2);
```

`convertParsecs(2)` treba da vrati `6`.

```js
assert.equal(convertParsecs(2), 6);
```

`convertParsecs(31)` treba da vrati `62`.

```js
assert.equal(convertParsecs(31), 62);
```

`convertParsecs(88)` treba da vrati `264`.

```js
assert.equal(convertParsecs(88), 264);
```

`convertParsecs(17)` treba da vrati `34`.

```js
assert.equal(convertParsecs(17), 34);
```

`convertParsecs(14)` treba da vrati `42`.

```js
assert.equal(convertParsecs(14), 42);
```

# --seed--

## --seed-contents--

```js
function convertParsecs(parsecs) {

  return parsecs;
}
```

# --solutions--

```js
function convertParsecs(parsecs) {
  if (parsecs % 2 !== 0) {
    return parsecs * 2;
  } else {
    return parsecs * 3;
  }
}
```