id: 69cfca90e8a0a6d4d6871c52
title: "Vežba 268: Narcisistički broj"
challengeType: 28
dashedName: challenge-268
---

# --description--

Ako je dat pozitivan cijeli broj, utvrdite da li je narcisistički broj.

- Broj je narcisistički ako je zbir svih njegovih cifara podignut na stepen jednak ukupnom broju cifra, jednak samom broju.

Na primjer, 153 ima 3 cifre, i $1^3 + 5^3 + 3^3 = 153$, pa je narcisistički.

# --hints--

`isNarcissistic(153)` bi trebalo da vrati `true`.

```js
assert.isTrue(isNarcissistic(153));
```

`isNarcissistic(154)` bi trebalo da vrati `false`.

```js
assert.isFalse(isNarcissistic(154));
```

`isNarcissistic(371)` bi trebalo da vrati `true`.

```js
assert.isTrue(isNarcissistic(371));
```

`isNarcissistic(512)` bi trebalo da vrati `false`.

```js
assert.isFalse(isNarcissistic(512));
```

`isNarcissistic(9)` bi trebalo da vrati `true`.

```js
assert.isTrue(isNarcissistic(9));
```

`isNarcissistic(11)` bi trebalo da vrati `false`.

```js
assert.isFalse(isNarcissistic(11));
```

`isNarcissistic(9474)` bi trebalo da vrati `true`.

```js
assert.isTrue(isNarcissistic(9474));
```

`isNarcissistic(6549)` bi trebalo da vrati `false`.

```js
assert.isFalse(isNarcissistic(6549));
```

# --seed--

## --seed-contents--

```js
function isNarcissistic(n) {

  return n;
}
```

# --solutions--

```js
function isNarcissistic(n) {
  const digits = String(n).split('');
  const power = digits.length;
  const sum = digits.reduce((acc, d) => acc + Math.pow(Number(d), power), 0);
  return sum === n;
}
```