id: 69cfca90e8a0a6d4d6871c4f
title: "Vežba 265: Najdublje ugnježđene zagrade"
challengeType: 28
dashedName: challenge-265
---

# --description--

Ako vam je dat niz znakova koji sadrži uparnjive zagrade, vratite sadržaj unutar najdublje ugnježdene zagrade.

- Zagrade mogu biti bilo koje od tri tipa: `()`, `[]`, i `{}`.
- Ulaz će imati samo jednu grupu koja je duboko ugnježđena.

Na primer, ako dobijete `"(hello (world))"`, vratite `"world"`.

# --hints--

`getDeepestBrackets("(hello (world))")` bi trebalo da vrati `"world"`.

```js
assert.equal(getDeepestBrackets("(hello (world))"), "world");
```

`getDeepestBrackets("[outer [inner] outer]")` bi trebalo da vrati `"inner"`.

```js
assert.equal(getDeepestBrackets("[outer [inner] outer]"), "inner");
```

`getDeepestBrackets("{a{b}c{d{e}f}g}")` bi trebalo da vrati `"e"`.

```js
assert.equal(getDeepestBrackets("{a{b}c{d{e}f}g}"), "e");
```

`getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]")` bi trebalo da vrati `"fox"`.

```js
assert.equal(getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]"), "fox");
```

`getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p")` bi trebalo da vrati `"C"`.

```js
assert.equal(getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p"), "C");
```

# --seed--

## --seed-contents--

```js
function getDeepestBrackets(str) {

  return str;
}
```

# --solutions--

```js
function getDeepestBrackets(str) {
  let maxDepth = 0, depth = 0, content = '';
  let start = 0;

  for (let i = 0; i < str.length; i++) {
    if ('([{'.includes(str[i])) {
      depth++;
      if (depth > maxDepth) {
        maxDepth = depth;
        start = i + 1;
      }
    } else if (')]}'.includes(str[i])) {
      if (depth === maxDepth) {
        content = str.slice(start, i);
      }
      depth--;
    }
  }
  return content;
}
```