id: 69d03e549613fbdbe21d11e0
title: "Vežba 271: Podsetnik za lek"
challengeType: 28
dashedName: challenge-271
---

# --description--

Uzimajući u obzir niz podataka o lijekovima i string koji predstavlja trenutno vrijeme, vratite sljedeći lek koji je potrebno uzeti i preostalo vrijeme do toga.

- Svaki lek je u formatu `[name, lastTaken]`, gdje je `name` naziv lijeka, a `lastTaken` vrijeme kada je zadnji put uzet.
- Sve vremena su data u formatu `"HH:MM"` (24 sata).

Koristite sljedeći raspored lijekova:

| Name | Raspored |
| - | - |
| Deployxitrin | 08:00, 16:00 |
| Debuggamanizole | 07:00, 13:00, 21:00 |
| Mergeflictamine | svakih 4 sata |

Vratite string u formatu `"{name} in Hh Mm"`. Na primjer, `"Debuggamanizole in 2h 0m"` ili `"Deployxitrin in 1h 5m"`.

# --hints--

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "10:00"]], "11:00")` treba da vrati `"Debuggamanizole in 2h 0m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "10:00"]], "11:00"), "Debuggamanizole in 2h 0m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "14:55")` treba da vrati `"Deployxitrin in 1h 5m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "14:55"), "Deployxitrin in 1h 5m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "17:15")` treba da vrati `"Mergeflictamine in 0h 45m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "13:00"], ["Mergeflictamine", "14:00"]], "17:15"), "Mergeflictamine in 0h 45m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "09:00"]], "12:59")` treba da vrati `"Debuggamanizole in 0h 1m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "09:00"]], "12:59"), "Debuggamanizole in 0h 1m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "21:00"], ["Mergeflictamine", "03:00"]], "06:55")` treba da vrati `"Debuggamanizole in 0h 5m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "21:00"], ["Mergeflictamine", "03:00"]], "06:55"), "Debuggamanizole in 0h 5m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00")` treba da vrati `"Mergeflictamine in 3h 30m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00"), "Mergeflictamine in 3h 30m");
```

`medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00")` treba da vrati `"Mergeflictamine in 3h 30m"`.

```js
assert.equal(medicationReminder([["Deployxitrin", "08:00"], ["Debuggamanizole", "07:00"], ["Mergeflictamine", "07:30"]], "08:00"), "Mergeflictamine in 3h 30m");
```

# --seed--

## --seed-contents--

```js
function medicationReminder(medications, currentTime) {

  return medications;
}
```

# --solutions--

```js
function medicationReminder(medications, currentTime) {
  const schedule = {
    Deployxitrin: { type: 'fixed', times: ['08:00', '16:00'] },
    Debuggamanizole: { type: 'fixed', times: ['07:00', '13:00', '21:00'] },
    Mergeflictamine: { type: 'interval', hours: 4 }
  };

  const toMinutes = t => {
    const [h, m] = t.split(':').map(Number);
    return h * 60 + m;
  };

  const current = toMinutes(currentTime);
  let earliest = Infinity;
  let nextMed = '';

  for (const [name, lastTaken] of medications) {
    const med = schedule[name];
    let nextDue;

    if (med.type === 'fixed') {
      const futureTimes = med.times
        .map(toMinutes)
        .filter(t => t > current);
      if (futureTimes.length === 0) continue;
      nextDue = Math.min(...futureTimes);
    } else {
      nextDue = toMinutes(lastTaken) + med.hours * 60;
    }

    const diff = nextDue - current;
    if (diff > 0 && diff < earliest) {
      earliest = diff;
      nextMed = name;
    }
  }

  const h = Math.floor(earliest / 60);
  const m = earliest % 60;
  return `${nextMed} in ${h}h ${m}m`;
}
```