id: 69e2383af7832c8032603b8f
title: "Vežba 273: Proverite važeći ISBN-13 broj"
challengeType: 28
dashedName: challenge-273
---

# --description--

Ako vam je dat niz znakova, utvrdite da li je važeći ISBN-13 broj.

Važeći ISBN-13:

- Sadrži samo cifre i crtice (-)
- Ima tačno 13 cifara nakon uklanjanja crtica
- Prolazi sledeće provere:
  1. Pomnožite svaku cifru sa 1 ili 3, uzastopno (prvu cifru pomnožiti sa 1, drugu sa 3, treću sa 1 i tako dalje).
  2. Zbir rezultata mora biti deljiv sa 10 bez ostatka.

# --hints--

`isValidIsbn13("9780306406157")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidIsbn13("9780306406157"));
```

`isValidIsbn13("97803064061570")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidIsbn13("97803064061570"));
```

`isValidIsbn13("978-0-13-595705-9")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidIsbn13("978-0-13-595705-9"));
```

`isValidIsbn13("978-030-64061A-4")` inapaswa kurudisha `false`.

```js
assert.isFalse(isValidIsbn13("978-030-64061A-4"));
```

`isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9")` inapaswa kurudisha `true`.

```js
assert.isTrue(isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn13(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn13(str) {
  const digits = str.replace(/-/g, '');
  if (!/^\d{13}$/.test(digits)) return false;
  const sum = digits.split('').reduce((acc, d, i) => acc + d * (i % 2 === 0 ? 1 : 3), 0);
  return sum % 10 === 0;
}
```