id: 69e2383af7832c8032603b94
title: "Vežba 278: Analizator kafe narudžbine"
challengeType: 28
dashedName: challenge-278
---

# --description--

Dato je niz niz znakova sa kafe narudžbinom, identifikujte sve stavke sa menija i vratite dobro formatiranu narudžbinu.

Koristite sledeće stavke sa menija i cene:

| Stavka | Cena | 
| - | - |
| `"cold brew"` | $4.50 |
| `"oat latte"` | $5.00 |
| `"cappuccino"` | $4.75 |
| `"espresso"` | $3.00 |
| `"vanilla syrup"` | $0.75 |
| `"caramel drizzle"` | $0.60 |
| `"extra shot"` | $0.50 |
| `"oat milk"` | $0.75 |
| `"cream"` | $0.75 |

Vratite niz znakova sa poklapanim stavkama spojene sa `" + "`, praćen dve dvotačke i razmak (`": "`), i ukupnom cenom.

Na primer, ako primite `"I'd like an oat latte with vanilla syrup and an extra shot please."`, vratite `"oat latte + vanilla syrup + extra shot: $6.25"`

Stavke moraju biti prikazane u redosledu kako su na meniju, a ukupna cena mora uvek imati dve decimalne začišlja.

# --hints--

`formatCoffeeOrder("I'd like an oat latte with vanilla syrup and an extra shot please.")` bi trebalo da vrati `"oat latte + vanilla syrup + extra shot: $6.25"`.

```js
assert.equal(formatCoffeeOrder("I'd like an oat latte with vanilla syrup and an extra shot please."), "oat latte + vanilla syrup + extra shot: $6.25");
```

`formatCoffeeOrder("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk.")` bi trebalo da vrati `"cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85"`.

```js
assert.equal(formatCoffeeOrder("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk."), "cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85");
```

`formatCoffeeOrder("Can I get a cold brew with some cream and an extra shot.")` bi trebalo da vrati `"cold brew + extra shot + cream: $5.75"`.

```js
assert.equal(formatCoffeeOrder("Can I get a cold brew with some cream and an extra shot."), "cold brew + extra shot + cream: $5.75");
```

`formatCoffeeOrder("Just an espresso please.")` bi trebalo da vrati `"espresso: $3.00"`.

```js
assert.equal(formatCoffeeOrder("Just an espresso please."), "espresso: $3.00");
```

`formatCoffeeOrder("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle.")` bi trebalo da vrati `"oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60"`.

```js
assert.equal(formatCoffeeOrder("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle."), "oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60");
```

# --seed--

## --seed-contents--

```js
function formatCoffeeOrder(order) {

  return order;
}
```

# --solutions--

```js
function formatCoffeeOrder(order) {
  const menu = [
    { name: "cold brew", price: 4.50 },
    { name: "oat latte", price: 5.00 },
    { name: "cappuccino", price: 4.75 },
    { name: "espresso", price: 3.00 },
    { name: "vanilla syrup", price: 0.75 },
    { name: "caramel drizzle", price: 0.60 },
    { name: "extra shot", price: 0.50 },
    { name: "oat milk", price: 0.75 },
    { name: "cream", price: 0.75 },
  ];

  const found = menu.filter(item => order.toLowerCase().includes(item.name));
  const total = found.reduce((sum, item) => sum + item.price, 0);
  return `${found.map(item => item.name).join(" + ")}: $${total.toFixed(2)}`;
}
```