id: 69e2383af7832c8032603b8e
title: "Vežba 272: Transponovana matrica"
challengeType: 28
dashedName: challenge-272
---

# --description--

Ako je matrica (niz podataka redova podataka), vrati verziju koja je promenila redove i kolone.

Da biste transponovali matricu, razmenite redove i kolone. Primer: vrednost na poziciji `[0, 1]` treba da se premesti na poziciju `[1, 0]`.

Na primer, ako je dato:

```json
[
  [1, 2, 3],
  [4, 5, 6]
]
```

Vratite:

```json
[
  [1, 4],
  [2, 5],
  [3, 6]
]
```

# --hints--

`transpose([[1, 2, 3], [4, 5, 6]])` bi trebalo da vrati `[[1, 4], [2, 5], [3, 6]]`.

```js
assert.deepEqual(transpose([[1, 2, 3], [4, 5, 6]]), [[1, 4], [2, 5], [3, 6]]);
```

`transpose([[1, 2], [3, 4], [5, 6]])` bi trebalo da vrati `[[1, 3, 5], [2, 4, 6]]`.

```js
assert.deepEqual(transpose([[1, 2], [3, 4], [5, 6]]), [[1, 3, 5], [2, 4, 6]]);
```

`transpose([[1, 2], [3, 4], [5, 6], [7, 8]])` bi trebalo da vrati `[[1, 3, 5, 7], [2, 4, 6, 8]]`.

```js
assert.deepEqual(transpose([[1, 2], [3, 4], [5, 6], [7, 8]]), [[1, 3, 5, 7], [2, 4, 6, 8]]);
```

`transpose([["a", "b", "c"], ["d", "e", "f"], ["g", "h", "i"], ["j", "k", "l"]])` bi trebalo da vrati `[["a", "d", "g", "j"], ["b", "e", "h", "k"], ["c", "f", "i", "l"]]`.

```js
assert.deepEqual(transpose([["a", "b", "c"], ["d", "e", "f"], ["g", "h", "i"], ["j", "k", "l"]]), [["a", "d", "g", "j"], ["b", "e", "h", "k"], ["c", "f", "i", "l"]]);
```

`transpose([[true, false, true, false], [false, true, false, true], [true, true, false, false], [false, false, true, true], [true, false, false, true]])` bi trebalo da vrati `[[true, false, true, false, true], [false, true, true, false, false], [true, false, false, true, false], [false, true, false, true, true]]`.

```js
assert.deepEqual(transpose([[true, false, true, false], [false, true, false, true], [true, true, false, false], [false, false, true, true], [true, false, false, true]]), [[true, false, true, false, true], [false, true, true, false, false], [true, false, false, true, false], [false, true, false, true, true]]);
```

# --seed--

## --seed-contents--

```js
function transpose(matrix) {

  return matrix;
}
```

# --solutions--

```js
function transpose(matrix) {

  return matrix[0].map((_, i) => matrix.map(row => row[i]));
}
```