id: 69e2383af7832c8032603b92
title: "Vežba 276: Element koji krši pravilo"
challengeType: 28
dashedName: challenge-276
---

# --description--

Dat niz celih brojeva je sortiran po rastućem redosledu, osim jednog elementa koji je na pogrešnom mestu. Vratite indeks tog elementa.

- Ako postoji više od jednog elementa koji se mogu smatrati van mesta mesta, vratite indeks prvog takvog elementa.

# --hints--

`findOffender([1, 6, 2, 3, 4, 5])` bi trebalo da vrati `1`.

```js
assert.equal(findOffender([1, 6, 2, 3, 4, 5]), 1);
```

`findOffender([1, 2, 3, 5, 4, 5])` bi trebalo da vrati `3`.

```js
assert.equal(findOffender([1, 2, 3, 5, 4, 5]), 3);
```

`findOffender([2, 1])` bi trebalo da vrati `0`.

```js
assert.equal(findOffender([2, 1]), 0);
```

`findOffender([2, 4, 1, 6, 8])` bi trebalo da vrati `2`.

```js
assert.equal(findOffender([2, 4, 1, 6, 8]), 2);
```

`findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91])` bi trebalo da vrati `6`.

```js
assert.equal(findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91]), 6);
```

# --seed--

## --seed-contents--

```js
function findOffender(arr) {

  return arr;
}
```

# --solutions--

```js
function findOffender(arr) {
  for (let i = 0; i < arr.length; i++) {
    const without = [...arr.slice(0, i), ...arr.slice(i + 1)];
    if (without.every((n, j) => j === 0 || without[j - 1] <= n)) return i;
  }
}
```