id: 69f35a5bb823ed620fcb7cb9
title: "Vežba 280: Datum Mongo ID"
challengeType: 28
dashedName: challenge-280
---

# --description--

Kada dobijete niz znakova MongoDB ID, vratite njegovo vreme kreiranja kao ISO 8601 niz znakova.

- MongoDB ID je 24-znak niz heksa. Prvih osam znakova predstavljaju Unix vremenski žig (u sekundama) predstavljen kao broj u bazi 16.

Na primer, `"6a094b50bcf86cd799439011"` ima vremenski žig `"6a094b50"` u heksadecimalnom sistemu, što je `1778994000` u decimalnom sistemu, predstavljajući vreme kreiranja `"2026-05-17T05:00:00.000Z"`.

# --hints--

`mongoIdToDate("6a094b50bcf86cd799439011")` bi trebalo da vrati `"2026-05-17T05:00:00.000Z"`.

```js
assert.equal(mongoIdToDate("6a094b50bcf86cd799439011"), "2026-05-17T05:00:00.000Z");
```

`mongoIdToDate("695344eb1f4a4c1123042128")` bi trebalo da vrati `"2025-12-30T03:20:11.000Z"`.

```js
assert.equal(mongoIdToDate("695344eb1f4a4c1123042128"), "2025-12-30T03:20:11.000Z");
```

`mongoIdToDate("386da62df34123ac54617e56")` bi trebalo da vrati `"2000-01-01T07:01:01.000Z"`.

```js
assert.equal(mongoIdToDate("386da62df34123ac54617e56"), "2000-01-01T07:01:01.000Z");
```

`mongoIdToDate("69f571c3d7711807afd3dd55")` bi trebalo da vrati `"2026-05-02T03:38:43.000Z"`.

```js
assert.equal(mongoIdToDate("69f571c3d7711807afd3dd55"), "2026-05-02T03:38:43.000Z");
```

`mongoIdToDate("68adce01c0e1144d0a90295a")` bi trebalo da vrati `"2025-08-26T15:08:49.000Z"`.

```js
assert.equal(mongoIdToDate("68adce01c0e1144d0a90295a"), "2025-08-26T15:08:49.000Z");
```

# --seed--

## --seed-contents--

```js
function mongoIdToDate(id) {

  return id;
}
```

# --solutions--

```js
function mongoIdToDate(id) {
  const timestamp = parseInt(id.slice(0, 8), 16);
  return new Date(timestamp * 1000).toISOString();
}
```