id: 69f35a5bb823ed620fcb7cbe
title: "Vežba 285: Vreme sastanka"
challengeType: 28
dashedName: challenge-285
---

# --description--

Ako vam je dat 3D niz podataka koji predstavlja prozor dostupnosti za više ljudi, vratite najranije vreme kada svi imaju sat slobodno. Ako takvo vreme ne postoji, vratite `"None"`.

- Dostupnost svake osobe je niz celobrojnih parova `[start, end]` u 24 sata. Na primer, `[10, 12]` znači da je osoba slobodna od sata 10 do 12. Vreme početka je između 0-23, a vreme završetka je između 1-24.

Na primer, ako dobijete:

```json
[
  [[10, 12], [15, 16]], // osoba 1
  [[11, 14], [15, 16]]  // osoba 2
]
```

Vratite `11`, početak njihovog prvog zajedničkog slobodnog sata.

# --hints--

`getMeetingTime([[[10, 12], [15, 16]], [[11, 14], [15, 16]]])` bi trebalo da vrati `11`.

```js
assert.equal(getMeetingTime([[[10, 12], [15, 16]], [[11, 14], [15, 16]]]), 11);
```

`getMeetingTime([[[9, 10], [12, 15]], [[10, 11], [13, 14]], [[9, 11], [10, 14]]])` bi trebalo da vrati `13`.

```js
assert.equal(getMeetingTime([[[9, 10], [12, 15]], [[10, 11], [13, 14]], [[9, 11], [10, 14]]]), 13);
```

`getMeetingTime([[[7, 8], [9, 11], [12, 14], [15, 16]], [[8, 11], [12, 13], [14, 15]]])` bi trebalo da vrati `9`.

```js
assert.equal(getMeetingTime([[[7, 8], [9, 11], [12, 14], [15, 16]], [[8, 11], [12, 13], [14, 15]]]), 9);
```

`getMeetingTime([[[7, 8], [10, 12], [13, 15]], [[8, 11], [12, 13], [14, 15]], [[6, 7], [8, 9], [12, 13]]])` bi trebalo da vrati `None`.

```js
assert.equal(getMeetingTime([[[7, 8], [10, 12], [13, 15]], [[8, 11], [12, 13], [14, 15]], [[6, 7], [8, 9], [12, 13]]]), "None");
```

`getMeetingTime([[[1, 3], [4, 6], [8, 10], [20, 23]], [[15, 16], [17, 18], [19, 22], [23, 24]], [[14, 16], [17, 23]], [[2, 4], [5, 6], [18, 19], [21, 22], [23, 24]]])` bi trebalo da vrati `21`.

```js
assert.equal(getMeetingTime([[[1, 3], [4, 6], [8, 10], [20, 23]], [[15, 16], [17, 18], [19, 22], [23, 24]], [[14, 16], [17, 23]], [[2, 4], [5, 6], [18, 19], [21, 22], [23, 24]]]), 21);
```

# --seed--

## --seed-contents--

```js
function getMeetingTime(availability) {

  return availability;
}
```

# --solutions--

```js
function getMeetingTime(availability) {
  for (let hour = 0; hour < 24; hour++) {
    const allFree = availability.every(person =>
      person.some(([start, end]) => start <= hour && hour + 1 <= end)
    );
    if (allFree) return hour;
  }
  return "None";
}
```