id: 69f35a5bb823ed620fcb7cbc
title: "Vežba 283: Zipovanje znakova iz niza"
challengeType: 28
dashedName: challenge-283
---

# --description--

Kada dobijete dva niza znakova, vratite novi niz koji kombinuje njihove znakove direktno po redu. Ako je jedan niz duži, dodajte preostale znakove na kraj.

S početi sa prvim znakom prvog niza.

# --hints--

`zipStrings("abc", "123")` bi trebalo da vrati `"a1b2c3"`.

```js
assert.equal(zipStrings("abc", "123"), "a1b2c3");
```

`zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz")` bi trebalo da vrati `"abcdefghijklmnopqrstuvwxyz"`.

```js
assert.equal(zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz"), "abcdefghijklmnopqrstuvwxyz");
```

`zipStrings("day", "night")` bi trebalo da vrati `"dnaiyght"`.

```js
assert.equal(zipStrings("day", "night"), "dnaiyght");
```

`zipStrings("python", "javascript")` bi trebalo da vrati `"pjyatvhaosncript"`.

```js
assert.equal(zipStrings("python", "javascript"), "pjyatvhaosncript");
```

`zipStrings("feCdCm", "reoeap")` bi trebalo da vrati `"freeCodeCamp"`.

```js
assert.equal(zipStrings("feCdCm", "reoeap"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function zipStrings(a, b) {

  return a;
}
```

# --solutions--

```js
function zipStrings(a, b) {
  let result = '';
  const len = Math.max(a.length, b.length);
  for (let i = 0; i < len; i++) {
    if (i < a.length) result += a[i];
    if (i < b.length) result += b[i];
  }
  return result;
}
```