id: 69f35a5bb823ed620fcb7cbd
title: "Vežba 284: I pre E"
challengeType: 28
dashedName: challenge-284
---

# --description--

Kada dobijete reč ili rečenicu, vratite ispravljenu verziju gde svaka reč prati pravilo "I pre E osim nakon C".

- Ako reč sadrži `"ei"` i nije prethodno napisana sa `"c"`, zamenite je sa `"ie"`.
- Ako reč sadrži `"ie"` i prethodi na `"c"`, zamenite je sa `"ei"`.
- Sve ostale reči ostaju nepromenjene.

# --hints--

`iBeforeE("beleive")` bi trebalo da vrati `"believe"`.

```js
assert.equal(iBeforeE("beleive"), "believe");
```

`iBeforeE("recieve")` bi trebalo da vrati `"receive"`.

```js
assert.equal(iBeforeE("recieve"), "receive");
```

`iBeforeE("we recieved a breif")` bi trebalo da vrati `"we received a brief"`.

```js
assert.equal(iBeforeE("we recieved a breif"), "we received a brief");
```

`iBeforeE("she beleived the friendly niece could percieve the greif")` bi trebalo da vrati `"she believed the friendly niece could perceive the grief"`.

```js
assert.equal(iBeforeE("she beleived the friendly niece could percieve the greif"), "she believed the friendly niece could perceive the grief");
```

`iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit")` bi trebalo da vrati `"we received relief after the thief gave us a brief piece of fierce deceit"`.

```js
assert.equal(iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit"), "we received relief after the thief gave us a brief piece of fierce deceit");
```

# --seed--

## --seed-contents--

```js
function iBeforeE(sentence) {

  return sentence;
}
```

# --solutions--

```js
function iBeforeE(sentence) {
  return sentence.split(' ').map(word => {
    word = word.replace(/([^c])ei/g, '$1ie');
    word = word.replace(/cie/g, 'cei');
    return word;
  }).join(' ');
}
```