id: 69f8c998d78ad3171a0713bf
title: "Vežba 292: Široki omjer stranica"
challengeType: 28
dashedName: challenge-292
---

# --description--

Ako vam se da dva niza karaktera sa različitim dimenzijama slika, vratite omjer stranica slike koja ima najveći omjer širine na visinu.

- Omjer karaktera koji je dat će biti u formatu `"ŠxV"`, na primer, `"1920x1080"`.
- Omjer stranica je odnos širine na visinu, pojednostavljen do najmanjih celih brojeva. Na primer, `"1920x1080"` se pojednostavljuje u `"16:9"`.
- Vratite niz karaktera u formatu `"Š:V"`, na primer, `"16:9"`.

# --hints--

`getWiderAspectRatio("1920x1080", "800x600")` bi trebalo da vrati `"16:9"`.

```js
assert.equal(getWiderAspectRatio("1920x1080", "800x600"), "16:9");
```

`getWiderAspectRatio("1080x1350", "2048x1536")` bi trebalo da vrati `"4:3"`.

```js
assert.equal(getWiderAspectRatio("1080x1350", "2048x1536"), "4:3");
```

`getWiderAspectRatio("640x480", "2440x1220")` bi trebalo da vrati `"2:1"`.

```js
assert.equal(getWiderAspectRatio("640x480", "2440x1220"), "2:1");
```

`getWiderAspectRatio("360x640", "1080x1920")` bi trebalo da vrati `"9:16"`.

```js
assert.equal(getWiderAspectRatio("360x640", "1080x1920"), "9:16");
```

`getWiderAspectRatio("3440x1440", "2048x858")` bi trebalo da vrati `"43:18"`.

```js
assert.equal(getWiderAspectRatio("3440x1440", "2048x858"), "43:18");
```

`getWiderAspectRatio("12345x61234", "12534x51234")` bi trebalo da vrati `"2089:8539"`.

```js
assert.equal(getWiderAspectRatio("12345x61234", "12534x51234"), "2089:8539");
```

# --seed--

## --seed-contents--

```js
function getWiderAspectRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function getWiderAspectRatio(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  function parse(str) {
    const [w, h] = str.split("x").map(Number);
    const g = gcd(w, h);
    return { ratio: w / h, simplified: `${w / g}:${h / g}` };
  }

  const imageA = parse(a);
  const imageB = parse(b);

  return imageA.ratio >= imageB.ratio ? imageA.simplified : imageB.simplified;
}
```