id: 69f8c998d78ad3171a0713be
title: "Vežba 291: Brojanje FizzBuzz-a"
challengeType: 28
dashedName: challenge-291
---

# --description--

Dajući početno i krajnje brojište, izračunajte broj pojavljivanja fizz-a i buzz-a u tom opsegu (uključujući te brojeve).

- Brojevi koji su deljivi sa 3 smatraju se fizz-om.
- Brojevi koji su deljivi sa 5 smatraju se buzz-om.
- Brojevi koji su deljivi i sa 3 i sa 5 smatraju se i fizz-om i buzz-om.

Vratite objekat ili mapu sa brojačima u formatu: `{ fizz, buzz }`.

# --hints--

`fizzBuzzCount(1, 11)` bi trebalo da vrati `{fizz: 3, buzz: 2}`.

```js
assert.deepEqual(fizzBuzzCount(1, 11), {fizz: 3, buzz: 2});
```

`fizzBuzzCount(14, 41)` bi trebalo da vrati `{fizz: 9, buzz: 6}`.

```js
assert.deepEqual(fizzBuzzCount(14, 41), {fizz: 9, buzz: 6});
```

`fizzBuzzCount(24, 100)` bi trebalo da vrati `{fizz: 26, buzz: 16}`.

```js
assert.deepEqual(fizzBuzzCount(24, 100), {fizz: 26, buzz: 16});
```

`fizzBuzzCount(-635, -14)` bi trebalo da vrati `{fizz: 207, buzz: 125}`.

```js
assert.deepEqual(fizzBuzzCount(-635, -14), {fizz: 207, buzz: 125});
```

`fizzBuzzCount(-5432, 6789)` bi trebalo da vrati `{fizz: 4074, buzz: 2444}`.

```js
assert.deepEqual(fizzBuzzCount(-5432, 6789), {fizz: 4074, buzz: 2444});
```

# --seed--

## --seed-contents--

```js
function fizzBuzzCount(start, end) {

  return start;
}
```

# --solutions--

```js
function fizzBuzzCount(start, end) {
  let fizz = 0, buzz = 0;

  for (let i = start; i <= end; i++) {
    if (i % 3 === 0) fizz++;
    if (i % 5 === 0) buzz++;
  }

  return { fizz, buzz };
}
```