id: 69f8c998d78ad3171a0713ba
title: "Vežba 287: Ispravljač Rimskih brojeva"
challengeType: 28
dashedName: challenge-287
---

# --description--

Dato niz loše formatiranih rimskih brojeva, vratite njegovu vrijednost u standardnu notaciju rimskih brojeva.

Unos će koristiti samo aditivnu notaciju, tako da svaki simbol dodaje svoju vrijednost ukupnom iznosu. Kao podsjetnik, ovo su simboli i njihove vrijednosti:

| Simbol | Vrijednost |
| ------ | --------- |
| `"I"`  | 1         |
| `"V"`  | 5         |
| `"X"`  | 10        |
| `"L"`  | 50        |
| `"C"`  | 100       |
| `"D"`  | 500       |
| `"M"`  | 1000      |

Prilikom ponovnog kodiranja, koristite najveći mogući simbol na svakom koraku, koristeći oduzimajuće parove (`"IV"`, `"IX"`, `"XL"`, `"XC"`, `"CD"`, `"CM"`) kada je potrebno.

# --hints--

`fixNumerals("XIIIII")` bi trebao vratiti `"XV"`.

```js
assert.equal(fixNumerals("XIIIII"), "XV");
```

`fixNumerals("IIIILX")` bi trebao vratiti `"LXIV"`.

```js
assert.equal(fixNumerals("IIIILX"), "LXIV");
```

`fixNumerals("XXVVVIIIII")` bi trebao vratiti `"XL"`.

```js
assert.equal(fixNumerals("XXVVVIIIII"), "XL");
```

`fixNumerals("MDCCLXXXXVIIII")` bi trebao vratiti `"MDCCXCIX"`.

```js
assert.equal(fixNumerals("MDCCLXXXXVIIII"), "MDCCXCIX");
```

`fixNumerals("IIIIVVVVXXXXLLLLCCDD")` bi trebao vratiti `"MCDLXIV"`.

```js
assert.equal(fixNumerals("IIIIVVVVXXXXLLLLCCDD"), "MCDLXIV");
```

`fixNumerals("ILCDMIVDIIXLCVCXDL")` bi trebao vratiti `"MMCMLXXXIV"`.

```js
assert.equal(fixNumerals("ILCDMIVDIIXLCVCXDL"), "MMCMLXXXIV");
```

# --seed--

## --seed-contents--

```js
function fixNumerals(str) {

  return str;
}
```

# --solutions--

```js
function fixNumerals(str) {
  const values = { I: 1, V: 5, X: 10, L: 50, C: 100, D: 500, M: 1000 };
  const encoding = [
    [1000, "M"], [900, "CM"], [500, "D"], [400, "CD"],
    [100, "C"], [90, "XC"], [50, "L"], [40, "XL"],
    [10, "X"], [9, "IX"], [5, "V"], [4, "IV"], [1, "I"]
  ];

  let total = [...str].reduce((sum, c) => sum + values[c], 0);

  let result = "";
  for (const [value, symbol] of encoding) {
    while (total >= value) {
      result += symbol;
      total -= value;
    }
  }
  return result;
}
```