id: 69f8c998d78ad3171a0713bd
title: "Zadatak 290: Pizza zabava"
challengeType: 28
dashedName: challenge-290
---

# --description--

Dato set sati koje je svako radilo danas, vratite broj pizzy koje treba naručiti za pizza zabavu.

- Podelite sate rada svakog sa 3 da biste dobili koliko komadića pice bi trebalo.
- Ne možete jesti man komad, pa zaokružite broj komadića svakoga na sledeći ceo broj.
- Svako dobija najmanje dva komadića.
- Pizza ima 8 komadića. Zaokružite ukupan broj pizzy na sledeću celobrojnu pizzu.

# --hints--

`getPizzasToOrder([8, 8, 8])` treba vratiti `2`.

```js
assert.equal(getPizzasToOrder([8, 8, 8]), 2);
```

`getPizzasToOrder([10, 9, 8, 2, 2, 6, 10])` treba vratiti `3`.

```js
assert.equal(getPizzasToOrder([10, 9, 8, 2, 2, 6, 10]), 3);
```

`getPizzasToOrder([1, 2, 3, 4, 5])` treba vratiti `2`.

```js
assert.equal(getPizzasToOrder([1, 2, 3, 4, 5]), 2);
```

`getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8])` treba vratiti `3`.

```js
assert.equal(getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8]), 3);
```

`getPizzasToOrder([9, 9, 6])` treba vratiti `1`.

```js
assert.equal(getPizzasToOrder([9, 9, 6]), 1);
```

`getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0])` treba vratiti `5`.

```js
assert.equal(getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5);
```

# --seed--

## --seed-contents--

```js
function getPizzasToOrder(hoursWorked) {

  return hoursWorked;
}
```

# --solutions--

```js
function getPizzasToOrder(hoursWorked) {
  const totalSlices = hoursWorked.reduce((sum, hours) => {
    return sum + Math.max(Math.ceil(hours / 3), 2);
  }, 0);

  return Math.ceil(totalSlices / 8);
}
```