id: 69f8c998d78ad3171a0713bb
title: "Vežba 288: Tajni broj"
challengeType: 28
dashedName: challenge-288
---

# --description--

Ako vam je dat tajni broj i pogađanje, utvrdite da li je pogađanje tačno.

Vratiti:

- `"higher"` ako je tajni broj veći od pogađanja.
- `"lower"` ako je tajni broj manji od pogađanja.
- `"you got it!"` ako je pogađanje tačno.

# --hints--

`guessNumber(50, 30)` bi trebalo da vrati `"higher"`.

```js
assert.equal(guessNumber(50, 30), "higher");
```

`guessNumber(85, 99)` bi trebalo da vrati `"lower"`.

```js
assert.equal(guessNumber(85, 99), "lower");
```

`guessNumber(2026, 2026)` bi trebalo da vrati `"you got it!"`.

```js
assert.equal(guessNumber(2026, 2026), "you got it!");
```

`guessNumber(92904, 11283)` bi trebalo da vrati `"higher"`.

```js
assert.equal(guessNumber(92904, 11283), "higher");
```

`guessNumber(230495, 423920)` bi trebalo da vrati `"lower"`.

```js
assert.equal(guessNumber(230495, 423920), "lower");
```

`guessNumber(120349, 120349)` bi trebalo da vrati `"you got it!"`.

```js
assert.equal(guessNumber(120349, 120349), "you got it!");
```

# --seed--

## --seed-contents--

```js
function guessNumber(secret, guess) {

  return secret;
}
```

# --solutions--

```js
function guessNumber(secret, guess) {
  if (guess === secret) return "you got it!";
  return guess < secret ? "higher" : "lower";
}
```