id: 69f90c1329a94b37e2a2086c
title: "Zadatak 293: Najbolja ruka"
challengeType: 28
dashedName: challenge-293
---

# --description--

Dato sa sastoji od pet sekvencijalnih karakterističkih nizova koji predstavljaju igrane karte; vratite naziv najbolje ruke.

- Svaka karta je predstavljena kao dvokarakterni niz: rang praćen šakama, npr., `"2h"`.
  - Rangi, od najnižeg do najvišeg, su: `"2"`, `"3"`, `"4"`, `"5"`, `"6"`, `"7"`, `"8"`, `"9"`, `"T"`, `"J"`, `"Q"`, `"K"`, i `"A"`.
  - Šake su: `"h"`, `"d"`, `"c"`, i `"s"`.
- Asovi (`"A"`) mogu se koristiti kao najviši ili najniži u sekvenci.

Ruke, od najlošije do najbolje, su:

| Naziv | Opis |
| - | - |
| `"High Card"` | Nema para niti bolje ruke |
| `"Pair"` | Jedan par iste rangirane karte |
| `"Two Pair"` | Dva para različitih ranžiranih karata |
| `"Three of a Kind"` | Tri karte istog ranga |
| `"Straight"` | Sekvenca od pet uzastopnih ranžirnih karata |
| `"Flush"` | Pet karata iste šake |
| `"Full House"` | Tri karte istog ranga i par drugog ranga |
| `"Four of a Kind"` | Četiri karte istog ranga |
| `"Straight Flush"` | Sekvenca od pet uzastopnih ranžirnih karata iste šake |
| `"Royal Flush"` | `"A"`, `"K"`, `"Q"`, `"J"`, `"T"` iste šake |

Vratite naziv najbolje ruke.

# --hints--

`getBestHand(["7s", "7h", "7d", "2c", "5h"])` bi trebalo da vrati `"Three of a Kind"`.

```js
assert.equal(getBestHand(["7s", "7h", "7d", "2c", "5h"]), "Three of a Kind");
```

`getBestHand(["Ks", "Kh", "Kd", "4s", "4h"])` bi trebalo da vrati `"Full House"`.

```js
assert.equal(getBestHand(["Ks", "Kh", "Kd", "4s", "4h"]), "Full House");
```

`getBestHand(["2h", "5h", "7h", "9h", "Jh"])` bi trebalo da vrati `"Flush"`.

```js
assert.equal(getBestHand(["2h", "5h", "7h", "9h", "Jh"]), "Flush");
```

`getBestHand(["As", "Ah", "Ad", "Ac", "Kh"])` bi trebalo da vrati `"Four of a Kind"`.

```js
assert.equal(getBestHand(["As", "Ah", "Ad", "Ac", "Kh"]), "Four of a Kind");
```

`getBestHand(["Ts", "Th", "9d", "9c", "8h"])` bi trebalo da vrati `"Two Pair"`.

```js
assert.equal(getBestHand(["Ts", "Th", "9d", "9c", "8h"]), "Two Pair");
```

`getBestHand(["9c", "8c", "7c", "6c", "5c"])` bi trebalo da vrati `"Straight Flush"`.

```js
assert.equal(getBestHand(["9c", "8c", "7c", "6c", "5c"]), "Straight Flush");
```

`getBestHand(["As", "Kh", "Jd", "8c", "5h"])` bi trebalo da vrati `"High Card"`.

```js
assert.equal(getBestHand(["As", "Kh", "Jd", "8c", "5h"]), "High Card");
```

`getBestHand(["As", "2h", "3d", "4c", "5h"])` bi trebalo da vrati `"Straight"`.

```js
assert.equal(getBestHand(["As", "2h", "3d", "4c", "5h"]), "Straight");
```

`getBestHand(["Ts", "Th", "7c", "6d", "5h"])` bi trebalo da vrati `"Pair"`.

```js
assert.equal(getBestHand(["Ts", "Th", "7c", "6d", "5h"]), "Pair");
```

`getBestHand(["As", "Ks", "Qs", "Js", "Ts"])` bi trebalo da vrati `"Royal Flush"`.

```js
assert.equal(getBestHand(["As", "Ks", "Qs", "Js", "Ts"]), "Royal Flush");
```

# --seed--

## --seed-contents--

```js
function getBestHand(cards) {

  return cards;
}
```

# --solutions--

```js
function getBestHand(cards) {
  const rankOrder = "23456789TJQKA";

  const ranks = cards.map(c => c[0]);
  const suits = cards.map(c => c[1]);

  const isFlush = suits.every(s => s === suits[0]);

  const rankIndices = ranks.map(r => rankOrder.indexOf(r)).sort((a, b) => a - b);

  function isStraight(indices) {
    const isRegular = indices.every((r, i) => i === 0 || r === indices[i - 1] + 1);
    const isLowAce = JSON.stringify(indices) === JSON.stringify([0, 1, 2, 3, 12]);
    return isRegular || isLowAce;
  }

  const straight = isStraight(rankIndices);

  const counts = {};
  ranks.forEach(r => counts[r] = (counts[r] || 0) + 1);
  const freq = Object.values(counts).sort((a, b) => b - a);

  if (isFlush && JSON.stringify(rankIndices) === JSON.stringify([8, 9, 10, 11, 12])) return "Royal Flush";
  if (isFlush && straight) return "Straight Flush";
  if (freq[0] === 4) return "Four of a Kind";
  if (freq[0] === 3 && freq[1] === 2) return "Full House";
  if (isFlush) return "Flush";
  if (straight) return "Straight";
  if (freq[0] === 3) return "Three of a Kind";
  if (freq[0] === 2 && freq[1] === 2) return "Two Pair";
  if (freq[0] === 2) return "Pair";
  return "High Card";
}
```