id: 69f90c1329a94b37e2a2086d
title: "Vežba 294: Kombinacije običnih zagrada"
challengeType: 28
dashedName: challenge-294
---

# --description--

Za dato prirodno broj, `n`, vratite broj validnih kombinacija sa parovima `n` običnih zagrada.

- Validna kombinacija je sekvenca karaktera gde svaka otvarajuća zagrada ima odgovarajuću zatvarajuću zagradu, i nijedna zatvarajuća zagrada ne nastaje pre odgovarajuće otvarajuće zagrade.

Na primer, za `2`, postoji 2 validne kombinacije:

```json
(())
()()
```

# --hints--

`getCombinations(2)` bi trebalo da vrati `2`.

```js
assert.equal(getCombinations(2), 2);
```

`getCombinations(3)` bi trebalo da vrati `5`.

```js
assert.equal(getCombinations(3), 5);
```

`getCombinations(5)` bi trebalo da vrati `42`.

```js
assert.equal(getCombinations(5), 42);
```

`getCombinations(8)` bi trebalo da vrati `1430`.

```js
assert.equal(getCombinations(8), 1430);
```

`getCombinations(13)` bi trebalo da vrati `742900`.

```js
assert.equal(getCombinations(13), 742900);
```

# --seed--

## --seed-contents--

```js
function getCombinations(n) {

  return n;
}
```

# --solutions--

```js
function getCombinations(n) {
  let count = 0;

  function generate(open, close) {
    if (open === 0 && close === 0) {
      count++;
      return;
    }
    if (open > 0) generate(open - 1, close);
    if (close > open) generate(open, close - 1);
  }

  generate(n, n);
  return count;
}
```