id: 6a0dcc730cb92a616f86f0bf
title: "Vežba 295: Provera strukture deo 1"
challengeType: 28
dashedName: challenge-295
---

# --description--

Ako vam je dat objekat (JavaScript) ili rečnik (Python), proverite da li odgovara sledećoj strukturi:

```json
{
  username: string
}
```

- Dodatni ključevi su dozvoljeni

# --hints--

`isValidSchema({ username: "bob" })` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidSchema({ username: "bob" }));
```

`isValidSchema({ username: "jen", posts: 30 })` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidSchema({ username: "jen", posts: 30 }));
```

`isValidSchema({ username: "" })` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidSchema({ username: "" }));
```

`isValidSchema({ username: 7 })` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidSchema({ username: 7 }));
```

`isValidSchema({ posts: 25 })` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidSchema({ posts: 25 }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  return typeof obj.username === 'string';
}
```