id: 6a0dcd03ee4e68698080ef6b
title: "Zadatak 307: Pravila urbanističkog planiranja"
challengeType: 28
dashedName: challenge-307
---

# --description--

Ako vam je dat dvodimenzionalni grid (matrica podataka sa redovima podataka) koji predstavlja raspored gradskih objekata, vratite koordinate svih objekata koji krše pravila planiranja.

Svaka ćelija u gridu ima jednu oznaku iz tabele ispod. Objekat krši pravilo ako bilo koji od njegovih susjeda (do četiri), vertikalno ili horizontalno, nije tip koji mu sme biti susjed.

| Simbol | Vrsta | Ne sme biti susjed sa |
| - | - | - |
| `"i"` | industrija | `"R"`, `"I"` |
| `"A"` | poljoprivreda | `"C"` |
| `"R"` | stambeno | `"i"`, `"C"` |
| `"I"` | institucije | `"i"` |
| `"C"` | komercijalno | `"R"`, `"A"` |
| `""` (prazan niz karaktera) | neizgrađeno | nema ograničenja |

Vratite koordinate svih ćelija koje krše pravila kao niz parova `[red, kol]`, u bilo kom redosledu. Ako ne postoji nijedno kršenje, vratite prazan niz podataka.

# --hints--

`getZoneViolations([["R", "C"], ["", "C"]])` bi trebalo da vrati `[[0, 0], [0, 1]]`.

```js
assert.deepEqual(getZoneViolations([["R", "C"], ["", "C"]]).sort(), [[0, 0], [0, 1]]);
```

`getZoneViolations([["", "i"], ["", "R"], ["R", "I"]])` bi trebalo da vrati `[[0, 1], [1, 1]]`.

```js
assert.deepEqual(getZoneViolations([["", "i"], ["", "R"], ["R", "I"]]).sort(), [[0, 1], [1, 1]]);
```

`getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` bi trebalo da vrati `[]`.

```js
assert.deepEqual(getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]]).sort(), []);
```

`getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` bi trebalo da vrati `[[0, 1], [0, 2], [0, 3]]`.

```js
assert.deepEqual(getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]]).sort(), [[0, 1], [0, 2], [0, 3]]);
```


`getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` bi trebalo da vrati `[[2, 3], [2, 4], [3, 1], [3, 2]]`.

```js
assert.deepEqual(getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]]).sort(), [[2, 3], [2, 4], [3, 1], [3, 2]]);
```

# --seed--

## --seed-contents--

```js
function getZoneViolations(grid) {

  return grid;
}
```

# --solutions--

```js
function getZoneViolations(grid) {
  const rules = {
    "i": ["R", "I"],
    "A": ["C"],
    "R": ["i", "C"],
    "I": ["i"],
    "C": ["R", "A"],
    "": []
  };

  const directions = [[-1, 0], [0, -1], [0, 1], [1, 0]];
  const violations = [];

  for (let row = 0; row < grid.length; row++) {
    for (let col = 0; col < grid[row].length; col++) {
      const cell = grid[row][col];
      const forbidden = rules[cell];

      for (const [dr, dc] of directions) {
        const nr = row + dr;
        const nc = col + dc;
        if (nr >= 0 && nr < grid.length && nc >= 0 && nc < grid[row].length) {
          if (forbidden.includes(grid[nr][nc])) {
            violations.push([row, col]);
            break;
          }
        }
      }
    }
  }

  return violations;
}
```