id: 6a0dcc730cb92a616f86f0c1
title: "Vežba 297: Proveri strukturu deo 3"
challengeType: 28
dashedName: challenge-297
---

# --description--

Ako vam je dat objekat (JavaScript) ili rečnik (Python), utvrdite da li odgovara sledećoj strukturi:

```json
Roles = "user" | "creator" | "moderator" | "staff" | "admin"

{
  username: string,
  posts: number,
  verified: boolean,
  role: Roles
}
```

- Simbol pipe (`|`) znači "ili". `role` mora biti jedna od navedenih vrednosti iz `Roles`.
- Dodatni ključevi su dozvoljeni

# --hints--

`isValidSchema({ username: "henry", posts: 0, verified: true, role: "staff" })` treba da vrati `true`.

```js
assert.isTrue(isValidSchema({ username: "henry", posts: 0, verified: true, role: "staff" }));
```

`isValidSchema({ username: "sara", posts: 45, verified: false, role: "creator", followers: 70 })` treba da vrati `true`.

```js
assert.isTrue(isValidSchema({ username: "sara", posts: 45, verified: false, role: "creator", followers: 70 }));
```

`isValidSchema({ username: "penelope", posts: 20, verified: true, role: "admin" })` treba da vrati `true`.

```js
assert.isTrue(isValidSchema({ username: "penelope", posts: 20, verified: true, role: "admin" }));
```

`isValidSchema({ username: "kevin", posts: 0, verified: false, role: "user" })` treba da vrati `true`.

```js
assert.isTrue(isValidSchema({ username: "kevin", posts: 0, verified: false, role: "user" }));
```

`isValidSchema({ username: "george", posts: 15, verified: true, role: "moderator" })` treba da vrati `true`.

```js
assert.isTrue(isValidSchema({ username: "george", posts: 15, verified: true, role: "moderator" }));
```

`isValidSchema({ username: "david", posts: 0, verified: false, role: "guest" })` treba da vrati `false`.

```js
assert.isFalse(isValidSchema({ username: "david", posts: 0, verified: false, role: "guest" }));
```

`isValidSchema({ username: "wendy", posts: 10, verified: true })` treba da vrati `false`.

```js
assert.isFalse(isValidSchema({ username: "wendy", posts: 10, verified: true }));
```

`isValidSchema({ username: "fabian", posts: 1, verified: true, role: true })` treba da vrati `false`.

```js
assert.isFalse(isValidSchema({ username: "fabian", posts: 1, verified: true, role: true }));
```

`isValidSchema({ username: 8, posts: 1, verified: true, role: "user" })` treba da vrati `false`.

```js
assert.isFalse(isValidSchema({ username: 8, posts: 1, verified: true, role: "user" }));
```

`isValidSchema({ username: "penny", posts: "10", verified: true, role: "staff" })` treba da vrati `false`.

```js
assert.isFalse(isValidSchema({ username: "penny", posts: "10", verified: true, role: "staff" }));
```

`isValidSchema({ username: "john", posts: "1", verified: "true", role: "admin" })` treba da vrati `false`.

```js
assert.isFalse(isValidSchema({ username: "john", posts: "1", verified: "true", role: "admin" }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  const roles = ["user", "creator", "moderator", "staff", "admin"];
  return (
    typeof obj.username === 'string' &&
    typeof obj.posts === 'number' &&
    typeof obj.verified === 'boolean' &&
    roles.includes(obj.role)
  );
}
```