id: 6a0dcd03ee4e68698080ef6c
title: "Vežba 308: Provera kreditne kartice"
challengeType: 28
dashedName: challenge-308
---

# --description--

Ako vam je dat niz brojeva kreditne kartice, utvrdite da li je broj validan koristeći sledeću metodu:

- Počnite od drugog broja odozgo (od kraja), pomnožite svaki drugi broj ka levi.
- Ako množenje broja rezultira brojem većim od 9, oduzmite 9.
- Saberi sve brojeve (oni koji su pomnoženi i oni koji nisu).
- Ako je zbir deljiv sa 10, broj je validan.

# --hints--

`isValidCard("4532015112830366")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidCard("4532015112830366"));
```

`isValidCard("5425233430109903")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidCard("5425233430109903"));
```

`isValidCard("371449635398431")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidCard("371449635398431"));
```

`isValidCard("6011111111111117")` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidCard("6011111111111117"));
```

`isValidCard("4532015112830367")` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidCard("4532015112830367"));
```

`isValidCard("1234567890123456")` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidCard("1234567890123456"));
```

`isValidCard("4532015112830368")` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidCard("4532015112830368"));
```

# --seed--

## --seed-contents--

```js
function isValidCard(number) {

  return number;
}
```

# --solutions--

```js
function isValidCard(number) {
  let sum = 0;

  for (let i = number.length - 1; i >= 0; i--) {
    let digit = parseInt(number[i]);

    if ((number.length - 1 - i) % 2 === 1) {
      digit *= 2;
      if (digit > 9) digit -= 9;
    }

    sum += digit;
  }

  return sum % 10 === 0;
}
```