id: 6a0dcd03ee4e68698080ef66
title: "Vežba 302: Jet Lag"
challengeType: 28
dashedName: challenge-302
---

# --description--

Dato polazno grad, dolazni grad, trajanje leta u satima i smer putovanja, vratite broj sati jet laga (jet lag) koje putnik doživljava.

Gradovi su uzeti iz sledeće liste koja uključuje njihove UTC odstupanja:

| Grad | Odstupanje |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

Da biste izračunali sate jet laga:

1. Pronađite razliku u vremenskom odstupanju između dva grada.
2. Odredite faktor smera. Ako putujete "east" (istočno), on je 1.5, inače je 1.0.
3. Dobijte sate jet laga po formuli: razlika u vremenskom odstupanju + (trajanje leta * 0.1) * faktor smera

Vratite sate jet laga zaokružene na jednu decimalnu tačku.

# --hints--

`getJetLagHours("Istanbul", "Hong Kong", 10, "east")` bi trebalo da vrati `6.5`.

```js
assert.equal(getJetLagHours("Istanbul", "Hong Kong", 10, "east"), 6.5);
```

`getJetLagHours("London", "New York", 8, "west")` bi trebalo da vrati `5.8`.

```js
assert.equal(getJetLagHours("London", "New York", 8, "west"), 5.8);
```

`getJetLagHours("Hong Kong", "Tokyo", 4, "east")` bi trebalo da vrati `1.6`.

```js
assert.equal(getJetLagHours("Hong Kong", "Tokyo", 4, "east"), 1.6);
```

`getJetLagHours("Dubai", "London", 7, "west")` bi trebalo da vrati `4.7`.

```js
assert.equal(getJetLagHours("Dubai", "London", 7, "west"), 4.7);
```

`getJetLagHours("Los Angeles", "Hong Kong", 15, "west")` bi trebalo da vrati `17.5`.

```js
assert.equal(getJetLagHours("Los Angeles", "Hong Kong", 15, "west"), 17.5);
```

`getJetLagHours("Tokyo", "Dubai", 9, "west")` bi trebalo da vrati `5.9`.

```js
assert.equal(getJetLagHours("Tokyo", "Dubai", 9, "west"), 5.9);
```

`getJetLagHours("New York", "Istanbul", 10, "east")` bi trebalo da vrati `9.5`.

```js
assert.equal(getJetLagHours("New York", "Istanbul", 10, "east"), 9.5);
```

# --seed--

## --seed-contents--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {

  return departureCity;
}
```

# --solutions--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {
  const offsets = {
    "Los Angeles": -8,
    "New York": -5,
    "London": 0,
    "Istanbul": +3,
    "Dubai": +4,
    "Hong Kong": +8,
    "Tokyo": +9
  };

  const depOffset = offsets[departureCity];
  const arrOffset = offsets[arrivalCity];

  const timezoneDiff = Math.abs(arrOffset - depOffset);
  const multiplier = direction === "east" ? 1.5 : 1.0;

  const score = timezoneDiff + (flightDuration * 0.1) * multiplier;
  return Math.round(score * 10) / 10;
}
```