id: 6a0dcc730cb92a616f86f0c5
title: "Vežba 301: Poslednji teret"
challengeType: 28
dashedName: challenge-301
---

# --description--

Uzimajući u obzir broj kašika deterdženta koje imate i niz podataka o kašikama koje ste koristili tokom proteklih dana, vratite broj punih dana za sapun koji vam preostaje.

Izračunajte prosečnu potrošnju po danu iz istorije upotrebe i pretpostavite tu potrošnju svaki dan od sada.

# --hints--

`lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2])` bi trebalo da vrati `5`.

```js
assert.equal(lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2]), 5);
```

`lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1])` bi trebalo da vrati `7`.

```js
assert.equal(lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1]), 7);
```

`lastLoadDate(33, [5, 0, 4, 3, 3, 2])` bi trebalo da vrati `11`.

```js
assert.equal(lastLoadDate(33, [5, 0, 4, 3, 3, 2]), 11);
```

`lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2])` bi trebalo da vrati `12`.

```js
assert.equal(lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2]), 12);
```

`lastLoadDate(20, [13, 9, 12, 10, 8])` bi trebalo da vrati `1`.

```js
assert.equal(lastLoadDate(20, [13, 9, 12, 10, 8]), 1);
```

# --seed--

## --seed-contents--

```js
function lastLoadDate(scoops, usage) {

  return scoops;
}
```

# --solutions--

```js
function lastLoadDate(scoops, usage) {
  const avg = usage.reduce((sum, n) => sum + n, 0) / usage.length;
  return Math.floor(scoops / avg);
}
```