id: 6a0dcc730cb92a616f86f0c0
title: "Vežba 296: Provera šeme deo 2"
challengeType: 28
dashedName: challenge-296
---

# --description--

Ako vam je dat objekat (JavaScript) ili rečnik (Python), utvrdite da li odgovara sledećoj strukturi:

```json
{
  username: string,
  posts: number,
  verified: boolean
}
```

- Dodatni ključevi su dozvoljeni

# --hints--

`isValidSchema({ username: "alice", posts: 10, verified: false })` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidSchema({ username: "alice", posts: 10, verified: false }));
```

`isValidSchema({ username: "carol", posts: 15, verified: true, followers: 25 })` bi trebalo da vrati `true`.

```js
assert.isTrue(isValidSchema({ username: "carol", posts: 15, verified: true, followers: 25 }));
```

`isValidSchema({ username: "frank", posts: "21", verified: true })` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidSchema({ username: "frank", posts: "21", verified: true }));
```

`isValidSchema({ username: "sam", posts: 17, verified: "false" })` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidSchema({ username: "sam", posts: 17, verified: "false" }));
```

`isValidSchema({ username: "bill", verified: true })` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidSchema({ username: "bill", verified: true }));
```

`isValidSchema({ username: "fred", verified: true })` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidSchema({ username: "fred", verified: true }));
```

`isValidSchema({ username: 5, posts: 10, verified: true })` bi trebalo da vrati `false`.

```js
assert.isFalse(isValidSchema({ username: 5, posts: 10, verified: true }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  return (
    typeof obj.username === 'string' &&
    typeof obj.posts === 'number' &&
    typeof obj.verified === 'boolean'
  );
}
```