id: 6a151d32d772271dd2448c2d
title: "Vežba 311: Čarobnjak reči"
challengeType: 28
dashedName: challenge-311
---

# --description--

Dato je sekvenca kodova reči; izračunajte ukupan rezultat.

Svaki slovo u sekvenci predstavlja reč:

| Kod | Reč | Kategorija | Osnovni rezultat |
| - | - | - | - |
| `"f"` | Vatra | Uništenje | 3 |
| `"l"` | Struja | Uništenje | 3 |
| `"i"` | Led | Kontrola | 2 |
| `"w"` | Vetar | Kontrola | 2 |
| `"h"` | Terapija | Obnavljanje | 1 |
| `"s"` | Čuvar | Obnavljanje | 1 |

Množilac se određuje na osnovu broja uzastopnih reči korišćenih iz različitih kategorija:

- Prva reč uvek dobija osnovni rezultat.
- Svaka naknadna reč iz različite kategorije od prethodne povećava množilac za 1.
- Korišćenje reči iz iste kategorije kao i prethodna vraća množilac na 1.
- Rezultat za svaku reč je njen osnovni rezultat pomnožen trenutnim množiocem.

Vratite ukupan rezultat sekvence reči.

# --hints--

`cast("fihwl")` bi trebalo da vrati `33`.

```js
assert.equal(cast("fihwl"), 33);
```

`cast("lwswfi")` bi trebalo da vrati `45`.

```js
assert.equal(cast("lwswfi"), 45);
```

`cast("wislhfl")` bi trebalo da vrati `37`.

```js
assert.equal(cast("wislhfl"), 37);
```

`cast("sihwlih")` bi trebalo da vrati `50`.

```js
assert.equal(cast("sihwlih"), 50);
```

`cast("wishlfihwslwifihl")` bi trebalo da vrati `101`.

```js
assert.equal(cast("wishlfihwslwifihl"), 101);
```

# --seed--

## --seed-contents--

```js
function cast(spells) {

  return spells;
}
```

# --solutions--

```js
function cast(spells) {
  const lookup = {
    f: { category: "destruction", score: 3 },
    l: { category: "destruction", score: 3 },
    i: { category: "control", score: 2 },
    w: { category: "control", score: 2 },
    h: { category: "restoration", score: 1 },
    s: { category: "restoration", score: 1 }
  };

  let total = 0;
  let multiplier = 1;
  let prevCategory = null;

  for (const spell of spells) {
    const { category, score } = lookup[spell];

    if (prevCategory && category !== prevCategory) {
      multiplier++;
    } else {
      multiplier = 1;
    }

    total += score * multiplier;
    prevCategory = category;
  }

  return total;
}
```