---
id: 587d7db7367417b2b2512b9c
title: Traži jednog ili više kriminalaca u praćenju
challengeType: 1
forumTopicId: 301343
dashedName: find-one-or-more-criminals-in-a-hunt
---

# --description--

Vreme je da malo staniš i probaš svoju novu veštinu pisanja uobičajenih fraza. Grupa kriminalaca pobegla je iz zatvora i trčala, ali ne znaš koliko ih je bilo. Međutim, znaš da ostaju blizu zajedno kada su u blizini drugih ljudi. Ti si odgovoran za pronalaženje svih kriminalaca odjednom.

Hapa kuna mfano wa kukagua jinsi ya kufanya hivi:

Standardni uzorak ``/z+/`` odgovara slovu ``z`` kada se pojavljuje jednom ili više puta zaredom. Pronaćet ćete podudarnosti u svim sledećim sekvencama:

```js
"z"
"zzzzzz"
"ABCzzzz"
"zzzzABC"
"abczzzzzzzzzzzzzzzzzzzzzabc"
```

Lakini haupati mechi katika mfuatano ufuatao kwa sababu hakuna herufi `z`:

```js
""
"ABC"
"abcabc"
```

# --instructions--

Napišite uobičajenu izjavu tipa „greedy“ koja traži jednog ili više krivaca unutar grupe drugih ljudi. Krivac je predstavljen velikim slovom `C`.

# --hints--

Vaša uobičajena upotreba bi trebalo da uporedi jednog kriminalca (`C`) u nizu `C`

```js
assert('C'.match(reCriminals) && 'C'.match(reCriminals)[0] == 'C');
```

Usemi wako wa kawaida unapaswa kulinganisha wahalifu wawili (`CC`) katika mfuatano `CC`

```js
assert('CC'.match(reCriminals) && 'CC'.match(reCriminals)[0] == 'CC');
```

Usemi wako wa kawaida unapaswa kulinganisha wahalifu watatu (`CCC`) katika mfuatano `P1P5P4CCCcP2P6P3`.

```js
assert(
  'P1P5P4CCCcP2P6P3'.match(reCriminals) &&
    'P1P5P4CCCcP2P6P3'.match(reCriminals)[0] == 'CCC'
);
```

Usemi wako wa kawaida unapaswa kulinganisha wahalifu watano (`CCCCC`) katika mfuatano `P6P2P7P4P5CCCCCP3P1`

```js
assert(
  'P6P2P7P4P5CCCCCP3P1'.match(reCriminals) &&
    'P6P2P7P4P5CCCCCP3P1'.match(reCriminals)[0] == 'CCCCC'
);
```

Usemi wako wa kawaida haupaswi kulinganisha wahalifu wowote katika mfuatano tupu `""`

```js
reCriminals.lastIndex = 0;
assert(!reCriminals.test(''));
```

Usemi wako wa kawaida haupaswi kulinganisha wahalifu wowote katika mfuatano `P1P2P3`

```js
reCriminals.lastIndex = 0;
assert(!reCriminals.test('P1P2P3'));
```

Standardna upotreba treba da uporedi pedeset kriminalaca (`CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCC`) u sekvenci `P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3`.

```js
assert(
  'P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3'.match(
    reCriminals
  ) &&
    'P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3'.match(
      reCriminals
    )[0] == 'CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCC'
);
```

# --seed--

## --seed-contents--

```js
let reCriminals = /./; // Change this line
```

# --solutions--

```js
let reCriminals = /C+/; // Change this line
```
