---
id: 587d7db5367417b2b2512b94
title: Uporedi bilo šta sa wildkardom za period
challengeType: 1
forumTopicId: 301348
dashedName: match-anything-with-wildcard-period
---

# --description--

Ponekad ne znaš (ili nema potrebe da znaš) tačne karaktere u svojim sistemima. Razmišljanje o svim poklapanjima, recimo, sa greškama u kucanju će potrajati dugo. Srećom, možeš uštedeti vreme korišćenjem wildkarta: `.`

Wildcard karakter `.` poklapa se sa bilo kojim pojedinačnim karakterom. Wildcard se takođe naziva `dot` i `period`. Možete koristiti wildcard karakter kao bilo koji drugi karakter u regularnom izrazu. Na primer, ako želite da poklopite `hug`, `huh`, `hut` i `hum`, možete koristiti regularni izraz `/hu./` da poklopite sva četiri elementa.

```js
let humStr = "I'll hum a song";
let hugStr = "Bear hug";
let huRegex = /hu./;
huRegex.test(humStr);
huRegex.test(hugStr);
```

Oba poziva od `test` će vratiti `true`.

# --instructions--

Kamilisha usemi wa kawaida `unRegex` ili ulingane na mfuatano wa herufi `run`, `sun`, `fun`, `pun`, `nun`, na `bun`. Usemi wako wa kawaida unapaswa kutumia herufi ya wildcard.

# --hints--

Treba da koristiš putanju `.test()`.

```js
assert(__helpers.removeJSComments(code).match(/\.test\(.*\)/));
```

Trebalo bi da koristiš znak za wildkard u svom običnom tekstu `unRegex`

```js
assert(/\./.test(unRegex.source));
```

Usemi wako wa kawaida `unRegex` unapaswa kulinganisha `run` katika mfuatano wa herufi `Let us go on a run.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Let us go on a run.'));
```

Vaša uobičajena vrednost `unRegex` mora da odgovara `sun` u nizu znakova `The sun is out today.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('The sun is out today.'));
```

Vaša uobičajena upotreba `unRegex` trebalo bi da odgovara `fun` u sekvenci karaktera `Coding is a lot of fun.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Coding is a lot of fun.'));
```

Vaše uobičajeno korišćenje `unRegex` mora da se poklapa sa `pun` u sekvenci slova `Seven days without a pun makes one weak.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Seven days without a pun makes one weak.'));
```

Vaša uobičajena vrednost `unRegex` mora odgovarati `nun` u nizu slova `One takes a vow to be a nun.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('One takes a vow to be a nun.'));
```

Vaš standardni `unRegex` mora da odgovara `bun` u nizu karaktera `She got fired from the hot dog stand for putting her hair in a bun.`

```js
unRegex.lastIndex = 0;
assert(
  unRegex.test(
    'She got fired from the hot dog stand for putting her hair in a bun.'
  )
);
```

Standardno korišćenje `unRegex` ne bi trebalo da upoređuje sekvencu slova `There is a bug in my code.`

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('There is a bug in my code.'));
```

Vaše uobičajeno korišćenje `unRegex` ne bi trebalo da uporedi niz znakova `Catch me if you can.`

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('Catch me if you can.'));
```

# --seed--

## --seed-contents--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /change/; // Change this line
let result = unRegex.test(exampleStr);
```

# --solutions--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /.un/; // Change this line
let result = unRegex.test(exampleStr);
```
