---
id: 587d7dbb367417b2b2512bab
title: Koristi grupe za hvatanje i menjanje
challengeType: 1
forumTopicId: 301368
dashedName: use-capture-groups-to-search-and-replace
---

# --description--

Pretraga je važna. Međutim, možete učiniti pretragu moćnijom kada takođe promenite (ili modifikujete) tekst koji poredite.

Unaweza kutafuta na kubadilisha maandishi katika mfuatano wa herufi kwa kutumia `.replace()` kwenye mfuatano wa herufi. Ingizo la `.replace()` ni kwanza mfano wa usemi wa kawaida unayotaka kutafuta. Kigezo cha pili ni mfuatano wa herufi wa kubadilisha matokeo au kitendakazi cha kufanya kitu fulani.

```js
let wrongText = "The sky is silver.";
let silverRegex = /silver/;
wrongText.replace(silverRegex, "blue");
```

Wito la `replace` litarudisha mfuatano wa herufi `The sky is blue.`.

Takođe, možete pronaći grupe za hvatanje u sekvenci zamene korišćenjem dolarskih znakova (`$`).

```js
"Code Camp".replace(/(\w+)\s(\w+)/, '$2 $1');
```

Poziv za `replace` će vratiti niz znakova `Camp Code`.

# --instructions--

Napišite redovni izraz (regex) `fixRegex` koristeći tri grupe za hvatanje koje će pronaći svaku reč u nizu znakova `one two three`. Zatim ažurirajte parametar `replaceText` da zamenite `one two three` sa nizom znakova `three two one` i stavite rezultat u parametar `result`. Uverite se da koristite grupe za hvatanje u nizu za zamenu koristeći sintaksu sa dolarovim znakom (`$`).

# --hints--

Trebalo bi da koristiš `.replace()` za pretraživanje i zamenu.

```js
assert(__helpers.removeJSComments(code).match(/\.replace\(.*\)/));
```

Vaš uobičajeni tekst mora da zameni niz karaktera `one two three` sa nizom karaktera `three two one`

```js
assert(result === 'three two one');
```

Hupaswi kubadilisha mstari wa mwisho.

```js
assert(__helpers.removeJSComments(code).match(/result\s*=\s*str\.replace\(.*?\)/));
```

`fixRegex` inapaswa kutumia angalau makundi matatu ya kunasa.

```js
assert(new RegExp(fixRegex.source + '|').exec('').length - 1 >= 3);
```

`replaceText` inapaswa kutumia mfuatano wa submatch ulio na mabano (yaani submatch ya nth yenye mabano, $n, inahusiana na kundi la nth la kunasa).

```js
{
  const re = /(\$\d{1,2})+(?:[\D]|\b)/g;
  assert(replaceText.match(re).length >= 3);
}
```

# --seed--

## --seed-contents--

```js
let str = "one two three";
let fixRegex = /change/; // Change this line
let replaceText = ""; // Change this line
let result = str.replace(fixRegex, replaceText);
```

# --solutions--

```js
let str = "one two three";
let fixRegex = /(\w+) (\w+) (\w+)/g; // Change this line
let replaceText = "$3 $2 $1"; // Change this line
let result = str.replace(fixRegex, replaceText);
```
