---
id: 587d8259367417b2b2512c86
title: Implementirajte sortiranje umetanjem
challengeType: 26
dashedName: implement-insertion-sort
---

# --description--

U ovom laboratorijumu, implementiraćete algoritam umetanje sortiranja. Ova metoda radi tako što gradi sortirovan niz na početku liste. On počinje sortirovanim nizom sa prvim elementom. Zatim inspekcira sledeći element i razmenjuje ga unazad u sortirovani niz dok ne dođe do sortirane pozicije. Nastavlja iteriranje kroz listu i razmenjivanje novih stavki unazad u sortirni deo sve dok ne stigne do kraja. Ovaj algoritam ima kvadratnu vremensku složenost u prosečnim i najgorem slučaju.

**Cilj:** Ispunite priče korisnika ispod i prođite sve testove kako biste završili laboratorijum.

**Priče korisnika:**

1. Treba da kreirate funkciju `insertionSort`.
1. Funkcija `insertionSort` treba da primi niz celih brojeva i vrati niz sa istim celim brojevima sortiranima od najmanjeg do najvećeg.

# --hints--

`insertionSort` should be a function.

```js
assert.isFunction(insertionSort);
```

`insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` should return an array that is unchanged except for order.

```js
assert.sameMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`insertionSort` should return a sorted array (least to greatest).

```js

assert.sameOrderedMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`insertionSort([5, 4, 33, 2, 8])` should return `[2, 4, 5, 8, 33]`.

```js
assert.sameOrderedMembers(insertionSort([5, 4, 33, 2, 8]), [2, 4, 5, 8, 33])
```

`insertionSort` should not use the built-in `.sort()` method.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    insertionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function insertionSort (array) {
  for (let currentIndex = 0; currentIndex < array.length; currentIndex++) {
    let current = array[currentIndex];
    let j = currentIndex - 1;
    while (j > -1 && array[j] > current) {
      array[j + 1] = array[j];
      j--;
    }
    array[j + 1] = current;
  }
  return array;
}
```
