---
id: 5f332a88dc25a0fd25c7687a
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Ime kafe je `CAMPER CAFE`. Stoga dodajte element `h1` unutar vašeg `main` elementa. Dajte mu ime kafe velikim slovima kako bi se istaklo.

# --hints--

You should have an opening `<h1>` tag.

```js
assert.match(code, /<h1>/i);
```

You should have a closing `</h1>` tag.

```js
assert.match(code, /<\/h1>/i);
```

You should not change your existing `main` element.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Your `h1` element should be nested in your `main` element.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

Your `h1` element should have the text `CAMPER CAFE`.

```js
assert.equal(document.querySelector("h1")?.innerText.trim(), "CAMPER CAFE");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```

