---
id: 5f33294a6af5e9188dbdb8f3
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovom radionici, veštaćete se osnovama CSS-a (Cascading Style Sheets) kreiranjem menija kafe.

Počnite dodavanjem nekog sadržaja menija. Dodajte element `main` unutar postojećeg elementa `body`. Na kraju će sadržati informacije o cenama za kafu i desertne specijalitete koje kafe nudi.

# --hints--

Your code should have an opening `<main>` tag.

```js
assert.match(code, /<main>/i);
```

Your code should have a closing `</main>` tag.

```js
assert.match(code, /<\/main>/i);
```

You should not change your `body` element. Make sure you don't accidentally delete your closing tag.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

Your `main` tag should be within your `body` tag.

```js
const main = document.querySelector('main');
assert.equal(main.parentElement.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
