---
id: 5f344f9c805cd193c33d829c
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

U prethodnim lekcijama, naučili ste kako da dodate `CSS` svojstva i vrednosti kao što je ovo:

```css
element {
 property: value;
}
```

Centrirajte sadržaj elementa `h1` postavljanjem njegovog svojstva `text-align` na vrednost `center`.

# --hints--

You should have an `h1` selector in your `style` element.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert.exists(hasSelector);
```

Your `text-align` property should have a value of `center`.

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert.isTrue(hasTextAlign);
```

Your `h1` selector should set the `text-align` property to `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert.equal(h1TextAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
--fcc-editable-region--
      
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

