---
id: 5f3477aefa51bfc29327200b
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Stilizovali ste tri elementa pišući CSS unutar tagova `style`. Ovo funkcioniše, ali pošto će biti mnogo više stilova, najbolje je da sve stilove stavite u zaseban fajl i povežete na njega.

Za vas je kreiran zaseban `styles.css` fajl. Možete prebacivati između fajlova pomoću kartica na vrhu uređivača.

Počnite prepisujući stilove koje ste kreirali u `styles.css` fajl. Uverite se da isključujete otvarajuće i zatvarajuće `style` tagove.

# --hints--

Your `styles.css` file should have the `h1, h2, p` type selector.

```js
assert.match(code.slice(code.indexOf('</html>')).replace(/[\s\n]*/g, ""), __helpers.concatRegex(__helpers.permutateRegex(['h1', 'h2', 'p'], { elementsSeparator: ',' }), /{/));
```
 
Your selector should set the `text-align` property to `center`.

```js
assert.match(code.slice(code.indexOf('</html>')), /text-align:\s*center;?/);
```

Your code should not contain the `<style>` tags.

```js
assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
