---
id: 5f35e5c4321f818cdc4bed30
title: Korak 25
challengeType: 0
dashedName: step-25
---

# --description--

Pošto sada sve izgleda dobro, vreme je da počnemo dodavati stavke menija.

Dodajte prazan element ``article`` ispod naslova ``Coffee``. On će sadržati ukus i cenu svake kafe koju trenutno nudite.

# --hints--

You should have an opening `<article>` tag.

```js
assert.match(code,/<article>/i);
```

You should have a closing `</article>` tag.

```js
assert.match(code, /<\/article>/i);
```

You should not change the existing `h2` element.

```js
assert.lengthOf(document.querySelectorAll('h2'), 1);
```

Your `article` element should come after your `h2` element.

```js
const article = document.querySelector('article');
assert.equal(article?.previousElementSibling?.tagName, 'H2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

