---
id: 5e44414f903586ffb414c950
title: Napravi kalkulator verovatnoce
challengeType: 23
saveSubmissionToDB: true
forumTopicId: 462364
dashedName: build-a-probability-calculator-project
---

# --description--

Zamisli da postoji sesir sa 5 plavih loptica, 4 crvene i 2 zelene. Kolika je verovatnoca da pri nasumicnom izvlacenju 4 loptice dobijes bar 1 crvenu i 2 zelene? Iako je moguce izracunati verovatnocu naprednom matematikom, jednostavniji pristup je da napises program koji ce izvesti mnogo eksperimenata i proceniti verovatnocu.

U ovom projektu ces napisati program koji procenjuje verovatnocu izvlacenja odredjenih loptica iz sesira.

Prvo napravi klasu `Hat` u fajlu `main.py`. Klasa treba da prima proizvoljan broj argumenata koji odredjuju broj loptica svake boje u sesiru. Na primer, objekat klase moze da se kreira na bilo koji od sledecih nacina:

```py
hat1 = Hat(yellow=3, blue=2, green=6)
hat2 = Hat(red=5, orange=4)
hat3 = Hat(red=5, orange=4, black=1, blue=0, pink=2, striped=9)
```

Sesir ce uvek imati bar jednu lopticu. Argumenti prosledjeni pri kreiranju objekta sesira treba da se pretvore u atribut instance `contents`. `contents` treba da bude lista stringova sa po jednim elementom za svaku lopticu u sesiru. Svaki element liste treba da bude naziv boje koji predstavlja jednu lopticu te boje. Na primer, ako je sesir `{'red': 2, 'blue': 1}`, onda `contents` treba da bude `['red', 'red', 'blue']`.

Klasa `Hat` treba da ima metodu `draw` koja prima argument sa brojem loptica koje treba izvuci iz sesira. Metoda treba nasumicno da ukloni loptice iz `contents` i da ih vrati kao listu stringova. Izvucene loptice se ne vracaju u sesir tokom izvlacenja, kao kod uzorkovanja bez vracanja. Ako je trazeni broj veci od broja dostupnih loptica, vrati sve loptice.

Zatim napravi funkciju `experiment` u `main.py` (van klase `Hat`). Ova funkcija treba da prima sledece argumente:

- `hat`: Objekat sesira sa lopticama koji se koristi u eksperimentu.
- `expected_balls`: Objekat koji opisuje tacnu grupu loptica koju pokusavas da izvuces u eksperimentima. Na primer, da procenis verovatnocu izvlacenja 2 plave i 1 crvene loptice, postavi `expected_balls` na `{'blue':2, 'red':1}`.
- `num_balls_drawn`: Broj loptica koje se izvlace iz sesira u svakom eksperimentu.
- `num_experiments`: Broj eksperimenata koji se izvode. (Sto je broj eksperimenata veci, procena verovatnoce je tacnija.)

Funkcija `experiment` treba da vrati verovatnocu.

Na primer, ako zelis da procenis verovatnocu da dobijes bar dve crvene i jednu zelenu lopticu kada izvlacis pet loptica iz sesira sa sest crnih, cetiri crvene i tri zelene. Uradices `N` eksperimenata, izbrojati koliko puta `M` dobijes bar dve crvene i jednu zelenu, pa proceniti verovatnocu kao `M/N`. Svaki eksperiment podrazumeva pocetak sa zadatim sesirom, izvlacenje odredjenog broja loptica i proveru da li su dobijene trazene loptice.

Ovako bi pozvao funkciju `experiment` za gornji primer sa 2000 eksperimenata:

```py
hat = Hat(black=6, red=4, green=3)
probability = experiment(hat=hat,
                  expected_balls={'red':2,'green':1},
                  num_balls_drawn=5,
                  num_experiments=2000)
```

Rezultat ce biti slican ovom:

```bash
0.356
```

Posto je proces nasumican, rezultat verovatnoce ce se blago razlikovati pri svakom pokretanju.

_Napomena: Razmisli o koriscenju modula koji su vec importovani na vrhu. Nemoj postavljati random seed unutar fajla._

Napomena: otvori konzolu pregledaca pomocu F12 da vidis detaljnije rezultate testova.

# --hints--

Kreiranje objekta `hat` treba da doda odgovarajuci sadrzaj.

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_class_contents(self):
        hat = probability_calculator.Hat(red=3,blue=2)
        actual = hat.contents
        expected = ["red","red","red","blue","blue"]
        self.assertEqual(actual, expected, 'Expected creation of hat object to add correct contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

Metoda `draw` u klasi `hat` treba da smanji broj elemenata u sadrzaju.


```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)

class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_draw(self):
        hat = probability_calculator.Hat(red=5,blue=2)
        actual = hat.draw(2)
        expected = ['blue', 'red']
        self.assertEqual(actual, expected, 'Expected hat draw to return two random items from hat contents.')
        actual = len(hat.contents)
        expected = 5
        self.assertEqual(actual, expected, 'Expected hat draw to reduce number of items in contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

Metoda `draw` treba ispravno da radi i kada je trazeni broj loptica veci od broja loptica u sesiru.

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_draw_2(self):
        hat = probability_calculator.Hat(yellow=5,red=1,green=3,blue=9,test=1)
        actual = sorted(hat.draw(20))
        expected = sorted(['yellow', 'yellow', 'yellow', 'yellow', 'yellow', 'red', 'green', 'green', 'green', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'test'])
        self.assertEqual(actual, expected, 'Expected hat draw to return all items from hat contents.')
        actual = len(hat.contents)
        expected = 0
        self.assertEqual(actual, expected, 'Expected hat draw to leave no items in contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

Metoda `experiment` treba da vrati odgovarajucu verovatnocu.


```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_prob_experiment(self):
        hat = probability_calculator.Hat(blue=3,red=2,green=6)
        probability = probability_calculator.experiment(hat=hat, expected_balls={"blue":2,"green":1}, num_balls_drawn=4, num_experiments=1000)
        actual = probability
        expected = 0.272
        self.assertAlmostEqual(actual, expected, delta = 0.01, msg = 'Expected experiment method to return a different probability.')
        hat = probability_calculator.Hat(yellow=5,red=1,green=3,blue=9,test=1)
        probability = probability_calculator.experiment(hat=hat, expected_balls={"yellow":2,"blue":3,"test":1}, num_balls_drawn=20, num_experiments=100)
        actual = probability
        expected = 1.0
        self.assertAlmostEqual(actual, expected, delta = 0.01, msg = 'Expected experiment method to return a different probability.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

# --seed--

## --seed-contents--

```py
import copy
import random

class Hat:
    pass

def experiment(hat, expected_balls, num_balls_drawn, num_experiments):
    pass
```

# --solutions--

```py
import copy
import random

class Hat:
    def __init__(self, **hat):
        self.hat = hat
        contents = []
        for i in hat:
            for j in range(hat[i]):
                contents.append(i)           
        self.contents = contents
               

    def draw(self, number): 
        drawn = []
        if number >= len(self.contents):
            drawn.extend(self.contents)
            self.contents = []
        else:
            for i in range(number):
                drawn.append(
                    self.contents.pop(random.randrange(len(self.contents)))
                )                
        return drawn

def experiment(hat, expected_balls, num_balls_drawn, num_experiments):    
    expected_balls_list = []
    drawn_list = []
    success = 0
    for i in expected_balls:
        for j in range(expected_balls[i]):
            expected_balls_list.append(i)
    for j in range(num_experiments):
        hat_copy = copy.deepcopy(hat)
        drawn_list.append(hat_copy.draw(num_balls_drawn))        
        exp_ball_list_copy = expected_balls_list[:]
        for k in range(len(drawn_list[j])):
            try:
                ind = exp_ball_list_copy.index(drawn_list[j][k])
                exp_ball_list_copy.pop(ind)
            except:
                continue
        if len(exp_ball_list_copy) == 0:
            success += 1        
    probability = success/num_experiments    
    return probability
```
