id: 587d7b85367417b2b2512b3a
title: Hvatanje argumenata prosleđenih u pogrešnom redosledu prilikom pozivanja funkcije
challengeType: 1
forumTopicId: 301184
dashedName: catch-arguments-passed-in-the-wrong-order-when-calling-a-function
---

# --description--

Nastavljajući sa diskusijom o pozivanju funkcija, sledeća greška koju treba uzeti u obzir je kada se argumenti funkcije proslede u pogrešnom redosledu. Ako su argumenti različitih tipova, kao što je funkcija koja očekuje niz podataka i ceo broj, to može rezultirati greškom tokom izvršenja. Ako su argumenti istog tipa (kao što su svi ceo brojevi, na primer), onda logika koda neće imati smisla. Uverite se da pružate sve potrebne argumente u ispravnom redosledu kako biste izbegli ove probleme.

# --instructions--

Funkcija `raiseToPower` uzima osnovni broj i podiže ga na stepen. Nažalost, ona nije pozvana ispravno - popravite kod tako da vrednost `power` bude 8 kako je očekivano.

# --hints--

Vaš kod bi trebalo da popravi parametar `power` tako da bude jednak 2 podignutom na stepen 3, a ne 3 podignutom na stepen 2.

```js
assert(power == 8);
```

Msimbo wako unapaswa kutumia mpangilio sahihi wa hoja kwa wito wa kitendakazi `raiseToPower`.

```js
assert(__helpers.removeJSComments(code).match(/raiseToPower\(\s*?base\s*?,\s*?exp\s*?\);/g));
```

# --seed--

## --seed-contents--

```js
function raiseToPower(b, e) {
  return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(exp, base);
console.log(power);
```

# --solutions--

```js
function raiseToPower(b, e) {
 return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(base, exp);
console.log(power);
```