id: 587d7b86367417b2b2512b3b
title: Pronađite greške "off by one" prilikom korišćenja indeksiranja
challengeType: 1
forumTopicId: 301189
dashedName: catch-off-by-one-errors-when-using-indexing
---

# --description--

<dfn>Off-by-one greške</dfn> (ponekad nazivane OBOE) nastaju kada pokušavate da pristupite specifičnom indeksu nizu karaktera ili nizu podataka (sa rezanjem ili dohvaćanjem dela), ili kada gledate na indeks tih elemenata. Indeksi u JavaScript-u počinju od nule, ne od jedinstva, što znači da je poslednji indeks uvek jedan manji od dužine elementa. Ako pokušate da pristupite indeksu koji odgovara dužini, program može baciti grešku "index out of range" ili ispisati `undefined`.

Kada koristite metode niza karaktera ili nizova podataka koje uzimaju indekse kao argumente, korisno je pročitati dokumentaciju i razumeti da li je uključeno (inclusive) ili ne. Evo primera off-by-one grešaka:

```js
let alphabet = "abcdefghijklmnopqrstuvwxyz";
let len = alphabet.length;
for (let i = 0; i <= len; i++) {
  console.log(alphabet[i]);
}
for (let j = 1; j < len; j++) {
  console.log(alphabet[j]);
}
for (let k = 0; k < len; k++) {
  console.log(alphabet[k]);
}
```

Prvi primer ovde se ponavlja jednom više nego što bi trebalo, a drugi se ponavlja malo manje (propušta prvi indeks, 0). Treći primer je ispravan.

# --instructions--

Ispravite dve greške indeksiranja u funkciji ispod kako bi sve brojevi od 1 do 5 bili ispisani u konzolu.

# --hints--

Vaš kod treba da postavi početno stanje petlje tako da počne sa prvim indeksom.

```js
assert(__helpers.removeJSComments(code).match(/i\s*?=\s*?0\s*?;/g).length == 1);
```

Vaš kod treba da ispravi početno stanje petlje tako da indeks petlje počne od 0.

```js
assert(!__helpers.removeJSComments(code).match(/i\s?=\s*?1\s*?;/g));
```

Vaš kod treba da postavi krajnje stanje petlje tako da se završi na poslednjem indeksu.

```js
assert(__helpers.removeJSComments(code).match(/i\s*<\s*len\s*;|i\s*<=\s*len\s*-\s*1\s*;/g).length == 1);
```

Vaš kod treba da ispravi krajnje stanje petlje tako da se završi jedan pre dužine.

```js
assert(!__helpers.removeJSComments(code).match(/i\s*?<=\s*?len;/g));
```

# --seed--

## --seed-contents--

```js
function countToFive() {
  let firstFive = "12345";
  let len = firstFive.length;
  // Only change code below this line
  for (let i = 1; i <= len; i++) {
  // Only change code above this line
    console.log(firstFive[i]);
  }
}

countToFive();
```

# --solutions--

```js
function countToFive() {
 let firstFive = "12345";
 let len = firstFive.length;
 // Only change code below this line
 for (let i = 0; i < len; i++) {
 // Only change code above this line
   console.log(firstFive[i]);
 }
}

countToFive();
```