id: 587d7b85367417b2b2512b39
title: Otkrivanje nedostatka zagrada za otvaranje i zatvaranje nakon poziva funkcije
challengeType: 1
forumTopicId: 301185
dashedName: catch-missing-open-and-closing-parenthesis-after-a-function-call
---

# --description--

Kada funkcija ili metoda ne prima argumente, možete zaboraviti da uključite obične zagrade za otvaranje i zatvaranje (prazne) kada je pozivate. Rezultat poziva funkcije često se sačuvava u promenljivoj za drugu upotrebu u vašem kodu. Ova greška može se otkriti upisivanjem vrednosti promenljivih (ili njihovog tipa) u konzolu i videćete da jedna predstavlja referencu na funkciju, umesto očekivane vrednosti koju bi funkcija trebalo da vrati.

Promenjive u sledećem primeru su različite:

```js
function myFunction() {
  return "You rock!";
}
let varOne = myFunction;
let varTwo = myFunction();
```

Ovde je `varOne` funkcija `myFunction`, a `varTwo` niz karaktera `You rock!`.

# --instructions--

Ispravite kod tako da promenljiva `result` dobije vrednost koju vraća poziv funkcije `getNine`.

# --hints--

Vaš kod treba da promeni promenljivu `result` tako da dobije broj koji funkcija `getNine` vraća.

```js
assert(result == 9);
```

Vaš kod treba da pozove funkciju `getNine`.

```js
assert(__helpers.removeJSComments(code).match(/getNine\(\)/g).length == 2);
```

# --seed--

## --seed-contents--

```js
function getNine() {
  let x = 6;
  let y = 3;
  return x + y;
}

let result = getNine;
console.log(result);
```

# --solutions--

```js
function getNine() {
 let x = 6;
 let y = 3;
 return x + y;
}

let result = getNine();
console.log(result);
```