id: 587d7b84367417b2b2512b37
title: Otkrij mešovano korišćenje jednostrukih i dvostrukih navodnika
challengeType: 1
forumTopicId: 301188
dashedName: catch-mixed-usage-of-single-and-double-quotes
---

# --description--

JavaScript dozvoljava korišćenje jednostrukih ('') i dvostrukih ("") navodnika za deklarisanje niza karaktera. Određivanje koji koristiti obično zavisi od ličnih preferencija, osim u nekim slučajevima.

Imati dve opcije je dobro kada niz karaktera sadrži skraćenice ili drugi deo teksta unutar navodnika. Pazite da ne zatvorite niz karaktera previše rano, jer to izaziva sintaksičku grešku.

Evo primera mešovitog korišćenja navodnika:

```js
const grouchoContraction = "I've had a perfectly wonderful evening, but this wasn't it.";
const quoteInString = "Groucho Marx once said 'Quote me as saying I was mis-quoted.'";
const uhOhGroucho = 'I've had a perfectly wonderful evening, but this wasn't it.';
```

Prva dva primera su tačna, ali treći nije.

Naravno, u redu je koristiti samo jedan stil navodnika. Možete izbeći navodnike unutar niza karaktera koristeći povratnik (`\`):

```js
const allSameQuotes = 'I\'ve had a perfectly wonderful evening, but this wasn\'t it.';
```

# --instructions--

Ispravite niz karaktera tako da koristi različite navodnike za vrednost `href`, ili izbegava postojeće navodnike. Omotajte dvostruke navodnike oko celog niza.

# --hints--

Vaš kod bi trebalo da popravi navodnike oko vrednosti `href` `#Home` menjanjem ili izbegavanjem.

```js
assert(__helpers.removeJSComments(code).match(/<a href=\s*?('|\\")#Home\1\s*?>/g));
```

Vaš kod bi trebalo da omotava dvostruke navodnike oko celog niza.

```js
assert(__helpers.removeJSComments(code).match(/"<p>.*?<\/p>";/g));
```

# --seed--

## --seed-contents--

```js
let innerHtml = "<p>Click here to <a href="#Home">return home</a></p>";
console.log(innerHtml);
```

# --solutions--

```js
let innerHtml = "<p>Click here to <a href=\"#Home\">return home</a></p>";
console.log(innerHtml);
```