id: 587d7b84367417b2b2512b36
title: 'Otkrivanje nezatvorenih zagradi, vitičastih zagrada, uglastih zagradi i navodnika'
challengeType: 1
forumTopicId: 301190
dashedName: catch-unclosed-parentheses-brackets-braces-and-quotes
---

# --description--

Još jedna sintaktička greška koju treba imati na umu je da sve obične zagrade, vitičaste zagrade, uglaste zagrade i navodnike moraju biti u paru. Zaboravljanje dela se dešava najčešće kada prepravljate postojeći kod i unosite elemente iste vrste parova. Budite pažljivi i kada formatirate unutar drugih zagrada koda, kao što je dodavanje funkcije povratne vrednosti kao argumenta u način.

Jedan način za izbegavanje ove greške je da čim napišete otvarajući karakter, odmah uključite par zatvarajućih karaktera, a zatim stavite placeholder unutar njih i nastavite sa pisanjem koda. Srećom, većina modernih editora koda automatski generiše drugi deo para.

# --instructions--

Ispravite dve greške u parovima u kodu.

# --hints--

Vaš kod bi trebalo da popravi nedostajući deo niza podataka.

```js
assert(__helpers.removeJSComments(code).match(/myArray\s*?=\s*?\[\s*?1\s*?,\s*?2\s*?,\s*?3\s*?\];/g));
```

Vaš kod bi trebalo da popravi nedostajući deo načina `.reduce()`. Rezultat u konzoli bi trebalo da pokaže da je `Sum of array values is: 6`.

```js
assert(arraySum === 6);
```

# --seed--

## --seed-contents--

```js
let myArray = [1, 2, 3;
let arraySum = myArray.reduce((previous, current =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```

# --solutions--

```js
let myArray = [1, 2, 3];
let arraySum = myArray.reduce((previous, current) =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```