id: 587d7b83367417b2b2512b33
title: Koristite JavaScript konzolu za proveru vrednosti promenljive
challengeType: 1
forumTopicId: 18372
dashedName: use-the-javascript-console-to-check-the-value-of-a-variable
---

# --description--

Svi Chrome i Firefox imaju odlične JavaScript konzole, poznate takođe kao alat za razvojno okruženje, za otklanjanje grešaka u vašem JavaScript-u.

Možete pronaći alati za razvojno okruženje u meniju Chrome ili u Mrežnoj konzoli (Developer Console) u meniju Firefoxa. Ako koristite drugi pretraživač ili mobilni telefon, veoma preporučujemo da pređete na desktop verzije Firefox-a ili Chrome-a.

Metoda `console.log()`, koja "ispisuje" rezultate onoga što se nalazi unutar njenih običnih zagrada, biće najvažniji alat za otklanjanje grešaka. Ubacivanje u ključne delove vašeg koda može vam pokazati srednje vrednosti promenljivih. Dobra praksa je da imate ideju o tome kakav rezultat bi trebalo da bude pre nego što proverite šta jeste. Imati sekcije za provere stanja vaše matematike u celom kodu pomoći će smanjiti mesto gde se greška nalazi.

Evo primera kako ispisati niz karaktera `Hello world!` u konzolu:

```js
console.log('Hello world!');
```

# --instructions--

Koristite metodu `console.log()` za ispisivanje vrednosti promenljive `a` tamo gde je prikazana u kodu.

# --hints--

Vaš kod bi trebalo da koristi `console.log()` za proveru vrednosti promenljive `a`.

```js
assert(__helpers.removeJSComments(code).match(/console\.log\(a\)/g));
```

# --seed--

## --seed-contents--

```js
let a = 5;
let b = 1;
a++;
// Only change code below this line


let sumAB = a + b;
console.log(sumAB);
```

# --solutions--

```js
var a = 5; console.log(a);
```