---
id: 587d7b84367417b2b2512b34
title: Koristite typeof za proveru tipa varijable
challengeType: 1
forumTopicId: 18374
dashedName: use-typeof-to-check-the-type-of-a-variable
---

# --description--

Možete koristiti `typeof` za proveru strukture podataka, ili tipa varijable. Ovo je važno prilikom otklanjanja grešaka kada koristite različite tipove podataka. Ako mislite da sabirate dva broja, ali jedan je niz znakova, rezultat može biti nepredvidiv. Greške u tipovima mogu se pojaviti u proračunima ili pozivima funkcija. Budite posebno pažljivi kada radite sa podacima iz eksternih izvora u formatu JavaScript Object Notation (JSON).

Evo primera koji koriste `typeof`:

```js
console.log(typeof "");
console.log(typeof 0);
console.log(typeof []);
console.log(typeof {});
```

Redosledom, konzola će prikazati niz znakova `string`, `number`, `object` i `object`.

JavaScript prepoznaje sedam osnovnih (nepromenljivih) tipova podataka: `Boolean`, `Null`, `Undefined`, `Number`, `String`, `Symbol` (noviji u ES6), i `BigInt` (noviji u ES2020), kao i jedan tip za promenljive objekte: `Object`. Zapamtite da su nizovi podataka tehnički tip objekta u JavaScriptu.

# --instructions--

Dodajte dve naredbe `console.log()` za proveru `typeof` dve varijable, `seven` i `three`, u kod.

# --hints--

Vaš kod bi trebalo da koristi `typeof` u dve naredbe `console.log()` za proveru tipa varijabli.

```js
assert(__helpers.removeJSComments(code).match(/console\.log\s*\(typeof[\( ].*\)?\)/g).length == 2);
```

Vaš kod bi trebalo da koristi `typeof` za proveru tipa varijable `seven`.

```js
assert(__helpers.removeJSComments(code).match(/typeof[\( ]seven\)?/g));
```

Vaš kod bi trebalo da koristi `typeof` za proveru tipa varijable `three`.

```js
assert(__helpers.removeJSComments(code).match(/typeof[\( ]three\)?/g));
```

# --seed--

## --seed-contents--

```js
let seven = 7;
let three = "3";
console.log(seven + three);
// Only change code below this line
```

# --solutions--

```js
let seven = 7;let three = "3";console.log(typeof seven);
console.log(typeof three);
```