id: 587d7b86367417b2b2512b3c
title: Koristite oprez prilikom ponovnog inicijalizovanja varijabli unutar petlje
challengeType: 1
forumTopicId: 301194
dashedName: use-caution-when-reinitializing-variables-inside-a-loop
---

# --description--

Ponekad je važno sačuvati informacije, povećati brojač ili ponovo inicijalizovati varijable unutar petlje. Problem može nastati kada bi varijable trebalo da se ponovo inicijalizuju, ali to ne rade, ili obrnuto. Ovo je posebno opasno ako slučajno ponovite varijablu koja se koristi za konačno stanje, što dovodi do beskonačne petlje.

Ispisivanje vrednosti varijabli u svakoj iteraciji pomoću `console.log()` može otkriti greškovo ponašanje povezano sa ponovnim inicijalizovanjem ili neuspešno ponovno inicijalizovanje varijable.

# --instructions--

Funkcija ovde treba da kreira dvodimenzionalni niz sa m redova i n kolona nula. Nažalost, ne vraća očekivani rezultat jer varijabla `row` nije ponovo inicijalizovana (resetovana na prazan niz) u spoljnjoj petlji. Ispravite kod tako da vrati ispravan dvodimenzionalni niz od 3x2 nula, koji izgleda kao `[[0, 0], [0, 0], [0, 0]]`.

# --hints--

Vaš kod treba da postavi varijablu `matrix` kao dvodimenzionalni niz sa 3 reda i po 2 kolone nula.

```js
assert(JSON.stringify(matrix) == '[[0,0],[0,0],[0,0]]');
```

Varijabla `matrix` treba da ima 3 reda.

```js
assert(matrix.length == 3);
```

Varijabla `matrix` treba da ima po 2 kolone u svakom redu.

```js
assert(
  matrix[0].length == 2 && matrix[1].length === 2 && matrix[2].length === 2
);
```

`zeroArray(4,3)` treba da vrati dvodimenzionalni niz sa 4 reda i po 3 kolone nula.

```js
assert(JSON.stringify(zeroArray(4,3)) == '[[0,0,0],[0,0,0],[0,0,0],[0,0,0]]');
```

# --seed--

## --seed-contents--

```js
function zeroArray(m, n) {
  // Creates a 2-D array with m rows and n columns of zeroes
  let newArray = [];
  let row = [];
  for (let i = 0; i < m; i++) {
    // Adds the m-th row into newArray

    for (let j = 0; j < n; j++) {
      // Pushes n zeroes into the current row to create the columns
      row.push(0);
    }
    // Pushes the current row, which now has n zeroes in it, to the array
    newArray.push(row);
  }
  return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```

# --solutions--

```js
function zeroArray(m, n) {
 // Creates a 2-D array with m rows and n columns of zeroes
 let newArray = [];
 for (let i = 0; i < m; i++) {
   let row = [];
   // Adds the m-th row into newArray

   for (let j = 0; j < n; j++) {
     // Pushes n zeroes into the current row to create the columns
     row.push(0);
   }
   // Pushes the current row, which now has n zeroes in it, to the array
   newArray.push(row);
 }
 return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```