---
id: 587d7fb1367417b2b2512bf1
title: Servirajte JSON na specifičnoj ruti
challengeType: 2
forumTopicId: 301517
dashedName: serve-json-on-a-specific-route
---

# --description--

Dok HTML server služi HTML-om, API služi podacima. <dfn>REST</dfn> API (Representational State Transfer) omogućava jednostavnu razmenu podataka bez klijenta koji mora da zna bilo šta o serveru. Klijent treba samo da zna gde je resurs (URL) i šta želi da uradi sa njim (metoda). GET metoda se koristi kada preuzimate informacije bez menjanja bilo čega. Danas, preporučeni format za slanje podataka preko interneta je JSON. Ukratko, JSON je jednostavan način za predstavljanje JavaScript objekta kao niza karaktera, kako bi se mogao lako preneti.

Napravimo jednostavan API tako što ćemo kreirati rutu koja odgovara sa JSON-om na ruti `/json`. Možete to uraditi na uobičajeni način, koristeći `app.get()`. Unutar handlera rute, koristite `res.json()`, prosleđujući objekat kao argument. Ova metoda završava zahtev-odgovor ciklus, vraćajući podatke. Iza kulisa, konvertuje JavaScript objekat u niz karaktera, zatim postavlja odgovarajuća zaglavlja koja obaveštavaju vaš pregledač da koristite JSON, i šalje podatke nazad. Objekat ima standardni format `{key: data}`. `data` može biti broj, niz karaktera, ugnježdeni objekat ili niz podataka. `data` takođe može biti promenljiva ili rezultat poziva funkcije, koja će biti izračunata pre nego što se konvertuje u niz karaktera.

# --instructions--

Servirajte objekat `{"message": "Hello json"}` kao odgovor, u JSON formatu, na GET zahtev na ruti `/json`. Zatim usmerite vaš pregledač na `your-app-url/json`, trebali biste da vidite poruku na ekranu.

# --hints--

Ruta `/json` treba da servira JSON objekat `{"message": "Hello json"}`

```js
  const response = await fetch(code + '/json');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.message,
    'Hello json',
    "The '/json' endpoint does not serve the right data"
  );
```
