---
id: 68551b677a5f6898af364b85
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovoj radionici, naučićete više o Tailwind CSS-u korišćenjem njegovih utilitarnih klasa za flexbox kako biste izgradili responsivnu komponentu poziva na akciju (CTA).

Omogućena vam je osnovna HTML struktura, pa prvo morate napraviti HTML elemente.

Počnite kreiranjem elementa `div`. Unutar tog `div`, kreirajte još dva odvojena `div` elementa.

# --hints--

You should create a `div` element.

```js
assert.exists(document.querySelector("div"));
```

You should have two more `div` elements inside the `div` element.

```js
assert.lengthOf(document.querySelectorAll("div > div"), 2)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
