---
id: 6858ef974c8ad5dcd7abb9f1
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Sadržaj unutar `div` može se bolje razmakniti i organizovati. Za ostatak razmaka i tu organizaciju, morate koristiti flexbox.

Dodajte klasu `flex` na `div` kako bi dva elementa deteta `div` stajala jedan pored drugog.

Takođe, dodajte `justify-around` kako bi prostor oko i između sadržaja bio ravnomerno raspoređen, i `items-center` kako bi ceo sadržaj na kraju bio centriran u `div`.

# --hints--

Your outer `div` element should have the utility class `flex`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex"));
```

Your outer `div` element should have the utility class `justify-around`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("justify-around"));
```

Your outer `div` element should have the utility class `items-center`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
