---
id: 6858f90c54647a48ddb8822d
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

U ovom trenutku, sadržaj, to jest dva `div` elementa unutar spoljašnjeg `div`, biće složeni jedan na drugog bez obzira na veličinu ekrana.

To je zato što imate `flex-col` klasu u omotajućem `div` elementu. Da biste to popravili, koristite `lg:` prefiks da dodate klasu `flex-row` na `div`.

To znači da će na uređajima sa većim ekranima, fleksibilni smer biti `row`. Da biste ponovo videli to u akciji, prilagodite pretraživač većoj širini.

# --hints--

Your outer `div` element should have the utility class `lg:flex-row`.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("lg:flex-row"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
