---
id: 68e51d8f163d3f8b70f9c4e5
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, veštaćete se rada sa `includes()` i `slice()` metodama.

Da biste počeli, kreirajte varijablu pod imenom `fccSentence` i dodelite joj string `"freeCodeCamp is a great place to learn web development."`.

# --hints--

You should declare a variable named `fccSentence`.

```js
assert.exists(fccSentence);
```

Your `fccSentence` variable should be a string.

```js
assert.isString(fccSentence);
```

Your `fccSentence` variable should be assigned the string `"freeCodeCamp is a great place to learn web development."`.

```js
assert.strictEqual(fccSentence, "freeCodeCamp is a great place to learn web development.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
