---
id: 68e6172875008fea742a63c9
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Kao što ste se setili iz prethodnih lekcija, metoda `includes()` proverava da li niz sadrži određeni podniz i vraća `true` ili `false`.

Evo primera korišćenja metode `includes()`:

```js
const text = "Hello World";
const hasHello = text.includes("Hello");
console.log(hasHello); // true
```

Kreirajte varijablu pod imenom `hasFreeCodeCamp`. Zatim joj dodelite rezultat korišćenjem metode `includes()` na `fccSentence` kako biste proverili da li sadrži `"freeCodeCamp"`.

# --hints--

You should declare a variable named `hasFreeCodeCamp`.

```js
assert.exists(hasFreeCodeCamp);
```

You should use `fccSentence.includes("freeCodeCamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

You should assign the result of `fccSentence.includes("freeCodeCamp")` to your `hasFreeCodeCamp` variable.

```js
assert.match(__helpers.removeJSComments(code), /hasFreeCodeCamp\s*=\s*fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

The result of `fccSentence.includes("freeCodeCamp")` should be `true`. 

```js
assert.isTrue(hasFreeCodeCamp);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

--fcc-editable-region--

--fcc-editable-region--
```
