---
id: 68e6172975008fea742a63cc
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Metoda `includes()` je osetljiva na veličinu slova, što znači da razlikuje između velikih i malih slova.

Kreirajte varijablu pod imenom `hasLowercaseFCC`. Zatim joj dodelite rezultat korišćenjem metode `includes()` za proveru da li `fccSentence` sadrži `"freecodecamp"` (sve mala slova).

Zatim koristite ili *template literal* ili konkatenaciju niza kako biste u konzolu zapisali poruku `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.`. Zamenite `<hasLowercaseFCC>` stvarnom vrednošću varijable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

You should declare a variable named `hasLowercaseFCC`.

```js
assert.exists(hasLowercaseFCC);
```

You should use `fccSentence.includes("freecodecamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

You should assign the result of `fccSentence.includes("freecodecamp")` to your `hasLowercaseFCC` variable.

```js
assert.match(__helpers.removeJSComments(code), /hasLowercaseFCC\s*=\s*fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

The result of `fccSentence.includes("freecodecamp")` should be `false`. 

```js
assert.isFalse(hasLowercaseFCC);
```

You should log the message that includes `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.`, where `<hasLowercaseFCC>` should be replaced with the actual value of the variable.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freecodecamp") returns false because includes is case-sensitive.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
