---
id: ab306dbdcc907c7ddfc30830
title: Kreirajte alat za dubinsko ravaniranje
challengeType: 26
dashedName: create-a-deep-flattening-tool
---

# --description--

U ovom laboratorijumu, implementiraćete algoritam ravnanja niza.

Ravnanje niza znači pretvaranje ugnježdene niza bilo koje dubine u jedan, jednodimenzionalni niz. Proces izvlači sve elemente po redu, otpakujući samo nizove. Ostali tipovi ostaju nepromenjeni.

Primer:

| Original | Flattened |
| --- | --- |
| `[[1], [], [2, [3]]]` | `[1, 2, 3]` |
| `[1, {"foo": "bar"}, [2]]` | `[1, {"foo": "bar"}, 2]` |
| `["baz", [1, 2], {}]` | `["baz", 1, 2, {}]` |

**Cilj:** Ispunite priče korisnika ispod i učinite da svi testovi prođu kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da imate funkciju pod nazivom `steamrollArray`.
1. Funkcija `steamrollArray` treba da prihvati jedan argument: ugnježdeni niz.
1. Funkcija bi trebalo da izravna ugnježdeni niz, uzimajući u obzir različite nivoe ugnježdenosti.
1. Vaše rešenje ne bi trebalo da koristi metode `Array.prototype.flat()` ili `Array.prototype.flatMap()`.
1. Ne sme se koristiti globalne varijable.

# --hints--

You should have a `steamrollArray` function.

```js
assert.isFunction(steamrollArray);
```

`steamrollArray([[["a"]], [["b"]]])` should return `["a", "b"]`.

```js
assert.deepEqual(steamrollArray([[['a']], [['b']]]), ['a', 'b']);
```

`steamrollArray([1, [2], [3, [[4]]]])` should return `[1, 2, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, [2], [3, [[4]]]]), [1, 2, 3, 4]);
```

`steamrollArray([1, [], [3, [[4]]]])` should return `[1, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, [], [3, [[4]]]]), [1, 3, 4]);
```

`steamrollArray([1, {}, [3, [[4]]]])` should return `[1, {}, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4]);
```

Your solution should not use the `Array.prototype.flat()` or `Array.prototype.flatMap()` methods.

```js
const arr = [1,2,3,[1,2,3,[1,2,3]]];

const spyFlat = __helpers.spyOn(Array.prototype, 'flat');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlat.calls);
} finally {
  spyFlat.restore();
}

const spyFlatMap = __helpers.spyOn(Array.prototype, 'flatMap');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlatMap.calls);
} finally {
  spyFlatMap.restore();
}
```

You should not use global variables.

```js
steamrollArray([1, {}, [3, [[4]]]])
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4])
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function steamrollArray(arr) {
  if (!Array.isArray(arr)) {
    return [arr];
  }
  const out = [];
  arr.forEach(function(e) {
    steamrollArray(e).forEach(function(v) {
      out.push(v);
    });
  });
  return out;
}
```
