---
id: 694648acde178bb8202d9514
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Sada treba da dodate uzrast učenika u izveštaj o ocenama. Za to ćete koristiti ceo broj (integer), koji je jedan od numeričkih tipova podataka u Pythonu.

Deklarišite varijablu pod imenom `age` i dodelite joj ceo broj (integer) vrednost `20`.

Zatim, ispišite vrednost i tip podataka `age` zajedno, razdvojeno zarezom. Proverite izlaz u terminalu koji prikazuje vrednost `age`, i njegov tip kao `int` (integer).

# --hints--

You should have an `age` variable.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("age")`) })
```

The variable `age` should store the value `20`. Do not surround the value with quotes.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("age").is_equivalent("age = 20")`) })
```

You should print `age` and `type(age)` using a comma separator.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(age, type(age))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

--fcc-editable-region--

--fcc-editable-region--

```
