---
id: 694648acde178bb8202d9516
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Izlaz je `False`, što pokazuje da `score` nije `int`.

Drugi uobičajeni tip broja u Pythonu je `float`, koji predstavlja broj sa decimalima. Zamenite `int` sa `float` u postojećem pozivu `isinstance()` da biste ovo potvrdili.

# --hints--

You should replace `int` with `float` in the existing `isinstance(score, int)` call.

```js
({ test: () => runPython(`
  assert not _Node(_code).has_call("print(isinstance(score, int))")
  assert _Node(_code).has_call("print(isinstance(score, float))")
`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
--fcc-editable-region--
print(isinstance(score, int))
--fcc-editable-region--

```
