---
id: 694648acde178bb8202d9512
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Izvod bi takođe trebalo da pokaže da li je učenik trenutno upisan. Ovo se može predstaviti korišćenjem booleanske vrednosti.

Booleanske vrednosti predstavljaju uslov da/ne, i često se koriste za donošenje odluka u kodu. Postoje samo dve booleanske vrednosti: `True` i `False`.

Deklarišite varijablu pod imenom `is_student` i dodelite joj vrednost `True`.

# --hints--

You should have an `is_student` variable.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("is_student")`) })
```

The variable `is_student` should store the value `True`. Do not surround the value with quotes.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("is_student").is_equivalent("is_student = True")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

--fcc-editable-region--

--fcc-editable-region--
```
