---
id: 694648acde178bb8202d9515
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Sada, dodajte ocenu učenika.

Deklarišite varijablu pod imenom `score` i dodelite joj vrednost `80.5`.

Iako su i oba `age` i `score` brojevi, možda nisu iste vrste. Python pruža funkciju nazvanu `isinstance()` za proveru toga.

```py
x = 10
print(isinstance(x, int)) # Output: True
```

Koristite `isinstance()` da proverite da li je `score` `int`, i odštampajte rezultat u terminal kao što je prikazano u gornjem primeru.

# --hints--

You should have a `score` variable.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("score")`) })
```

The variable `score` should store the value `80.5`. Do not surround the value with quotes.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("score").is_equivalent("score = 80.5")`) })
```

You should print the result of `isinstance(score, int)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(isinstance(score, int))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

--fcc-editable-region--

--fcc-editable-region--

```
