---
id: 694648acde178bb8202d9513
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Ime studenta bi trebalo da prati isti format kao i ostali podaci.

Uklonite ranije izlazne podatke varijable `name`. Zatim, štampajte `name` i `type(name)` zajedno na jednoj liniji, razdvojene zarezom, kao u prethodnom koraku.

# --hints--

You should not have `print(name)` in your code.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(name)")`)})
```

You should not have `print(type(name))` in your code.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(type(name))")`)})
```

You should print `name` and `type(name)` in the same `print()` call using a comma separator.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name, type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--
print(name)
print(type(name))
--fcc-editable-region--

is_student = True
print(is_student, type(is_student))
```
