---
id: 69461e4f98594012939bc1c6
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovoj radionici, vežbaće tipove podataka Pythona kreiranjem jednostavnog štampača ocena.

Kao što ste naučili u prethodnim lekcijama, varijable se dodeljuju na ovaj način:

```py
greeting = 'Hello World'
```

U gornjem primeru, varijabla `greeting` ima vrednost niza (stringa), što je sekvenca karaktera okruženih navodnicima.

Kreirajte varijablu nazvanu `name` i dodelite joj niz (string) `Alice`. Zapamtite da vrednost okružujete ili jednostrukim ili dvostrukim navodnicima, kao što je prikazano u primeru.

**Važno:** Python koristi uvlačenje (razmaci na početku linije) za organizovanje koda. Za ovu radionicu, uverite se da nema dodatnih razmaka na početku svake linije koda. Dodavanje dodatnih razmaka će uzrokovati `IndentationError` i sprečiti vaš kod od pokretanja.

# --hints--

You should have a `name` variable.

```js
({ test: () => runPython(`assert _Node(_code).has_variable('name')`) })
```

The variable `name` should store the value `'Alice'`.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("name").is_equivalent("name = 'Alice'")`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
