---
id: 694648abde178bb8202d9511
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Trebalo bi da sada vidite ime studenta ispisano u terminalu.

Python pruža funkciju pod nazivom `type()` koju možete koristiti za proveru tipa vrednosti.

```py
platform = 'freeCodeCamp'
print(type(platform)) # Output: <class 'str'>
```

U gornjem primeru, izlaz `<class 'str'>` znači da je promenljiva prosleđena funkciji `type()` string.

Koristite funkciju `type()` sa `name` kao njenim argumentom i ispišite izlaz kao u primeru. Proverite izlaz u terminalu koji pokazuje da je `name` tipa `str` (string).

# --hints--

You should print the result of calling `type(name)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
--fcc-editable-region--

--fcc-editable-region--
```
