---
id: 6947a3e516bb8a4a681274a7
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Funkcija `print()` može prikazati više od jedne vrednosti u jednom trenutku. Razdvojte vrednosti zarezom (`,`) da ih ispišete na istoj liniji.

```py
subject = 'Python'
print(subject, type(subject)) # Output: Python <class 'str'>
```

Ispišite i `is_student` i `type(is_student)` na istoj liniji koristeći zarez `,` kao što je prikazano u primeru. Zatim, proverite izlaz u terminalu koji pokazuje vrednost `is_student`, i njegov tip kao `bool` (boolean).

# --hints--

You should print `is_student` and `type(is_student)` using a comma separator.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(is_student, type(is_student))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

is_student = True
--fcc-editable-region--

--fcc-editable-region--
```
