---
id: 6947a0c2b6de035dd8d3eea1
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Možete ispisati vrednost varijable koristeći funkciju `print()`.

```py
greeting = 'Hello World'
print(greeting) # Output: Hello World
```

Naučićete više o funkcijama u narednim lekcijama. Za sada, znajte da je funkcija ponovno upotrebljiv blok koda koji se može pozvati ili aktivirati kako bi se izvršio njegov kod, a argumente mu se može proslediti.

U gornjem primeru, `print(greeting)` je poziv funkcije, a `greeting` je argument te funkcije.

Pogledajte primer i ispišite varijablu `name`. Proverite izlaz u terminalu.

# --hints--

You should call `print()` with `name` as its argument.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name)")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--

--fcc-editable-region--
```
