---
id: 5a24c314108439a4d4036146
title: Mapa poziva dispečiranja po parametrima
challengeType: 6
forumTopicId: 301432
dashedName: map-dispatch-to-props
---

# --description--

Kitendakazi cha `mapDispatchToProps()` hutumika kutoa watengenezaji wa vitendo maalum kwa sehemu zako za React ili waweze kutuma vitendo dhidi ya hifadhi ya Redux. Ni sawa na muundo wa kitendakazi cha `mapStateToProps()` ulichokiandika katika changamoto iliyopita. Hurejesha kitu kinachoonyesha ramani ya vitendo vya dispatch kwa majina ya sifa, ambayo yanakuwa `props` za sehemu. Hata hivyo, badala ya kurejesha kipande cha `state`, kila sifa hurejesha kitendakazi kinachofanya wito wa `dispatch` kwa mtengenezaji wa kitendo na data yoyote muhimu ya kitendo. Una upatikanaji wa `dispatch` hii kwa sababu hupitishwa kwa `mapDispatchToProps()` kama kigezo unapoeleza kitendakazi, kama vile ulivyopitisha `state` kwa `mapStateToProps()`. Chini ya pazia, React Redux inatumia `store.dispatch()` ya Redux kufanya dispatch hizi kwa `mapDispatchToProps()`. Hii ni sawa na jinsi inavyotumia `store.subscribe()` kwa sehemu zinazopangwa kwa `state`.

Kwa mfano, una mtengenezaji wa kitendo cha `loginUser()` anayechukua `username` kama mzigo wa kitendo. Kitu kinachorejeshwa kutoka `mapDispatchToProps()` kwa mtengenezaji huyu wa kitendo kingekuwa kama ifuatavyo:

```jsx
{
  submitLoginUser: function(username) {
    dispatch(loginUser(username));
  }
}
```

# --instructions--

Kodni editor pruža fabriku akcija nazvanu `addMessage()`. Napisite funkciju za `mapDispatchToProps()` koja uzima `dispatch` kao argument, a zatim vraća objekat. Taj objekat treba da ima svojstvo `submitNewMessage` postavljeno za dispatch funkciju, koja uzima parametar za dodavanje nove poruke kada šalje `addMessage()`.

# --hints--

`addMessage` treba da vrati nešto sa ključevima `type` i `message`.

```js
assert(
  (function () {
    const addMessageTest = addMessage();
    return (
      addMessageTest.hasOwnProperty('type') &&
      addMessageTest.hasOwnProperty('message')
    );
  })()
);
```

`mapDispatchToProps` inapaswa kuwa kitendakazi.

```js
assert(typeof mapDispatchToProps === 'function');
```

`mapDispatchToProps` trebalo bi da vrati nešto.

```js
assert(typeof mapDispatchToProps() === 'object');
```

Slanje `addMessage` na `submitNewMessage` iz `mapDispatchToProps` treba da vrati poruku dispečerskoj funkciji.

```js
assert(
  (function () {
    let testAction;
    const dispatch = (fn) => {
      testAction = fn;
    };
    let dispatchFn = mapDispatchToProps(dispatch);
    dispatchFn.submitNewMessage('__TEST__MESSAGE__');
    return (
      testAction.type === 'ADD' && testAction.message === '__TEST__MESSAGE__'
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

// Change code below this line
```

# --solutions--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

// Change code below this line

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: function(message) {
      dispatch(addMessage(message));
    }
  }
};
```
