---
id: 5a24c314108439a4d4036147
title: Unganisha Redux na React
challengeType: 6
forumTopicId: 301426
dashedName: connect-redux-to-react
---

# --description--

Sada ste napisali funkciju od `mapStateToProps()` i `mapDispatchToProps()`, možete koristiti ovo da spojite `state` sa `dispatch` na `props` jedne od vaših komponenti React. Metoda `connect` iz Redux-a React može obraditi ovaj zadatak. Ova metoda prima dva opcionalna argumenta, `mapStateToProps()` i `mapDispatchToProps()`. To je opciono jer možete imati komponentu koja samo zahteva pristup `state`, ali ne mora da dispatch-uje akcije, ili obrnuto.

Da biste koristili ovu metodu, prosledite funkcije kao argumente, i odmah pozovite rezultat sa vašim delom. Ova sintaksa je malo jedinstvena i izgleda ovako:

```js
connect(mapStateToProps, mapDispatchToProps)(MyComponent)
```

**Napomena:** Ako želite da uklonite jedan od argumenata za putanju `connect`, koristite `null` umesto toga.

# --instructions--

Mhariri wa msimbo una vitendakazi vya `mapStateToProps()` na `mapDispatchToProps()` na sehemu mpya ya React iitwayo `Presentational`. Unganisha sehemu hii na Redux kwa njia ya `connect` kutoka kwenye kitu cha ulimwengu wa `ReactRedux`, na fanya wito mara moja kwa sehemu ya `Presentational`. Weka matokeo katika `const` mpya iitwayo `ConnectedComponent` inayowakilisha sehemu iliyounganishwa. Hiyo tu, sasa umeunganishwa na Redux! Jaribu kubadilisha moja ya hoja za `connect` kwa `null` na angalia matokeo ya majaribio.

# --hints--

Ovaj deo `Presentational` treba da prikaže.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('Presentational').length === 1;
  })()
);
```

Sehemu ya `Presentational` inapaswa kupokea vigezo `messages` kupitia `connect`.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const props = mockedComponent.find('Presentational').props();
    return props.messages === '__INITIAL__STATE__';
  })()
);
```

Sekcija `Presentational` treba da primi parametre `submitNewMessage` preko `connect`.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const props = mockedComponent.find('Presentational').props();
    return typeof props.submitNewMessage === 'function';
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message));
    }
  }
};

class Presentational extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h3>This is a Presentational Component</h3>
  }
};

const connect = ReactRedux.connect;
// Change code below this line
```

# --solutions--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message));
    }
  }
};

class Presentational extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h3>This is a Presentational Component</h3>
  }
};

const connect = ReactRedux.connect;
// Change code below this line

const ConnectedComponent = connect(mapStateToProps, mapDispatchToProps)(Presentational);
```
