---
id: 5a24c314108439a4d4036145
title: Mapa stanja po parametrima
challengeType: 6
forumTopicId: 301433
dashedName: map-state-to-props
---

# --description--

Deo `Provider` vam omogućava da pružite `state` i `dispatch` za vaše sekcije od React, ali morate tačno definisati koje uslove i akcije želite. Na ovaj način osiguravate da svaka sekcija ima pristup samo neophodnom stanju. Postižete ovo kreiranjem dve funkcije: `mapStateToProps()` i `mapDispatchToProps()`.

U ovim zadacima, deklariš koji tip stanja želiš da pristupiš i koje praktične komponente moraš da pošalješ. Kada su ti zadaci gotovi, videćeš kako da koristiš React Redux `connect` da ih povežeš sa svojim delovima u drugom izazovu.

**Kumbuka:** Chini ya pazia, React Redux hutumia njia ya `store.subscribe()` kutekeleza `mapStateToProps()`.

# --instructions--

Unda kitendakazi `mapStateToProps()`. Kitendakazi hiki kinapaswa kuchukua `state` kama hoja, kisha kurudisha kitu kinachoonyesha ramani ya hali hiyo kwa majina maalum ya sifa. Sifa hizi zitakuwa na upatikanaji kwa sehemu yako kupitia `props`. Kwa kuwa mfano huu unahifadhi hali yote ya programu katika safu moja ya data, unaweza kupitisha hali yote hiyo kwa sehemu yako. Unda sifa `messages` katika kitu kinachorudishwa, na uiweke kuwa `state`.

# --hints--

const `state` Trebalo bi da bude prazan red.

```js
assert(Array.isArray(state) && state.length === 0);
```

`mapStateToProps` Trebalo bi da bude funkcija.

```js
assert(typeof mapStateToProps === 'function');
```

`mapStateToProps` inapaswa kurudisha kitu.

```js
assert(typeof mapStateToProps() === 'object');
```

Kupitisha safu kama hali kwa `mapStateToProps` inapaswa kurudisha safu hii ikiwa imepewa thamani kwa ufunguo wa `messages`.

```js
assert(mapStateToProps(['messages']).messages.pop() === 'messages');
```

# --seed--

## --seed-contents--

```jsx
const state = [];

// Change code below this line
```

# --solutions--

```jsx
const state = [];

// Change code below this line

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};
```
