---
id: 689f7996426c1534fa5bea66
title: Izgradite prevodilac golf rezultata
challengeType: 26
dashedName: build-a-golf-score-translator
---

# --description--

U igri Golfa, svaka rupa ima `par`, što znači prosečan broj `strokes` koji se očekuje da golfer napravi kako bi potopio loptu u rupu i završio igru. U zavisnosti od toga koliko su vaše `strokes` iznad ili ispod `par`, postoji različit nadimak.

U ovom laboratorijumu, napisaćete funkciju koja konvertuje `par` i `strokes` u njihov nadimak. 

**Cilj:** Ispunite korisničke priče ispod i prođite sve testove kako biste završili laboratorijum.

**Korisničke priče:**

1. Treba da kreirate funkciju pod nazivom `golfScore`.
1. `golfScore` treba da primi dva numerička argumenta, koji su par za teren i broj napravljenih poteza.
1. `golfScore` treba da vrati string.
1. `golfScore` treba da vrati `"Hole-in-one!"` ako je `strokes` jednak `1`. 
1. `golfScore` treba da vrati `"Eagle"` ako je `strokes` manji ili jednak `par` minus `2`.
1. `golfScore` treba da vrati `"Birdie"` ako je `strokes` jednak `par` minus `1`.
1. `golfScore` treba da vrati `"Par"` ako je `strokes` jednak `par`. 
1. `golfScore` treba da vrati `"Bogey"` ako je `strokes` jednak `par` plus `1`. 
1. `golfScore` treba da vrati `"Double Bogey"` ako je `strokes` jednak `par` plus `2`. 
1. `golfScore` treba da vrati `"Go Home!"` ako je `strokes` veći ili jednak `par` plus `3`.

# --hints--

You should create a function named `golfScore`.

```js
assert.isFunction(golfScore);
```

`golfScore` should take two parameters.

```js
assert.lengthOf(golfScore, 2);
```

`golfScore` should return a string

```js
assert.isString(golfScore(4,1)); 
```

`golfScore(1, 1)` should return the string `Hole-in-one!`

```js
assert.strictEqual(golfScore(1,1), 'Hole-in-one!');
```

`golfScore(3, 1)` should return the string `Hole-in-one!`

```js
assert.strictEqual(golfScore(3,1), 'Hole-in-one!');
```

`golfScore(4, 1)` should return the string `Hole-in-one!`

```js
assert.strictEqual(golfScore(4,1), 'Hole-in-one!');
```

`golfScore(5, 1)` should return the string `Hole-in-one!`

```js
assert.strictEqual(golfScore(5,1), 'Hole-in-one!');
```

`golfScore(4, 2)` should return the string `Eagle`

```js
assert.strictEqual(golfScore(4,2), 'Eagle');
```

`golfScore(5, 2)` should return the string `Eagle`

```js
assert.strictEqual(golfScore(5,2), 'Eagle');
```

`golfScore(3, 2)` should return the string `Birdie`

```js
assert.strictEqual(golfScore(3,2), 'Birdie');
```

`golfScore(4, 3)` should return the string `Birdie`

```js
assert.strictEqual(golfScore(4,3), 'Birdie');
```

`golfScore(5, 4)` should return the string `Birdie`

```js
assert.strictEqual(golfScore(5,4), 'Birdie');
```

`golfScore(3, 3)` should return the string `Par`

```js
assert.strictEqual(golfScore(3,3), 'Par');
```

`golfScore(4, 4)` should return the string `Par`

```js
assert.strictEqual(golfScore(4,4), 'Par');
```

`golfScore(5, 5)` should return the string `Par`

```js
assert.strictEqual(golfScore(5,5), 'Par');
```

`golfScore(3, 4)` should return the string `Bogey`

```js
assert.strictEqual(golfScore(3,4), 'Bogey');
```

`golfScore(4, 5)` should return the string `Bogey`

```js
assert.strictEqual(golfScore(4,5), 'Bogey');
```

`golfScore(5, 6)` should return the string `Bogey`

```js
assert.strictEqual(golfScore(5,6), 'Bogey');
```

`golfScore(3, 5)` should return the string `Double Bogey`

```js
assert.strictEqual(golfScore(3,5), 'Double Bogey');
```

`golfScore(4, 6)` should return the string `Double Bogey`

```js
assert.strictEqual(golfScore(4,6), 'Double Bogey');
```

`golfScore(5, 7)` should return the string `Double Bogey`

```js
assert.strictEqual(golfScore(5,7), 'Double Bogey');
```

`golfScore(3, 7)` should return the string `Go Home!`

```js
assert.strictEqual(golfScore(3,7), 'Go Home!');
```

`golfScore(4, 8)` should return the string `Go Home!`

```js
assert.strictEqual(golfScore(4,8), 'Go Home!');
```

`golfScore(5, 9)` should return the string `Go Home!`

```js
assert.strictEqual(golfScore(5,9), 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
