---
id: 699f1d72e906e4616155290d
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, nastavićete da učite o tipskim anotacijama i kako rade sa funkcijama, razvijajući siguran matematički alat.

Počnite kreiranjem funkcije nazvane `square` koja prima parametar i vraća kvadrat broja.

# --hints--

You should have a function called `square`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.square);
```

Your `square` function should have one parameter.

```js
const explorer = await __helpers.Explorer(code);
assert.lengthOf(explorer.allFunctions.square.parameters, 1);
```

Your `square` function should return the square of a number.

```js
assert.strictEqual(square(2), 4);
assert.strictEqual(square(3), 9);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
