---
id: 699f5063d2393c674b1f45dc
title: Korak 16
challengeType: 1
dashedName: step-16
---

# --description--

Kao što ste se setili iz prethodnog koraka, ako pozovete funkciju sa pogrešnim brojem argumenata, TypeScript će baciti grešku.

Postoji način da označite parametre kao opcione dodavanjem `?` nakon imena parametra. To vam omogućava da pozovete funkciju sa manje argumenata nego parametara.

Evo primera:

```ts
function greet(name: string, greeting?: string) {
  if (greeting) {
    return `${greeting}, ${name}!`;
  }
  return `Hello, ${name}!`;
}
```

U ovom primeru, ako se parametar `greeting` ne pruži prilikom pozivanja funkcije `greet`, biće `undefined`, i funkcija će vratiti podrazumevano pozdravljanje.

Ažurirajte svoju funkciju `raiseTo` da učinite parametar `exponent` opcionalnim.

# --hints--

Your `raiseTo` function should have the `exponent` parameter marked as optional.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[1].matches("exponent?: number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(2, 3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
