---
id: 699f38937b053e4679878ff1
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Kada radite sa čistim JavaScript-om (vanilla JavaScript), ništa vas ne sprečava da pozovete funkciju `square` sa stringom ili nizom, što bi dovelo do neočekivanih rezultata.

Da biste ovo ilustrovali, promenite poziv funkcije da koristi string `"something"` umesto broja `5`.

U sledećoj lekciji, naučićete kako da koristite TypeScript za dodavanje sigurnosti tipa vašem kodu i sprečavanje ovakvih grešaka.

# --hints--

Your `square` function call should have an argument of `"something"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.result.value.matches("square('something')"));
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
