---
id: 699f458e47fb88c756fdba81
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Trenutno, parametar `numbers` ima implicitni tip `any`. Da biste to popravili, možete eksplicitno definisati parametar kao niz brojeva.

Zapamtite iz prethodnih lekcija da možete tipizovati niz brojeva na ovaj način:

```ts
const numbers: number[] = [1, 2, 3];
```

Ažurirajte svoju funkciju `getAverage` da eksplicitno tipizira parametar `numbers` kao niz brojeva.

# --hints--

Your `getAverage` function should have a parameter named `numbers` that is explicitly typed as an array of numbers.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.getAverage.parameters;
assert.isTrue(parameters[0].matches("numbers: number[]"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}
--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
