---
id: 699f52220de88c8091032076
title: Korak 18
challengeType: 1
dashedName: step-18
---

# --description--

Ako otvorite konzolu, rezultat je `NaN` jer drugi argument za `exponent` nikada nije dostavljen. Potrebno je nekoliko ažuriranja za ovu funkciju `raiseTo` koja će ovo popraviti.

Prva promena je da se vrati `base ** 2`.

U sledećem koraku, videćete kako da rukujete opcionalnim parametrom `exponent`.

# --hints--

Your `raiseTo` function should return a number.

```js
assert.isNumber(raiseTo(4));
```

Your `raiseTo` function should return the value of `base` raised to the power of `2`.

```js
assert.strictEqual(raiseTo(4), 16);
assert.strictEqual(raiseTo(5), 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
