---
id: 699f3f636aa840b7201aa37c
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Ako otvorite konzolu, trebalo bi da vidite sledeću poruku greške TypeScript-a:

```md
Argument of type 'string' is not assignable to parameter of type 'number'
```

Sada kada je parametar eksplicitno tipizovan kao broj, TypeScript može da uhvati grešku pokušaja prosleđivanja string argumenata funkciji `square`. 

Promenite argument prosleđen funkciji `square` u broj `5` da biste popravili grešku i videli ispravan izlaz u konzoli.

# --hints--

Your `square` function call should have an argument of `5` instead of `"something"`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");
--fcc-editable-region--

console.log(result);
```

