---
id: 699f3cb8d057f29954205d7d
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Kao što ste se setili iz ranijih lekcija, možete dodati tipne anotacije funkcionim parametrima na ovaj način:

```ts
function add(num1: number, num2: number) {
  return num1 + num2;
}
```

U ovom primeru, TypeScript će prikazati grešku ako pokušate da pozovete funkciju `add` sa argumentima koji nisu brojevi.

Ažurirajte svoju funkciju `square` da uključite tipnu anotaciju za parametar `num`. Tip bi trebalo da bude `number`.

# --hints--

You should add a type annotation to the `num` parameter in the `square` function. The type should be a `number`.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.square;
assert.isTrue(parameters[0].matches("num: number"));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
