---
id: 699f3738921fb79a651ded56
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Sada je vreme da pozovete svoju funkciju.

Kreirajte varijablu nazvanu `result` i dodelite joj poziv funkcije `square` sa argumentom `5`.

Ispod toga, logujte vrednost `result` na konzolu.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should have a variable called `result`.

```js
assert.isDefined(result);
```

Your `result` variable should be a number.

```js
assert.isNumber(result);
```

Your `result` variable should be the square of `5`.

```js
assert.strictEqual(result, 25);
```

You should log the `result` variable to the console.

```js
assert.equal(getLogs()[0], result);
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
