---
id: 699f4861dd5b5f164e1bda4a
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

Još jedna prednost korišćenja TypeScript-a je osiguranje da su tipovi povratnih vrednosti striktno definisani i ono što očekujete. Ponekad TypeScript može zaključiti tip povratne vrednosti funkcije, ali ga možete eksplicitno specificirati.

Da bismo ovo ilustrovali, promenite tip povratne vrednosti funkcije `getAverage` u string `"something"`.

# --hints--

Your `getAverage` function should return a string.

```js
assert.isString(getAverage([2, 14, 26, 8]));
```

Your `getAverage` function should return the string `"something"`.

```js
assert.strictEqual(getAverage([2, 14, 26, 8]), "something");
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  --fcc-editable-region--
  return sum / numbers.length;
  --fcc-editable-region--
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
