---
id: 699f412114c87221f16a99e3
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Još jedna prednost korišćenja TypeScript-a je što će uhvatiti greške ako prosledite previše ili premalo argumenata funkcijskom pozivu.

Da bismo ovo ilustrovali, ažurirajte vaš poziv funkcije ``square`` da uzme dva argumenta od tipova ``5`` i ``10``.

# --hints--

Your `square` function call should have two arguments of `5` and `10`.

```js
const explorer = await __helpers.Explorer(code);
const { result } = explorer.variables;
assert.isTrue(result.value.matches("square(5, 10)"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5, 10);

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
