---
id: 699f52fac083214b6bf51de6
title: Korak 19
challengeType: 1
dashedName: step-19
---

# --description--

U ovom poslednjem koraku, moraćete da dodate uslovnu proveru za parametar `exponent` u funkciji `raiseTo`.

Ako je naveden argument `exponent`, funkcija bi trebalo da vrati `base` podignut na stepen `exponent`. Ako nije naveden argument `exponent`, trebalo bi da vrati `base ** 2`.

Ovim poslednjim korakom, vaš tip-sigurni matematički alat je završen!

# --hints--

Your `raiseTo` function should return `base ** exponent` when `exponent` argument is provided.

```js
assert.strictEqual(raiseTo(3, 3), 27);
```

Your `raiseTo` function should return `base ** 2` when `exponent` argument is not provided.

```js
assert.strictEqual(raiseTo(3), 9);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {

  return base ** 2;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  if (exponent) {
    return base ** exponent;
  }

  return base ** 2;
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
