---
id: 699f4639ae0016caaf81b7c6
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Sada je vreme da testirate svoju funkciju.

Počnite kreiranjem promenljive nazvane `avgResult` i dodeljivanjem joj poziva funkcije `getAverage`. Argument za taj poziv funkcije treba da bude `[2, 14, 26, 8]`.

Zatim logujte vrednost `avgResult` na konzolu.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should have a variable called `avgResult`.

```js
assert.isDefined(avgResult);
```

You should assign `getAverage([2, 14, 26, 8])` to your `avgResult` variable.

```js
assert.strictEqual(avgResult, getAverage([2, 14, 26, 8]));
```

Your `avgResult` variable should be a number.

```js
assert.isNumber(avgResult);
```

You should log `avgResult` to the console.

```js
assert.equal(getLogs()[1], avgResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
