---
id: 66c06fad3475cd92421b9ac2
title: Izgradite kalkulator za prolećnu godišnjicu
challengeType: 26
dashedName: build-a-leap-year-calculator
---

# --description--

Zelena godina je godina koja je deljiva sa `4`, osim za godine koje su deljive sa `100` i nisu deljive sa `400`. Na primer, `2000` je zelena godina, ali `1900` nije. Takođe, zelena godina ima dodatni dan u februaru, koji je 29. dan meseca.

**Cilj:** Ispunite priče korisnika ispod i prođite sve testove kako biste završili laboratoriju.

**Priče korisnika:**

1. Definišite funkciju nazvanu `isLeapYear` koja prima broj kao argument.
2. Napolju od funkcije, deklarišite varijablu `year` koja čuva vrednost godine koju želite proveriti.
3. Unutar funkcije, koristite `if/ else` izjavu ili ternarni operator da proverite da li je godina zelena.
4. Da biste proverili da li je godina zelena, ispunite sledeće uslove:

   - Ako je godina deljiva sa `4`, onda je zelena godina.
   - Osim ako nije i deljiva sa `100`, onda nije zelena godina.
   - Osim ako je takođe deljiva sa `400`, onda je zelena godina.

5. Ako je godina zelena, vratite `[year] is a leap year.`. U suprotnom, vratite `[year] is not a leap year.`. Zamenićete `[year]` parametrom definisanim u funkciji `isLeapYear`.
6. Trebalo bi da pozovete funkciju `isLeapYear` sa `year` kao argumentom i dodelite rezultat varijabli nazvanoj `result`.
7. Trebalo bi da ispišete varijablu `result` na konzolu koristeći `console.log()`.

# --hints--

You should define a function named `isLeapYear`.

```js
assert.isFunction(isLeapYear);
```

The `isLeapYear` function should take a number as an argument.

```js
assert.match(isLeapYear.toString(), /\s*isLeapYear\(\s*\w+\s*\)/);
```

You should declare a variable `year` and assign it a value to check if it is a leap year.

```js
assert.isDefined(year);
assert.isNumber(year);
assert.isAtLeast(year, 0);
assert.isTrue(Number.isInteger(year));
```

The `year` variable shouldn't be empty.

```js
assert.isNotNull(year);
```

With `2024` as the value of the `year` variable, the `result` should be `2024 is a leap year.`

```js
assert.strictEqual(isLeapYear(2024), '2024 is a leap year.');
```

With `2000` as the value of the `year` variable, the `result` should be `2000 is a leap year.`

```js
assert.strictEqual(isLeapYear(2000), '2000 is a leap year.');
```

With `1900` as the value of the `year` variable, the `result` should be `1900 is not a leap year.`

```js
assert.strictEqual(isLeapYear(1900), '1900 is not a leap year.');

```

You should call the `isLeapYear` function and pass `year` as a parameter.

```js
assert.match(__helpers.removeJSComments(code), /isLeapYear\(\s*year\s*\)/);
```

You should declare a `result` variable.

```js
assert.isDefined(result);
```

You should store the result of calling the `isLeapYear` function in a variable named `result`.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*isLeapYear\(\s*year\s*\)/);
```

You should output the `result` to the console using `console.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function isLeapYear(year) {
  if ((year % 4 === 0 && year % 100 !== 0) || year % 400 === 0) {
    return(year + " is a leap year.");
  } else {
    return(year + " is not a leap year.");
  }
}

const year = 2024;
const result = isLeapYear(year);
console.log(result);
```

