---
id: 61696ef7ac756c829f9e4048
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

U ovom projektu ćete raditi sa eksternom CSS datotekom za stilizovanje stranice. Datoteka `styles.css` je već kreirana za vas. Ali pre nego što ga možete koristiti, moraćete da ga povežete sa stranicom.

Ugnjezdite element `link` unutar elementa `head`. Dajte mu atribut `rel` postavljen na `"stylesheet"` i atribut `href` postavljen na `"styles.css"`.

# --hints--

You should have one `link` element.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Your `link` element should be a void element, it should not have a closing tag `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Your `link` element should have a `rel` attribute with the value `"stylesheet"`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

Your `link` element should have an `href` attribute with the value `"styles.css"`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

Your `link` element should be inside the `head` element.

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
