---
id: 616971b3cd990186b66c99a1
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Sada što je vaš eksterni CSS fajl podešen, možete početi sa stilizovanjem stranice.

Kao podsetnik, evo kako da ciljate paragraf element i poravnate ga desno:

```css
p {
  text-align: right;
}
```

Kreirajte novo CSS pravilo koje cilja element `h1`, i postavite njegovu svojstvo `text-align` na `center`.

# --hints--

You should use an `h1` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

Your `h1` CSS rule should have a `text-align` property set to `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
