---
id: 616d4a84b756d9c4b8255093
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Vreme je da dodamo boju markeru. Zapamtite da je jedan način za dodavanje boje elementu korišćenje ključne reči `<dfn>color</dfn>` kao što su `black`, `cyan` ili `yellow`.

Kao podsetnik, evo kako da ciljate klasu `freecodecamp`:

```css
.freecodecamp {
  
}
```

Kreirajte novo CSS pravilo koje cilja klasu `marker`, i postavite njegovu svojstvo `background-color` na `red`.

**Napomena:** Nećete videti nikakve promene nakon dodavanja CSS-a.

# --hints--

You should create a class selector to target the `marker` class.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.marker'));
```

Your `.marker` CSS rule should have a `background-color` property set to `red`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor, 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
