---
id: 616d3a67ccf800ad94ec89ae
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Sada ćete dodati neke elemente koje ćete na kraju stilizovati u kolor markere.

Prvo, unutar elementa `body`, dodajte element `div` sa atributom `class` postavljenim na `container`. Uverite se da je element `div` postavljen ispod elementa `h1`.

# --hints--

You should have an opening `div` tag.

```js
assert.match(code, /<div.*>/i);
```

You should have a closing `div` tag.

```js
assert.match(code, /<\/div\s*>/i);
```

Your `div` element should be within your `body` element.

```js
assert.strictEqual(document.querySelector('div')?.parentElement?.localName, 'body');
```

Your `div` element should be after the `h1` element.

```js
assert.exists(document.querySelector('h1 + div'));
```

Your `div` element should have a `class` attribute set to `container`.

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}
```
