---
id: 616d525007b8c5d8b3308b1c
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Vaš marker bi izgledao bolje da je centriran na stranici. Jednostavan način za to je korišćenje `margin` <dfn>shorthand svojstva</dfn>.

Možete postaviti oblast margina elemenata zasebno sa svojstvima kao što su `margin-top` i `margin-left`, a `margin` shorthand svojstvo olakšava postavljanje više oblasti margina istovremeno.

Da biste centrirali vaš marker na stranici, postavite njegovo svojstvo `margin` na `auto`. Ovo postavlja `margin-top`, `margin-right`, `margin-bottom` i `margin-left` sve na `auto`.

# --hints--

Your `.marker` CSS rule should have a `margin` property set to `auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
