---
id: 61764c602bee6974e7790f35
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Da biste markerima dali različite boje, moraćete dodati jedinstvenu klasu svakom od njih. Možete dodati više klasa elementu navodeći ih u atributu `class` i razdvajajući ih razmakom. Na primer, sledeći kod dodaje klase `animal` i `dog` na element `div`:

```html
<div class="animal dog">
```

Ako dodate više klasa na HTML element, stilovi prvih klasa koje navedete u listu stilova (stylesheet) mogu biti nadjačani kasnijim klasama.

Da biste počeli, dodajte klasu `one` prvom marker elementu `div`.

# --hints--

You should add the class `one` to the first marker `div` element.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

Your first marker `div` should have the classes `marker` and `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('marker'));
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
--fcc-editable-region--
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
