---
id: 617b8e0d93a8d10d9a90e720
title: Korak 18
challengeType: 0
dashedName: step-18
---

# --description--

Primećujete da je `background-color` za vaš marker i dalje crven. To je zato što ste vrednost crvene komponente funkcije `rgb` postavili na maksimum od `255`, ili 100% crven, i postavili obe zelene i plave vrednosti na `0`.

Sada koristite funkciju `rgb` da postavite ostale boje.

U CSS pravilu `.two`, koristite funkciju `rgb` da postavite `background-color` na maksimalnu vrednost za zelenu, i `0` za ostale vrednosti. I u CSS pravilu `.three`, koristite funkciju `rgb` da postavite `background-color` na maksimalnu vrednost za plavu, i `0` za ostale vrednosti.

# --hints--

Your `.two` CSS rule should not use the `green` color keyword to set the `background-color` property.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'green');
```

Your `.two` CSS rule should have a `background-color` property set to `rgb(0, 255, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 0)');
```

Your `.three` CSS rule should not use the `blue` color keyword to set the `background-color` property.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'blue');
```

Your `.three` CSS rule should have a `background-color` property set to `rgb(0, 0, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
