---
id: 617bc3386dc7d07d6469bf20
title: Korak 27
challengeType: 0
dashedName: step-27
---

# --description--

Primećujete da ste za kreiranje narandžaste povećali intenzitet crvene i smanjili intenzitet zelene u vrednostima `rgb`. To je zato što se narandžasta kreira mešanjem crvene sa žutom.

Da biste stvorili tercijarni boju prolećno zelena, kombinujte cijan sa zelenom. Ažurirajte funkciju `rgb` u CSS pravilu `.two` postavljanjem zelene na njenu maksimalnu vrednost i plave na `127`.

# --hints--

Your `.two` CSS rule should have a `background-color` property set to `rgb(0, 255, 127)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 255, 255);
--fcc-editable-region--
}

.three {
  background-color: rgb(255, 0, 255);
}

```
