---
id: 617b954d9f4f6217a749380e
title: Korak 20
challengeType: 0
dashedName: step-20
---

# --description--

Sada dodajte malo više vertikalnog prostora između vaših markera i ivice elementa `container` u kojem se nalaze.

U CSS pravilu `.container`, koristite skraćeno svojstvo (shorthand) `padding` da dodate `10px` za padding na vrhu i dnu, i postavite padding sa leve i desne strane na `0`. Ovo radi slično skraćenom svojstvu (shorthand) `margin` koje ste ranije koristili.

# --hints--

You should not remove the `background-color` property and its value from the `.container` CSS rule.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(0, 0, 0)');
```

Your `.container` CSS rule should have a `padding` property set to `10px 0`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.padding, '10px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
--fcc-editable-region--

--fcc-editable-region--
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 127, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
