---
id: 617bd6ec666b1da2587e4e37
title: Korak 29
challengeType: 0
dashedName: step-29
---

# --description--

Postoje još tri tercijarne boje: chartreuse zelena (zeleno + žuto), azurno (plavo + ciano) i ružičasta/roze (crveno + magenta).

Da biste kreirali chartreuse zelenu, ažurirajte funkciju `rgb` u CSS pravilu `.one` tako da je crvena na `127`, a postavite zeleno na maksimalnu vrednost.

Za azurno, ažurirajte funkciju `rgb` u CSS pravilu `.two` tako da je zeleno na `127` i plavo na maksimalnu vrednost.

A za ružičastu (koja se ponekad naziva svetlo roze), ažurirajte funkciju `rgb` u CSS pravilu `.three` tako da je plavo na `127`, a crveno na maksimalnu vrednost.

# --hints--

Your `.one` CSS rule should have a `background-color` property set to `rgb(127, 255, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(127, 255, 0)');
```

Your `.two` CSS rule should have a `background-color` property set to `rgb(0, 127, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 127, 255)');
```

Your `.three` CSS rule should have a `background-color` property set to `rgb(255, 0, 127)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(255, 0, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
  background-color: rgb(127, 0, 255);
}
--fcc-editable-region--

```
