---
id: 617b97abd9f3f61d1590b815
title: Korak 21
challengeType: 0
dashedName: step-21
---

# --description--

U aditivnom RGB modelu boja, <dfn>primarne boje</dfn> su boje koje, kada se kombinuju, stvaraju čistu belu. Ali da bi se ovo dogodilo, svaka boja mora biti na svojoj najvišoj intenzitetnosti.

Pre nego što kombinujete boje, vratite svoj zeleni marker na čistu zelenu boju. Za funkciju `rgb` u CSS pravilu `.two`, vratite zeleno na maksimalnu vrednost `255`.

# --hints--

Your `.two` CSS rule should have a `background-color` property set to `rgb(0, 255, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 127, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 255);
}

```
