---
id: 617b8b38f32bf2080a140675
title: Korak 17
challengeType: 0
dashedName: step-17
---

# --description--

Kao što ste naučili u prethodnim lekcijama, funkcija je deo koda koji može uzeti ulaz i izvršiti određenu akciju. Funkcija CSS `rgb` prima vrednosti, ili <dfn>argumente</dfn>, za crvenu, zelenu i plavu boju, i proizvodi boju:

```css
rgb(red, green, blue);
```

Svaka vrednost crvene, zelene i plave boje je broj od `0` do `255`. `0` znači da postoji 0% te boje, i to je crno. `255` znači da postoji 100% te boje.

U CSS pravilu `.one`, zamenite ključnu reč boje `red` funkcijom `rgb`. Za funkciju `rgb`, postavite vrednost za crvenu na `255`, vrednost za zelenu na `0` i vrednost za plavu na `0`.

# --hints--

Your `.one` CSS rule should not use the `red` color keyword to set the `background-color` property.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'red');
```

Your `.one` CSS rule should have a `background-color` property set to `rgb(255, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
--fcc-editable-region--
  background-color: red;
--fcc-editable-region--
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
