---
id: 617b65579ce424bf5f02ca73
title: Korak 16
challengeType: 0
dashedName: step-16
---

# --description--

Kao što ste naučili u lekcijama, postoje dva glavna modela boja: <dfn>aditivni</dfn> RGB (crvena, zelena, plava) model, koji se koristi u elektronskim uređajima, i <dfn>subtraktivni</dfn> CMYK (cijan, magenta, žuta, crna) model, koji se koristi za štampu.

U ovom projektu, radićete sa RGB modelom. To znači da boje počinju kao crne i menjaju se kako se uvode različiti nivoi crvene, zelene i plave. Lako način da to vidite je pomoću CSS funkcije `rgb`.

Kreirajte novo CSS pravilo koje cilja klasu `container` i podesite njegovu `background-color` na crnu sa `rgb(0, 0, 0)`.

# --hints--

You should use a class selector to target the class `container`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.container'));
```

Your `.container` CSS rule should have a `background-color` property set to `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
