---
id: 618a0927005553b74bfa05ae
title: Korak 31
challengeType: 0
dashedName: step-31
---

# --description--

Točkić boja je krug gde su slične boje, ili <dfn>nijanze</dfn>, blizu jedna jedne, a različite su dalje jedna od druge. Na primer, čisti crveni se nalazi između nijansi ružičastog i narandžastog.

Dve boje koje su nasprotne jedna drugoj na točkiću boja nazivaju se <dfn>komplementarne boje</dfn>. Ako se dve komplementarne boje kombinuju, proizvode sivu boju. Ali kada su postavljene jedna pored druge, ove boje stvaraju snažan vizuelni kontrast i izgledaju svetlije.

U funkciji `rgb` za CSS pravilo `.one`, postavite vrednost crvene na maksimum od `255` da biste proizveli čisti crveni. U funkciji `rgb` za CSS pravilo `.two`, postavite vrednosti za zelenu i plavu na maksimum od `255` da biste proizveli cijan.

# --hints--

Your `.one` CSS rule should have a `background-color` property set to `rgb(255, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 0, 0)');
```

Your `.two` CSS rule should have a `background-color` property set to `rgb(0, 255, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
