---
id: 618a00ed1ca871a2b3aca0cb
title: Korak 30
challengeType: 0
dashedName: step-30
---

# --description--

Sada kada ste prošli kroz sve primarne, sekundarne i tercijarne boje na točkići boja, lakše će vam biti razumeti druge koncepte teorije boja i kako utiču na dizajn.

Prvo, u CSS pravilima `.one`, `.two` i `.three`, prilagodite vrednosti u funkciji `rgb` tako da je `background-color` svakog elementa postavljen na čistu crnu boju. Zapamtite da funkcija `rgb` koristi additivni model boja, gde boje počinju kao crne i menjaju se kako rastu vrednosti crvene, zelene i plave.

# --hints--

Your `.one` CSS rule should have a `background-color` property set to `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(0, 0, 0)');
```

Your `.two` CSS rule should have a `background-color` property set to `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 0, 0)');
```

Your `.three` CSS rule should have a `background-color` property set to `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
