---
id: 619b761916dac02643940022
title: Korak 41
challengeType: 0
dashedName: step-41
---

# --description--

Vrlo uobičajen način za primenu boje na element pomoću CSS-a je korišćenje `<dfn>heksadecimalnih</dfn>` ili heksadecimalnih vrednosti. Iako zvuče komplikovano, oni su zapravo samo druga forma RGB vrednosti.

Heksadecimalne vrednosti boje počinju sa `#` karakterom i koriste šest karaktera od 0-9 i A-F. Prvi par karaktera predstavlja crvenu boju, drugi par predstavlja zelenu, a treći par predstavlja plavu. Na primer, `#4B5320`.

U klasnom selektoru `.green`, podesite svojstvo `background-color` na heksadecimalni kod boje sa vrednostima `00` za crvenu, `FF` za zelenu i `00` za plavu.

# --hints--

Your `.green` CSS rule should have a `background-color` property set to `#00FF00`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor, 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
--fcc-editable-region--
  background-color: rgb(0, 0, 0);
--fcc-editable-region--
}

.blue {
  background-color: rgb(0, 0, 0);
}

```
