---
id: 619b7fd56aa2253778bcf5f7
title: Korak 43
challengeType: 0
dashedName: step-43
---

# --description--

Model boja <dfn>HSL</dfn>, ili nijansa, zasićenost i svetlina, je još jedan način za predstavljanje boja.

Funkcija CSS `hsl` prima 3 vrednosti: broj od 0 do 360 za nijansu, procenat od 0 do 100 za zasićenost i procenat od 0 do 100 za svetlinu.

Ako zamislite točkić boja, nijansa crvenila je na 0 stepeni, zelena je na 120 stepeni, a plava je na 240 stepeni.

Zasićenost je intenzitet boje od 0%, ili sive, do 100% za čistu boju. Morate dodati znak procenata `%` vrednostima zasićenosti i svetline.

Svetlina je koliko boja izgleda svetlo, od 0%, ili potpunog crnog, do 100%, potpunog belog, pri čemu je 50% neutralno.

U CSS pravilu `.blue`, koristite funkciju `hsl` da promenite svojstvo `background-color` u čistu plavu boju. Postavite nijansu na `240`, zasićenost na `100%` i svetlinu na `50%`.

# --hints--

Your `.blue` CSS rule should have a `background-color` property set to `hsl(240, 100%, 50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

.blue {
--fcc-editable-region--
  background-color: rgb(0, 0, 0);
--fcc-editable-region--
}

```
