---
id: 61a489b8579e87364b2d2cdb
title: Korak 44
challengeType: 0
dashedName: step-44
---

# --description--

Vežbali ste nekoliko načina za postavljanje ravnih boja u CSS-u, ali možete koristiti i prelaz boje, ili <dfn>gradijent</dfn>, na elementu.

Gradijent je kada jedna boja prelazi u drugu. CSS funkcija `linear-gradient` omogućava vam da kontrolišete pravac tranzicije duž linije i koje boje se koriste.

Jedna stvar koju treba zapamtiti je da funkcija `linear-gradient` zapravo kreira element `image`, i obično je uparena sa svojstvom `background`, koje može prihvatiti sliku kao vrednost.

U CSS pravilu `.red`, promenite svojstvo `background-color` u `background`.

# --hints--

Your `.red` CSS rule should have a `background` property with the value `rgb(255, 0, 0)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:rgb\(255,0,0\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background-color: rgb(255, 0, 0);
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
