---
id: 61a493ead935c148d2b76312
title: Korak 45
challengeType: 0
dashedName: step-45
---

# --description--

Funkcija `linear-gradient` je veoma fleksibilna — evo osnovne sintakse koju ćete pregledati u ovom tutorijalu:

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection` je smer linije korišćen za tranziciju. `color1` i `color2` su argumenti boje, koji su boje koje će se koristiti u samoj tranziciji. Ovo mogu biti bilo kog tipa boje, uključujući ključne reči boja, hex, `rgb`, ili `hsl`.

Sada ćete primeniti crveno-zeleni gradijent duž linije od 90 stepeni na prvi marker.

Prvo, u `.red` CSS pravilu, postavite svojstvo `background` na `linear-gradient()`, i prosledite mu vrednost `90deg` kao `gradientDirection`.

# --hints--

Your `.red` CSS rule should have a `background` property with the value `linear-gradient(90deg)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background: rgb(255, 0, 0);
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
