---
id: 61b0a44a6b865738ba49b9fb
title: Korak 81
challengeType: 0
dashedName: step-81
---

# --description--

Poslednja stvar koju ćete uraditi je dodavanje blage senke svakom markeru kako bi izgledali još realnije.

Svojstvo `box-shadow` omogućava vam da primenite jednu ili više senki oko elementa. Evo osnovne sintakse:

```css
box-shadow: offsetX offsetY color;
```

Evo kako funkcionišu vrednosti `offsetX` i `offsetY`:

* oba `offsetX` i `offsetY` prihvataju numeričke vrednosti u `px` i druge CSS jedinice
* pozitivna vrednost `offsetX` pomera senku udesno, a negativna vrednost pomeri je ulevo
* pozitivna vrednost `offsetY` pomera senku nadole, a negativna vrednost pomeri je nagore
* ako želite vrednost nula (`0`) za bilo koji ili oba `offsetX` i `offsetY`, ne morate dodavati jedinicu. Svaki pretraživač razume da nula znači bez promene.

Visina i širina senke određuju se visinom i širinom elementa na koji je primenjena. Takođe možete koristiti opcionalnu vrednost `spreadRadius` da proširite opseg senke. Više o tome kasnije.

Počnite dodavanjem jednostavne senke crvenom markeru.

U CSS pravilu `.red`, dodajte svojstvo `box-shadow` sa vrednostima `5px` za `offsetX`, `5px` za `offsetY` i `red` za `color`.

# --hints--

Your `.red` CSS rule should have a `box-shadow` shorthand property and with the value `5px 5px red`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow, 'red 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
--fcc-editable-region--
  
--fcc-editable-region--
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
