---
id: 61b0936d9e7fc020b43b1bb8
title: Korak 73
challengeType: 0
dashedName: step-73
---

# --description--

Izgleda da je vaš rukav nestao, ali ne brinite — on je još uvek tu. Ono što se dogodilo jeste da vaša nova kapa `div` zauzima celu širinu markera i gura rukav na sledeći red.

To je zato što podrazumevana svojstvo `display` za elemente `div` jeste `block`. Dakle, kada su dva `block` elementa jedan pored drugog, oni se slažu kao stvarni blokovi. Na primer, vaši markeri elementi su sve složeni jedan preko drugog.

Da biste pozicionirali dva `div` elementa na istom redu, podesite njihove `display` svojstva na `inline-block`.

Kreirajte novo pravilo za ciljanje kako klase `cap` tako i `sleeve`, i podesite `display` na `inline-block`.

# --hints--

You should use a class selector to target both the `cap` and `sleeve` classes.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve') || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap'));
```

Your CSS rule should have a `display` property set to `inline-block`.

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve')?.display === 'inline-block' || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap')?.display === 'inline-block');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
}

--fcc-editable-region--

--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
