---
id: 61b093329e7fc020b43b1bb5
title: Korak 70
challengeType: 0
dashedName: step-70
---

# --description--

Koristili ste funkciju `rgb` za postavljanje boja. Da biste dodali alfa kanal boji `rgb`, koristite funkciju `rgba` umesto toga.

Funkcija `rgba` radi baš kao funkcija `rgb`, ali uzima još jedan broj od `0` do `1.0` za alfa kanal:

```css
rgba(redValue, greenValue, blueValue, alphaValue);
```

Takođe možete koristiti alfa kanal sa bojama `hsl` i `hex`. Uskoro ćete videti kako to učiniti.

U pravilu `.sleeve`, koristite funkciju `rgba` da postavite svojstvo `background-color` na čistu belu sa opacititetom od 50%.

# --hints--

Your `.sleeve` CSS rule should have a `background-color` property set to `rgba(255, 255, 255, 0.5)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.sleeve')?.backgroundColor, 'rgba(255, 255, 255, 0.5)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.sleeve {
  width: 110px;
  height: 25px;
--fcc-editable-region--
  background-color: white;
--fcc-editable-region--
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
