---
id: 61b0a1b2af494934b7ec1a72
title: Korak 80
challengeType: 0
dashedName: step-80
---

# --description--

Odlično. Tvoj crveni marker izgleda dobro. Sada sve što treba da uradiš je dodati kape i rukave na ostale markere.

Na zeleni i plavi marker dodaj kapu i rukav. Možeš jednostavno kopirati elemente `div` sa crvenog markera i zalepiti ih u drugi dva markera.

# --hints--

Your green marker `div` should contain two `div` elements.

```js
const greenMarkerChildren = [...document.querySelector('.green')?.children];
greenMarkerChildren.forEach(child => assert.strictEqual(child?.localName, 'div'))
assert.lengthOf(greenMarkerChildren, 2);
```

Your green marker's cap `div` element should be before the sleeve `div` element.

```js
const greenMarkerChildren = [...document.querySelector('.green')?.children];
const greenMarkerCap = document.querySelector('.green .cap');
const greenMarkerSleeve = document.querySelector('.green .sleeve');
assert.isBelow(greenMarkerChildren.indexOf(greenMarkerCap), greenMarkerChildren.indexOf(greenMarkerSleeve));
```

Your blue marker `div` should contain two `div` elements.

```js
const blueMarkerChildren = [...document.querySelector('.blue')?.children];
blueMarkerChildren.forEach(child => assert.strictEqual(child?.localName, 'div'));
assert.lengthOf(blueMarkerChildren, 2);
```

Your blue marker's cap `div` element should be before the sleeve `div` element.

```js
const blueMarkerChildren = [...document.querySelector('.blue')?.children];
const blueMarkerCap = document.querySelector('.blue .cap');
const blueMarkerSleeve = document.querySelector('.blue .sleeve');
assert.isBelow(blueMarkerChildren.indexOf(blueMarkerCap), blueMarkerChildren.indexOf(blueMarkerSleeve));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      --fcc-editable-region--
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
      --fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
