---
id: 66fe4f33a2cc9b33f4d5cd9b
title: Izgradite generator permutacija
challengeType: 26
dashedName: build-a-permutation-generator
---

# --description--

U ovom laboratorijumu, kreiraćete generator permutacija koji će uzeti string i vratiti sve moguće permutacije karaktera u tom stringu. Na primer, moguće permutacije stringa `cat` su `cat`, `cta`, `act`, `atc`, `tac` i `tca`.

Rekurzivni način kreiranja permutacija stringa radi skladištenjem fiksiranog početnog dela stringa (prefiksa) i kreiranjem permutacija ostatka.

Na primer, razmotrite reč `machine`. Prva runda kreiranja permutacija bi se izvršila fiksiranjem `m` kao prefiksa stringa, a zatim kreiranjem permutacija ostatka stringa, `achine`.

Za ostatak stringa, permutacije se nastavljaju na isti način. Dodaje se jedan slovo prefiksu, možda `c`, tako da prefiks postaje `mc`. Zatim, svaka od permutacija `ahine` konkatinuje se sa prefiksom.

Ovo se nastavlja sve dok prefiks ne sadrži sva slova, a ostatak stringa nije prazan, što znači da je kreirana jedna permutacija.

**Cilj:** Ispunite korisničke priče ispod i učinite da svi testovi prođu kako biste završili laboratoriju.

**Korisničke priče:**

1. Trebalo bi da kreirate funkciju pod nazivom `permuteString`.
2. Funkcija `permuteString` treba da uzme jedan parametar, string, i zatim dva parametra sa podrazumevanom vrednošću: prefiksna vrednost i prazan niz za skladištenje i vraćanje rezultata. Prefiksna vrednost bi trebalo da bude string koji se koristi za akumuliranje karaktera radi formiranja permutacije. Funkcija će biti pozvana sa jednim argumentom, kao što je `permuteString("cat")`.
3. Unutar funkcije, trebalo bi da proverite da li je dužina prosleđenog stringa `0`. Ako jeste, dodajte trenutni prefiks u rezultate i vratite rezultate.
4. Iterirajte kroz svaki karakter u ulaznom stringu i za svaku iteraciju, uklonite trenutni karakter iz stringa i pozovite funkciju `permuteString` rekurzivno sa ažuriranim argumentima kako biste izgradili preostale permutacije.
5. Trebalo bi da vratite konačan niz rezultata.
6. Trebalo bi da osigurate da su permutacije jedinstvene uklanjanjem duplikata.

# --hints--

You should have a function `permuteString`.

```js
assert.isFunction(permuteString);
```

You should use recursion in your `permuteString` function.

```js
assert.match(permuteString.toString(), /permuteString\s*\(/);
```

`permuteString("far")` should return `[ "far", "fra", "afr", "arf", "rfa", "raf" ]`.

```js
assert.sameMembers(permuteString("far"), [ "far", "fra", "afr", "arf", "rfa", "raf" ]);
```

`permuteString("fcc")` should return `[ "fcc", "cfc", "ccf" ]`.

```js
assert.sameMembers(permuteString("fcc"), [ "fcc", "cfc", "ccf" ]);
```

`permuteString("p")` should return `[ "p" ]`.

```js
assert.deepStrictEqual(permuteString("p"), ["p"]);
```

`permuteString("")` should return `[""]`.

```js
let emptyArr= permuteString("");

// 1 because the empty string is being pushed and is a permutation of itself
assert.lengthOf(emptyArr, 1);

```

`permuteString("walk")` should return `["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]`.
`.

```js
assert.sameMembers(permuteString("walk"), ["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]);
```

`permuteString` should return the correct results.

```js
// tests to prevent hard coding
assert.sameMembers(permuteString("road"), [ "road",
  "roda",
  "raod",
  "rado",
  "rdoa",
  "rdao",
  "orad",
  "orda",
  "oard",
  "oadr",
  "odra",
  "odar",
  "arod",
  "ardo",
  "aord",
  "aodr",
  "adro",
  "ador",
  "droa",
  "drao",
  "dora",
  "doar",
  "daro",
  "daor" ]);

assert.sameMembers(permuteString("fog"), [ "fog",
  "fgo",
  "ofg",
  "ogf",
  "gfo",
  "gof" ]);

assert.sameMembers(permuteString("bird"), [ "bird",
  "bidr",
  "brid",
  "brdi",
  "bdir",
  "bdri",
  "ibrd",
  "ibdr",
  "irbd",
  "irdb",
  "idbr",
  "idrb",
  "rbid",
  "rbdi",
  "ribd",
  "ridb",
  "rdbi",
  "rdib",
  "dbir",
  "dbri",
  "dibr",
  "dirb",
  "drbi",
  "drib" ]);

```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const permuteString = (str, prefix = "", result = []) => {
  if (str.length === 0) {
    result.push(prefix);
    return result;
  }

  // Sort the string initially to group duplicates
  const sortedStr = str.split("").sort().join("");

  for (let i = 0; i < sortedStr.length; i++) {
    if (i > 0 && sortedStr[i] === sortedStr[i - 1]) {
      // Skip duplicate characters
      continue;
    }

    const rem = sortedStr.slice(0, i) + sortedStr.slice(i + 1);
    permuteString(rem, prefix + sortedStr[i], result);
  }

  return result;
};
```
