---
id: 66ada3f46945763dd97f43f8
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Ranije ste kreirali varijable `bot` i `botLocation`. Sada ćete ih koristiti za ispis novih poruka na konzolu.

U prethodnim lekcijama, naučili ste kako da radite sa konkatenacijom stringova koristeći operator `+` za spajanje stringova zajedno na ovaj način:

```js
let firstName = "John";

console.log("Hello, my name is " + firstName + ".");
// result: "Hello, my name is John."
```

Zapamtite da morate biti pažljivi sa razmakama prilikom konkatenacije stringova sa varijablama.

Kreirajte varijablu nazvanu `botIntroduction`.

Zatim koristite konkatenaciju stringova sa operatorom `+` da spojite string `"My name is "`, praćen varijablom `bot`, a zatim tačkom (`.`). 

Dodelite ovu vrednost varijablji `botIntroduction`.

Zatim, logujte (ili: Ispišite) varijablu `botIntroduction` na konzolu.

# --hints--

You should have a variable called `botIntroduction`.

```js
assert.isNotNull(botIntroduction);
```

Your `botIntroduction` variable should hold the value of a string.

```js
assert.isString(botIntroduction);
```

You should use string concatenation with the `+` operator to join the string `"My name is "` with the `bot` variable followed by a period (`.`). Be mindful of spaces. Assign this concatenated string to the `botIntroduction` variable.

```js
assert.equal(botIntroduction, "My name is teacherBot.");
```

You should have a fourth `console.log()` statement in your code.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 4);
```

Your `console` statement should output the `botIntroduction` variable.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botIntroduction\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

--fcc-editable-region--

--fcc-editable-region--
```
