---
id: 66ad8ab945d266383f318cbf
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Sada bi trebalo da vidite prvu poruku od bota u konzoli. 

Vreme je dodati drugu poruku od bota.

Dodajte još jednu izjavu ``console.log`` da bi se poruka ``"I am excited to talk to you."`` ispisala u konzolu.

# --hints--

You should have a second `console.log()` statement in your code.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 2);
```

Your `console` statement should have the message `"I am excited to talk to you."`. Don't forget the quotes.

```js
assert.match(__helpers.removeJSComments(code), /("|')I\s+am\s+excited\s+to\s+talk\s+to\s+you.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```
