---
id: 66ad998285df023b23dacdd3
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Sada je vreme da se inicijalizuje promenljiva `botLocation`.

Kada morate deklarisati promenljive sa više reči, možete koristiti konvenciju imenovanja <dfn>camelCase</dfn>.

Kada koristite `camelCase`, prva reč je u malim slovima, a prvi slovo svake naknadne reči je veliko. 

Evo primera:

```js
let thisIsCamelCase = "Alice the camel is actually a horse.";
```

Deklarišite i dodelite string `"the universe"` promenljivoj `botLocation` na istoj liniji koristeći ključnu reč `let`.

# --hints--

Use the `let` keyword to declare a variable.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

Your variable should be named `botLocation`.

```js
assert.isNotNull(botLocation);
```

You should assign a string to your `botLocation` variable.

```js
assert.isString(botLocation);
```

Your `botLocation` variable should be assigned the string `"the universe"`.

```js
assert.equal(botLocation, "the universe");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
