---
id: 66cc12fa504b0479dac479a0
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Prednost funkcija je što se mogu pozvati sa različitim argumentima, omogućavajući vam da ponovo koristite isti kod sa različitim vrednostima.

Evo nekoliko primera poziva funkcija sa različitim string argumentima:

```js
// function definition
function greetUser(name) {
  console.log(`Hello ${name}!`);
}

// function calls
greetUser("John"); // "Hello John!"
greetUser("Jane"); // "Hello Jane!"
```

Dodajte `console.log` koji poziva funkciju `calculateSum` sa argumentima `2` i `5`.

# --hints--

You should have a `console.log` that calls the `calculateSum` function with the arguments `2` and `5`.

```js
assert.match(code, /console\.log\(calculateSum\(\s*2\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
