---
id: 66cc2739f687e484d50bb6f1
title: Korak 18
challengeType: 1
dashedName: step-18
---

# --description--

Poslednji deo funkcionalnosti za vaš kalkulator biće funkcija kvadratnog korena.

Kvadrat je broj pomnožen sam sobom. Dakle, kvadratni koren bi bio broj koji, kada se pomnoži sam sobom, daje originalan broj.

Na primer, kvadratni koren od `9` je `3` jer je `3 * 3 = 9`.

Da biste dobili kvadratni koren broja u JavaScript-u, možete koristiti metod <dfn>Math.sqrt()</dfn>.

Evo primera:

```js
// result: 3
Math.sqrt(9);
```

Deklarišite funkciju nazvanu `calculateSquareRoot` koja prima parametar `num` i vraća kvadratni koren od `num`.

# --hints--

You should have a function called `calculateSquareRoot`.

```js
assert.isFunction(calculateSquareRoot);
```

Your `calculateSquareRoot` function should take a `num` parameter.

```js
const regex = __helpers.functionRegex('calculateSquareRoot', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

Your `calculateSquareRoot` function should return the square root of `num`.

```js
assert.equal(calculateSquareRoot(9), 3);
assert.equal(calculateSquareRoot(16), 4);
assert.equal(calculateSquareRoot(25), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

--fcc-editable-region--

--fcc-editable-region--
```
