---
id: 66cc06012e46aa6bc9b8c001
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

Na ovom radionici, pregledaćete kako da radite sa funkcijama izgradnjom aplikacije kalkulatora.

U prethodnim lekcijama, naučili ste kako da deklarišete funkcije ovako:

```js
// regular function
function myFunction() {

}

// arrow function
const myArrowFunction = () => {

}
```

Počnite kreiranjem funkcije nazvane `addTwoAndSeven`.

Možete izabrati da koristite standardnu sintaksu funkcija ili sintaksu strelične funkcije.

# --hints--

You should declare a function called `addTwoAndSeven`.

```js
assert.isFunction(addTwoAndSeven);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
