---
id: 66cc0a9e06e00b75d6782be9
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

U prethodnim lekcijama ste naučili kako da pozovete (izvršite) funkciju. Pozivanje funkcije znači izvršiti kod unutar funkcije.

Evo podsetnika kako da pozovete funkciju:

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

Dodajte izjavu `console.log`, a unutar nje pozovite funkciju `addTwoAndSeven`.

Sada bi trebalo da vidite zbir `2` i `7` logovan u konzolu.

# --hints--

You should have a `console.log` in your code.

```js
assert.match(code, /console\.log(.*)/);
```

You should call the `addTwoAndSeven` function inside the `console.log`.

```js
assert.match(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

--fcc-editable-region--

--fcc-editable-region--
```
