---
id: 66cc14ca14e65e7a9ebecb80
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Sada što vaš kalkulator može da sabere dva broja, vreme je da kreirate funkciju koja će oduzimati dva broja.

Deklarišite funkciju nazvanu `calculateDifference` koja prima dva parametra, `num1` i `num2`.

Unutar funkcije, vratite razliku između ta dva parametra.

# --hints--

You should have a function called `calculateDifference`.

```js
assert.isFunction(calculateDifference);
```

Your `calculateDifference` function should have two parameters called `num1` and `num2`.

```js
const regex = __helpers.functionRegex('calculateDifference', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

Your `calculateDifference` function should return the difference between the two parameters.

```js
assert.equal(calculateDifference(5, 2), 3);
assert.equal(calculateDifference(10, 5), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

--fcc-editable-region--

--fcc-editable-region--
```
