---
id: 66cc0de2920bec775f610424
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Iako funkcije rade kako je očekivano, postoji mnogo ponavljanja u vašem kodu. 

Pošto gradite kalkulator, ne želite da kreirate funkciju za svaku moguću kombinaciju brojeva koje želite sabrati.

Bolje bi bilo kreirati jednu ponovno upotrebljivu funkciju koja može sabrati bilo koja dva broja.

Izbrišite sav kod koji ste do sada napisali.

U sledećem koraku, pregledaćete kako da radite sa parametrima i argumentima u funkcijama.

# --hints--

You should not have a function called `addTwoAndSeven`.

```js
assert.notMatch(code, /function\s+addTwoAndSeven\s*\(\s*\)\s*{/);
```

You should not have a function call to `addTwoAndSeven`.

```js
assert.notMatch(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

You should not have a function called `addThreeAndFour`.

```js
assert.notMatch(code, /function\s+addThreeAndFour\s*\(\s*\)\s*{/);
```

You should not have a function call to `addThreeAndFour`.

```js
assert.notMatch(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
