---
id: 66cc10958d72e57907f898b2
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Unutar vaše funkcije ``calculateSum``, moraćete da vratite zbir dva parametra.

# --hints--

Your `calculateSum` function should return the sum of the two parameters.

```js
assert.equal(calculateSum(1, 2), 3);
assert.equal(calculateSum(5, 5), 10);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
