---
id: 66cc095dfe1682753d2ab030
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Jedan od koncepata koje ste naučili u prethodnim lekcijama je kako vratiti vrednosti iz funkcije.

Evo podsetnika o tome kako vratiti vrednost iz funkcije:

```js
function myFunction() {
  return "Hello World";
}
```

Unutar vaše funkcije `addTwoAndSeven`, vratite zbir `2` i `7`.

# --hints--

Your `addTwoAndSeven` function should return the sum of `2` and `7`.

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

Your `addTwoAndSeven` function should return the sum using the `+` operator.

```js
assert.match(code, /return\s*(2\s*\+\s*7|7\s*\+\s*2)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {  
--fcc-editable-region--

--fcc-editable-region--
}
```
