---
id: 66cc0f1ae40802781b2ea972
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

U prethodnim lekcijama, naučili ste kako da radite sa parametrima funkcija.

Parametar funkcije je promenljiva koja se definiše u deklaraciji funkcije i deluje kao zamenski deo (placeholder).

Evo primera funkcije koja ima parametar:

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

Deklarišite funkciju nazvanu `calculateSum` koja prima dva parametra, `num1` i `num2`.

# --hints--

You should declare a function called `calculateSum`.

```js
assert.isFunction(calculateSum);
```

Your `calculateSum` function should take two parameters, `num1` and `num2`.

```js
const regex = __helpers.functionRegex('calculateSum', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
