---
id: 66cdf76685e4cb5a8726e27b
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Da biste dalje testirali svoju funkciju `calculateSum`, moraćete je pozvati sa različitim argumentima.

Dodajte još jedan `console.log` koji poziva funkciju `calculateSum` sa argumentima `10` i `10`.

Ispod tog `console.log`, dodajte još jedan `console.log` koji poziva funkciju `calculateSum` sa argumentima `5` i `5`.

# --hints--

You should have a `console.log` that calls the `calculateSum` function with the arguments `10` and `10`.

```js
assert.match(code, /console\.log\(calculateSum\(10\s*,\s*10\)\)/);
```

You should have a `console.log` that calls the `calculateSum` function with the arguments `5` and `5`.

```js
assert.match(code, /console\.log\(calculateSum\(5\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

console.log(calculateSum(2, 5));
--fcc-editable-region--

--fcc-editable-region--
```
