---
id: 5dfb6250eacea3f48c6300b2
title: Korak 24
challengeType: 0
dashedName: step-24
---

# --description--

Nakon neuređene liste, dodaj novu sliku sa vrednošću atributa elementa `src` postavljeno na:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

I vrednost svojstva elementa `alt` je:

`A slice of lasagna on a plate.`

# --hints--

Kuwa na kipengele cha `img` mara moja baada ya lebo ya kufunga `</ul>`.

```js
assert.equal(document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.nodeName, 'IMG');
```

Nova slika nema atribut polja `alt`. Uverite se da postoji razmak nakon imena oznake otvaranja i/ili da postoji razmak pre svih imena atributa polja.

```js
assert.isTrue(document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.hasAttribute('alt'));
```

Nova slika mora imati vrednost od `alt` od `A slice of lasagna on a plate.`. Proverite da je vrednost svojstva elementa `alt` okružena navodnicima.

```js
assert.match(
  document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.getAttribute('alt')?.replace(/\s+/g, ' '), /^A slice of lasagna on a plate\.?$/i 
);
```

Nova slika nema atribut elementa `src`. Uverite se da postoji razmak nakon naziva otvarajuće oznake i/ili da postoji razmak pre svih imena atributa elemenata.

```js
assert.isTrue(document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.hasAttribute('src'));
```

Picha mpya inapaswa kuwa na thamani ya `src` ya `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`. Hakikisha thamani ya sifa ya kipengele `src` imezungukwa na alama za nukuu.

```js
assert.strictEqual(
  document.querySelectorAll('section')[1].querySelector('ul').nextElementSibling?.getAttribute('src'),
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Ingawa umeweka sifa ya kipengele `src` ya picha mpya kuwa URL sahihi, inashauriwa kila mara kuzizunguka thamani ya sifa ya kipengele na alama za nukuu.

```js
assert.isTrue(!/\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
