---
id: 68c18b1d0fce081f85312d97
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

U prethodnoj lekciji ste naučili da je atribut prefiksiran sa jednom donjom crtom namenjen za interno korišćenje po konvenciji.

Izmenite oba atributa `name` i `level` u `_name` i `_level`, jer se oni ne bi trebalo menjati izvan svoje klase.

Napomena je da ovo ne sprečava pristup ili modifikaciju atributa izvan klase. Takođe, u Pythonu uvek postoji način za pristup privatnim atributima (prefiksiranim sa dvostrukom donjom crtom).

# --hints--

Your `__init__` method should set `self._name` to `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._name = name")`)) })
```

Your `__init__` method should set `self._level` to `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level

charlie_brown = Employee('Charlie Brown', 'trainee')
--fcc-editable-region--
```
