---
id: 68c15a2a1f0e70ca7154aba5
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovoj radionici, krećete sa izradi sistema za praćenje plata zaposlenima.

Počnite kreiranjem klase pod nazivom `Employee`. Unutar nje kreirajte metod `__init__` sa parametrima `self`, `name` i `level`. Unutar metoda `__init__`, dodelite `name` i `level` instancijalnim atributima sa istim imenom.

# --hints--

You should have a class named `Employee`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Employee")`)) })
```

Your `Employee` class should have an `__init__` method with three parameters `self`, `name`, and `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_args("self, name, level")`)) })
```

Your `__init__` method should set `self.name` to `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

Your `__init__` method should set `self.level` to `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
