---
id: 68c28d659d345270d4a61084
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Metoda `__repr__` je posebna metoda koja bi trebalo da vrati string reprezentaciju objekta koji se može koristiti za njegovo instanciranje.

Na primer, metoda `__repr__` klase `Employee('Charlie', 'developer')` bi trebalo da vrati string `Employee('Charlie', 'developer')`, koji je isti string korišćen za kreiranje objekta.

Dajte svojoj klasi `Employee` metodu `__repr__` sa parametrom `self`, i naterajte je da vrati string koji se može koristiti za instanciranje objekta.

# --hints--

Your `Employee` class should have a `__repr__` method with a `self` parameter.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__repr__").has_args("self")`)) })
```

Your `__repr__` method should return the string used to instantiate the object.

```js
({ test: () => runPython(`
  empl = Employee("Frank", "dreamer")
  assert repr(empl) == 'Employee("Frank", "dreamer")' or repr(empl) == "Employee('Frank', 'dreamer')"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
