---
id: 68c7f7012a700243eff1cbc0
title: Korak 16
challengeType: 20
dashedName: step-16
---

# --description--

Nivo koji se koristi za instanciranje zaposlenog treba da bude izabran među specifične nivoe. Koristićete `_base_salaries` za validaciju nivoa i postavljanje prave plate za zaposlenog.

Nakon vaše postojeće `if` izjave, kreirajte još jednu `if` koja proverava da li `level` nije u `Employee._base_salaries`.

Unutar nove `if` izjave, pokrenite (ili: izbacite) `ValueError` sa porukom `Invalid value '{level}' for 'level' attribute.`, gde bi `{level}` trebalo da bude zamenjen vrednošću argumenta `level`.

# --hints--

When `level` is not in `Employee._base_salaries`, you should raise a `ValueError` with the message `Invalid value '{level}' for 'level' attribute.`, where `{level}` should be replaced by the value of the `level` argument.

```js
({ test: () => runPython(`
  try:
    Employee("Frank", "dreamer")
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
    assert False, "Expected to raise ValueError with invalid level"
`) })
```

You should not raise any exception when `level` is in `Employee._base_salaries`.

```js
({ test: () => runPython(`
  levels = [
    "trainee",
    "junior",
    "mid-level",
    "senior"
  ]
  for level in levels:
    try:
      Employee("Frank", level)
    except Exception:
      assert False, f"Expected not to raise ValueError with level '{level}'"
    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
--fcc-editable-region--
    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
