---
id: 60f02e7361b68405e27b62a5
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovom radionici, naučićete kako da stilizujete forme dizajniranjem obrasca za registraciju.

Sav HTML šablon (`DOCTYPE`, `html`, `head`, i `body`) je dostavljen za vas.

Unutar `body`, dodajte naslov kako biste dali kontekst obrascu korišćenjem `h1` elementa sa tekstom `Registration Form`.

# --hints--

You should add the `h1` within the `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

You should give the `h1` the text `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
