---
id: 60f1922fcbd2410527b3bd89
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

U prethodnim lekcijama, naučili ste kako da radite sa viewport jedinicama.

Zapamtite da jedinica `vh` označava visinu viewport-a, i jednaka je 1% od `height` viewport-a. To čini relativnim u odnosu na visinu viewport-a.

Vreme je da oživite projekat nekim CSS-om. Počnite tako što ćete `body` dati `width` od `100%`, i `height` od `100vh`.

# --hints--

You should use the `body` element selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

You should give the `body` a `width` of `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

You should give the `body` a `height` of `100vh`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
