---
id: 60f80e0081e0f2052ae5b505
title: Korak 14
challengeType: 0
dashedName: step-14
---

# --description--

Navođenje atributa `type` za element `input` važno je kako bi pretraživač znao kakve vrste podataka da očekuje. Ako nije navedeno `type`, pretraživač će podrazumevano koristiti `text`.

Prva dva `input` elementa treba da imaju atribut `type` od `text`, treći atribut `type` od `email`, a četvrti atribut `type` od `password`.

Tip `email` dozvoljava samo e-mail adrese sa `@` i `.` u domeni.
Tip `password` maskira unos i upozorava ako sajt ne koristi HTTPS.

# --hints--

You should give the first `input` element a `type` attribute of `text`.

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

You should give the second `input` element a `type` attribute of `text`.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

You should give the third `input` element a `type` attribute of `email`.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

You should give the fourth `input` element a `type` attribute of `password`.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label for="first-name">Enter Your First Name: <input id="first-name" /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" /></label>
      <label for="email">Enter Your Email: <input id="email" /></label>
      <label for="new-password">Create a New Password: <input id="new-password" /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
