---
id: 60f83e7bfc09900959f41e20
title: Korak 17
challengeType: 0
dashedName: step-17
---

# --description--

Određene vrednosti atributa `type` dolaze sa ugrađenom validacijom forme. Na primer, `type="email"` zahteva da je vrednost validna e-mail adresa.

Dodajte prilagođenu validaciju elementu za lozinku `input`, dodavanjem atributa `minlength` sa vrednošću `8`. Time se sprečava slanje unosa koji imaju manje od 8 karaktera.

# --hints--

You should give the password `input` element a `minlength` attribute.

```js
assert.notEqual(document.querySelector('input[type="password"]')?.minLength, -1);
```

You should give the `minlength` attribute a value of `8`.

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, 8);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
--fcc-editable-region--
      <label for="new-password">Create a New Password: <input id="new-password" type="password" required /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
