---
id: 60fab4a123ce4b04526b082b
title: Korak 24
challengeType: 0
dashedName: step-24
---

# --description--

Trenutno korisnici mogu poslati formular bez označavanja radijskih dugmadi. Iako ste ranije koristili atribut `required` da biste ukazali na to da je unos obavezan, neće raditi u ovom slučaju jer dodavanje `required` oba unosa će preneti pogrešne informacije korisnicima.

Da biste ovo rešili, možete pružiti kontekst onoga što je potrebno dodavanjem elementa `legend` sa tekstom `Account type (required)` pre elemenata `label` unutar drugog `fieldset`. Zatim dodajte atribut `checked` na unos `Personal` kako biste osigurali da je formular poslat sa potrebnim podacima.

# --hints--

Your `input` elements should remain `type` of `radio`.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-child(2) input[type="radio"]'), 2);
```

You should add a `legend` element within the second `fieldset` element.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(2) > legend'), 1);
```

You should add `Account type (required)` text to the `legend` element.

```js
assert.equal(document.querySelector('fieldset:nth-of-type(2) > legend')?.textContent, 'Account type (required)');
```

You should give an attribute of `checked` to the `Personal` input.

```js
assert.isTrue(document.querySelector('fieldset:nth-child(2) input[type="radio"]')?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
--fcc-editable-region--
      
--fcc-editable-region--
      <label><input type="radio" name="account-type" value="personal" /> Personal</label>
      <label><input type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
