---
id: 63541ef4f96cd82e8e6c788a
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Atribut `method` određuje kako poslati podatke forme na URL naveden u atributu `action`. Podaci forme mogu se poslati putem zahteva `GET` kao URL parametri (sa `method="get"`) ili putem zahteva `POST` kao podaci u telu zahteva (sa `method="post"`).

Postavite atribut `method` da pošaljete podatke forme putem zahteva `POST`.

**NAPOMENA**: Vi ćete saznati više o metodama `GET` i `POST` u nadolazećim lekcijama i radionicama.

# --hints--

You shouldn't add a new `form` element.

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

Your `form` element should have a `method` attribute.

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

Your `method` attribute should be set to `post`.

```js
assert.match(document.querySelector('form')?.getAttribute('method'), /^post$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
--fcc-editable-region--
  <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
