---
id: 6989f1ba358acba651d5ee6b
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Sada ćete postaviti strukturu podataka za skladištenje vaših rezultata. Kreirajte varijablu pod imenom `result` i inicijalizujte je praznim nizom. Ovaj niz će skladištiti sve važeće kombinacije zagradi koje generišete.

Ažurirajte vaš `return` izraz da vrati `result` umesto praznog niza.

# --hints--

You should declare a variable named `result` in your `genParentheses` function.

```js
assert.match(genParentheses.toString(), /result/);
```

You should initialize `result` to an empty array.

```js
assert.match(genParentheses.toString(), /result\s*=\s*\[\]/);
```

You should return `result` at the end of your function.

```js
assert.match(genParentheses.toString(), /return\s*result/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
--fcc-editable-region--
  
  return [];
--fcc-editable-region--
}
```
