---
id: 6989f207306332c06b0add39
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Za pristup pretrazi širine (breadth-first search), koristićeš red (queue) za praćenje različitih stanja dok gradiš kombinacije zagradi. Svako stanje će biti predstavljeno nizom koji sadrži tri elementa:

- Trenutni string u izgradnji
- Broj korišćenih otvarajućih zagradi
- Broj korišćenih zatvarajućih zagradi

Kreiraj varijablu pod nazivom `queue` i inicijalizuj je nizom koji sadrži jedan niz: `['', 0, 0]`. Ovo predstavlja početno stanje sa praznim stringom i nula korišćenih zagradi.

# --hints--

You should declare a variable named `queue` in your `genParentheses` function.

```js
assert.match(genParentheses.toString(), /queue/);
```

You should initialize `queue` with an array containing the array `['', 0, 0]`.

```js
assert.match(genParentheses.toString(), /queue\s*=\s*\[\s*\['',\s*0,\s*0\]\s*\]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

--fcc-editable-region--
  
--fcc-editable-region--
  let result = [];
  
  return result;
}
```
