---
id: 6989f5bf7f1a1d0fc8ebd1e8
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

Sada morate da obradite dodavanje zatvorene zagrade. Ključno pravilo je da možete dodati zatvorenu zagradu samo ako održava ravnotežu, što znači da mora biti više korišćenih otvarajućih zagradi nego zatvarajućih.

Još unutar bloka `else`, dodajte drugu `if` instancu da proverite da li je `closesUsed` manje od `opensUsed`.

Ako je ovaj uslov tačan, dodajte još jedan novi niz u `queue`: `[current + ')', opensUsed, closesUsed + 1]`. Ovo predstavlja stanje nakon dodavanja zatvorene zagrade.

# --hints--

You should have a second `if` statement with the condition `closesUsed < opensUsed` inside your `else` block.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*closesUsed\s*<\s*opensUsed\s*\)/);
```

You should push `[current + ')', opensUsed, closesUsed + 1]` to the queue when the condition is true.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\)['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
    if (current.length === 2 * pairs) {
      result.push(current);
    } else {
      if (opensUsed < pairs) {
        queue.push([current + '(', opensUsed + 1, closesUsed]);
      }
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
  return result;
}

console.log(genParentheses(1));
```
