---
id: 6989f2e35e6941e13e445156
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Sada ćete implementirati glavnu BFS petlju. Kreirajte `while` petlju koja se izvršava dok `queue` nije prazan, odnosno, `queue.length` je veći od `0`.

Unutar petlje, logujte `queue` na konzolu.

# --hints--

You should create a `while` loop with the condition `queue.length > 0`.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
```

You should log `queue` to the console inside your `while` loop.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
assert.match(genParentheses.toString(), /console\.log\s*\(\s*queue\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
--fcc-editable-region--
  
--fcc-editable-region--
  return result;
}
```
