---
id: 6989eb65324d97775c4965fe
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovom radionici, implementiraćete funkciju koja generiše sve validne kombinacije zagradi koristeći pristup širinskom pretraživanju (BFS). Na primer, validne kombinacije za par zagrada su `(())` i `()()`.

Počnite kreiranjem funkcije pod nazivom `genParentheses` sa jednim parametrom `pairs`. Za sada, vratite prazan niz iz funkcije.

# --hints--

You should define a function named `genParentheses`.

```js
assert.isFunction(genParentheses);
```

Your `genParentheses` function should have a single parameter named `pairs`.

```js
const genParenthesesParams = __helpers.getFunctionParams(genParentheses.toString());
assert.equal(genParenthesesParams[0].name, 'pairs')
```

Your function should return an empty array.

```js
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
