---
id: 6989f48c3f0292c84f8d8938
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Sada morate identifikovati kada ste izgradili kompletnu kombinaciju zagradi. Kompletnu kombinaciju ima dužina jednaka dvostrukom broja parova (jer svaki par doprinosi jednu otvarajuću i jednu zatvarajuću zagradu).

Unutar vaše `while` petlje, dodajte `if` izjavu da proverite da li je `current.length === 2 * pairs`. Kada je ovo uslov tačan, dodajte (push) `current` u niz `result`.

# --hints--

You should have an `if` statement with the condition `current.length === 2 * pairs` inside your `while` loop.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)/);
```

You should push `current` to the `result` array inside your new `if` statement.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?result\.push\s*\(\s*current\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
--fcc-editable-region--
    
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
