---
id: 6989ef885db3a3488a13b426
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Pre implementacije osnovnog algoritma, morate validirati ulaz. Parametar `pairs` bi trebalo da bude ceo broj, jer predstavlja broj parova zagradi koje treba generisati.

Dodajte izjavu `if` na početak vaše funkcije da proverite da li `pairs` nije ceo broj. Za to koristite metod `Number.isInteger()`.

Ako je uslov tačan, vratite string `The number of pairs should be an integer`.

# --hints--

You should have an if statement with the condition `!Number.isInteger(pairs)` inside the `genParentheses` function.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*!\s*Number\.isInteger\s*\(\s*pairs\s*\)\s*\)/);
```

You should return a string when the input is not an integer.

```js
let result = genParentheses("i");
assert.isString(result);
```

You should return the string `The number of pairs should be an integer` when the input is a string.

```js
let result = genParentheses("i");
assert.strictEqual(result, 'The number of pairs should be an integer');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
--fcc-editable-region--
  
--fcc-editable-region--
  
  return [];
}
```

