---
id: 6989f10a72ad424d4ad6d514
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Zatim, morate da validirate da je broj parova najmanje jedan, jer ne možete generisati kombinacije zagradi sa nula ili negativan broj parova.

Dodajte još jednu `if` instancu da proverite da li je `pairs` manje od `1`. Ako je ovaj uslov tačan, vratite niz znakova (string) `The number of pairs should be at least 1`.

# --hints--

You should have a second `if` statement with the condition `pairs < 1`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*pairs\s*<\s*1\s*\)/);
```

You should return the string `The number of pairs should be at least 1` when `pairs` is less than `1`.

```js
let result = genParentheses(0);
assert.strictEqual(result, 'The number of pairs should be at least 1');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  
  return [];
}
```
