---
id: af7588ade1100bde429baf20
title: Izgraditi detektor nedostajućih slova
challengeType: 26
dashedName: build-a-missing-letter-detector
---

# --description--

**Cilj:** Ispunite priče korisnika ispod i osigurajte da svi testovi prođu kako biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da imate funkciju pod nazivom `fearNotLetter`.
1. Funkcija `fearNotLetter` treba da prihvati jedan argument: niz znakova koji predstavlja opseg slova u alfabetskom redosledu, a koji može imati jedno nedostajuće slovo.
1. Funkcija treba da pronađe nedostajuće slovo u prosleđenom opsegu slova i vrati ga.
1. Ako su sva slova prisutna u opsegu, funkcija treba da vrati `undefined`.

# --hints--

You should have a `fearNotLetter` function.

```js
assert.isFunction(fearNotLetter);
```

`fearNotLetter("abce")` should return the string `d`.

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

`fearNotLetter("abcdefghjklmno")` should return the string `i`.

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

`fearNotLetter("stvwx")` should return the string `u`.

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

`fearNotLetter("bcdf")` should return the string `e`.

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

`fearNotLetter("abcdefghijklmnopqrstuvwxyz")` should return `undefined`.

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    let letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
