---
id: 694c84330cbe604821b9deef
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Kontrolni centar zahteva metod za dodavanje novih članova u spisak. Kreirajte praznu funkciju pod nazivom `addCrewMember` koja prima dva parametra, `crew` i `astronaut`.

# --hints--

You should create a function named `addCrewMember`.

```js
assert.isFunction(addCrewMember);
```

Your `addCrewMember` function should have parameters `crew` and `astronaut`.

```js
const regex = __helpers.functionRegex('addCrewMember', ['crew', 'astronaut']);
assert.match(__helpers.removeJSComments(addCrewMember.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
