---
id: 694c9c186c089f7bdff0f3c6
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Generalno, dobra je praksa da validirate svoje ulazne podatke unutar funkcije. Koristite petlju `for` da iterirate kroz niz `crew` i proverite da li bilo koji član već ima isti `id` kao ulazni `astronaut`. Ako se pronađe duplikat, logujte poruku o grešci `console.log("Duplicate ID: " + astronaut.id)` i zatim pozovite `return` da biste rano izašli iz funkcije.

Evo primera ove često korišćene tehnike:

```js
if (existingMember.id === newMember.id) {
    console.log("Duplicate ID: " + newMember.id);
    return;
}
```

# --hints--

Your function should use a `for` loop to scan existing `crew` members (any valid `for (...)` form is fine).

```js
assert.match(__helpers.removeJSComments(code), /for\s*\([^)]*\)/);
```

Inside the `for` loop, you should check for a duplicate `id`. If one is found, log `"Duplicate ID: " + astronaut.id"` to the console and then call `return` to exit the function.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewWithDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 },
    { id: 3 }
  ];
  addCrewMember(crewWithDuplicate, { id: 3 });

  assert.lengthOf(
    spy.calls,
    1,
    "Duplicate ID input should trigger exactly one console.log and then exit"
  );
  assert.match(
    spy.calls[0][0],
    /duplicate\s+id\s*:\s*3/i,
    "Console message should mention 'duplicate id'"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

When given valid inputs, your `addCrewMember` function should not invoke a `console.log()` error message or an early `return` exit.

```js
const spy = __helpers.spyOn(console, "log");

try {
  const crewNoDuplicate = [
    { id: 1 },
    { id: 2 },
    { id: 3 },
    { id: 4 }
  ];
  addCrewMember(crewNoDuplicate, { id: 5 });

  assert.lengthOf(
    spy.calls,
    0,
    "A truly new astronaut should produce zero console.log calls"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
--fcc-editable-region--

--fcc-editable-region--
}
```
