---
id: 694c40d86fe6a78d20b6559c
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, koristićete JavaScript da izgradite spisak astronauta za svemirsku misiju. 

Da biste počeli, kreirajte prazan niz za vaš spisak pod nazivom `squad`.

# --hints--

You should create a `squad` variable.

```js
assert.isDefined(squad);
```

Your `squad` variable should be an array.

```js
assert.isArray(squad);
```

Your `squad` array should be empty.

```js
assert.deepEqual(squad, []);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
