---
id: 694d8ddd106340237efb0c80
title: Korak 12
challengeType: 1
dashedName: step-12
---

# --description--

Metoda ``splice()`` može modifikovati nizove dodavanjem ili uklanjanjem elemenata na bilo kojoj poziciji, uključujući sredinu. Pošto ``splice()`` vraća niz koji sadrži uklonjene elemente, možete ga koristiti za zamenu dva elementa u nizu bez mutiranja originalnog. Ovo se može uraditi u jednoj liniji korišćenjem sledeće tehnike:

```js
// swap elements at i and j without mutating the original
const copy = array.slice();
copy[i] = copy.splice(j, 1, copy[i])[0];
```

Evo primera:

```js
const originalArray = [12, 97, 68, 55];
const copyArray = originalArray.slice();
copyArray[1] = copyArray.splice(3, 1, copyArray[1])[0];
console.log(copyArray); // [12, 55, 68, 97]
```

Tehnika, primenjena u trećoj liniji, radi na sledeći način:

- ``splice(3, 1, copyArray[1])`` uklanja element sa indeksa ``3`` (``55``)

- On ubacuje element iz ``copyArray[1]`` (``97``) na indeks ``3``

- ``splice()`` vraća ``[55]``, niz koji sadrži uklonjeni element

- ``[0]`` izvlači ``55`` iz tog niza

- Ta vrednost (``55``) se ponovo dodeljuje ``copyArray[1]``, čime se završava zamena

Koristite ovu tehniku sa nizom ``updatedCrew`` da zamenite astronautima na pozicijama ``fromIndex`` i ``toIndex``.

# --hints--

You should perform the swap by calling the `splice()` method.

```js
assert.match(swapCrewMembers.toString(), /splice\s*\(/);
```

You should call `splice()` with the `updatedCrew` array to perform the swap. You should use the one-line technique.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(swapCrewMembers.toString()));
const regex = /updatedCrew\[(fromIndex|toIndex)\]=updatedCrew\.splice\((fromIndex|toIndex),1,updatedCrew\[\1\]\)\[0\];?\s*\}$/;
assert.match(cleaned, regex);
```

Your `swapCrewMembers` function should not mutate the original input `crew` array.

```js
const demo = [{ id: 1 }, { id: 2 }, { id: 3 }];
swapCrewMembers(demo, 0, 2);
assert.deepEqual(demo.map(a => a.id), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
--fcc-editable-region--

--fcc-editable-region--
}
```
