---
id: 68e45a2b4ab9b2f48b1cc6df
title: Napravite stranicu za podeЕЎavanje cena paketa
challengeType: 25
dashedName: design-a-pricing-plans-layout-page
demoType: onClick
---

# --description--

**Cilj:** Popunite korisniДЌke priДЌe navedene ispod i proД‘ite sve testove da biste zavrЕЎili laboratoriju.

**PriДЌe korisnika:**

1. Ukurasa wako unapaswa kuwa na kipengele cha `h1` chenye maandishi ya `Pricing Plans`.
2. Ukurasa wako unapaswa kuwa na kipengele cha `div` chenye darasa la vitu la `pricing-container` chini ya kipengele cha `h1`.
    - Kichaguzi cha `.pricing-container` kinapaswa kuwa na sifa ya `display` yenye thamani ya `flex` na sifa ya `flex-wrap` yenye thamani ya `wrap`.
3. Ndani ya kipengele cha `.pricing-container`, unapaswa kuwa na vipengele vitatu vya `div` vyenye darasa la vitu la `pricing-card` kuwakilisha mipango ya bei.
    - Kimoja cha vipengele vya `.pricing-card` kinapaswa kuwa na darasa la vitu la `basic-plan` pamoja na darasa la vitu la `pricing-card`.
    - Kimoja cha vipengele vya `.pricing-card` kinapaswa kuwa na darasa la vitu la `pro-plan` pamoja na darasa la vitu la `pricing-card`.
    - Kimoja cha vipengele vya `.pricing-card` kinapaswa kuwa na darasa la vitu la `premium-plan` pamoja na darasa la vitu la `pricing-card`.
4. Kipengele chako cha `.basic-plan` kinapaswa kuwa na kipengele cha `h2` chenye maandishi ya `Basic`.
5. Kipengele chako cha `.basic-plan` kinapaswa kuwa na kipengele cha `p` chenye maandishi ya `$9/month`.
6. Kipengele chako cha `.pro-plan` kinapaswa kuwa na kipengele cha `h2` chenye maandishi ya `Pro`.
7. Kipengele chako cha `.pro-plan` kinapaswa kuwa na kipengele cha `p` chenye maandishi ya `$19/month`.
8. Kipengele chako cha `.premium-plan` kinapaswa kuwa na kipengele cha `h2` chenye maandishi ya `Premium`.
9. Kipengele chako cha `.premium-plan` kinapaswa kuwa na kipengele cha `p` chenye maandishi ya `$29/month`.
10. Kila kipengele cha `.pricing-card` kinapaswa:
    - Kutumia Flexbox kupanga yaliyomo kwa safu na kuhalalisha nafasi kati ya watoto kwa kutumia `space-between`.
    - Weka sifa ya `flex` kuwa `0 0 200px` ili kiwe na upana thabiti na kuzuia kukua au kupungua katika mpangilio.
    - Weka sifa ya `border` kuwa `2px solid black` kuona jinsi kadi tofauti zinavyotumia nafasi.
11. Kipengele cha `.basic-plan` kinapaswa kuonekana kwanza katika mpangilio. Unapaswa kutumia sifa ya `order` kwa hili.
12. Kipengele cha `.pro-plan` kinapaswa kuonekana cha pili katika mpangilio. Unapaswa kutumia sifa ya `order` na kuweka sifa yake ya `flex-grow` kuwa `2` ili ichukue nafasi zaidi kuliko mipango mingine.
13. Kipengele cha `.premium-plan` kinapaswa kuonekana mwisho katika mpangilio. Unapaswa kutumia sifa ya `order` kwa hili.

**Napomena:** Uverite se da ste povezali svoj stilski list na HTML i koristili CSS.

# --hints--

Trebalo bi da imaЕЎ element `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Stranica vam treba da ima element sa tekstom `Pricing Plans`, koji je `h1`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Pricing Plans');
```

Trebalo bi da imaЕЎ element `div` ispod elementa `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

Tvoj element za `div` treba da ima klasu objekata `pricing-container`.

```js
assert.exists(document.querySelector('h1 + div.pricing-container'));
```

Trebalo bi da imaЕЎ biraДЌ za `.pricing-container`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-container'));
```

Tvoj selektor za `pricing-container` mora imati svojstvo `display` sa vrednoЕЎД‡u `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.display, 'flex');
```

VaЕЎ selektor za `.pricing-container` mora imati atribut `flex-wrap` sa vrednoЕЎД‡u `wrap`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-container')?.flexWrap, 'wrap');
```

Unapaswa kuwa na vipengele vitatu vya `div` vyenye darasa la vitu la `pricing-card` ndani ya kipengele cha `.pricing-container`.

```js
const divEls = document.querySelectorAll('.pricing-container .pricing-card');
assert.lengthOf(divEls, 3);
```

Komponenta vaЕЎih `.pricing-card` treba da ima klasu objekata od `basic-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.basic-plan'));
```

Komponenta vaЕЎih elemenata za `.pricing-card` treba imati klasu objekata za `premium-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.premium-plan'));
```

Komponenta vaЕЎih `.pricing-card` treba da ima klasu od `pro-plan`.

```js
assert.exists(document.querySelector('div.pricing-card.pro-plan'));
```

Trebalo bi da imate biraДЌ za `.pricing-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.pricing-card'));
```

VaЕЎ selektor za `.pricing-card` mora imati svojstvo `display` sa vrednoЕЎД‡u `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.display, 'flex');
```

VaЕЎ selektor za `.pricing-card` mora imati svojstvo od `flex-direction` sa vrednoЕЎД‡u `column`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flexDirection, 'column');
```

Tvoj selector za `.pricing-card` mora imati svojstvo `justify-content` sa vrednoЕЎД‡u `space-between`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.justifyContent, 'space-between');
```

Tvoj selektor za `.pricing-card` mora imati svojstvo `flex` sa vrednoЕЎД‡u `0 0 200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.flex, '0 0 200px');
```

VaЕЎ selektor za `.pricing-card` mora imati svojstvo `border` sa vrednoЕЎД‡u `2px solid black`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pricing-card')?.border, '2px solid black');
```

Tvoj selektor za `.basic-plan` mora imati svojstvo `order` sa vrednoЕЎД‡u `0`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.basic-plan')?.order, '0');
```

VaЕЎ element za `.basic-plan` treba da ima element `h2` sa tekstom `Basic`.

```js
const h2 = document.querySelector('.basic-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Basic');
```

Komponenta za vaЕЎ `.basic-plan` treba da ima komponentu `p` sa tekstom `$9/month`.

```js
const p = document.querySelector('.basic-plan > p');
assert.strictEqual(p?.textContent.trim(), '$9/month');
```

VaЕЎ element za `.pro-plan` treba da ima komponentu `h2` sa tekstom od `Pro`.

```js
const h2 = document.querySelector('.pro-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Pro');
```

VaЕЎ element od `.pro-plan` treba da ima element od `p` sa tekstom `$19/month`.

```js
const p = document.querySelector('.pro-plan > p');
assert.strictEqual(p?.textContent.trim(), '$19/month');
```

VaЕЎ atribut `.pro-plan` mora imati svojstvo `order` sa vrednoЕЎД‡u od `1`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.order, '1');
```

VaЕЎe polje `.pro-plan` mora imati svojstvo `flex-grow` sa vrednoЕЎД‡u `2`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.pro-plan')?.flexGrow, '2');
```

VaЕЎ element od `.premium-plan` mora da ima element od `h2` sa tekstom `Premium`.

```js
const h2 = document.querySelector('.premium-plan > h2');
assert.strictEqual(h2?.textContent.trim(), 'Premium');
```

VaЕЎ element od `.premium-plan` mora imati element od `p` sa tekstom `$29/month`.

```js
const p = document.querySelector('.premium-plan > p');
assert.strictEqual(p?.textContent.trim(), '$29/month');
```

VaЕЎ element za `premium-plan` treba biti poslednji element u nizu i imati svojstvo `order` sa vrednoЕЎД‡u `2`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.premium-plan')?.order, '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Pricing Plans Layout Page</title>
  </head>
  <body></body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Pricing Plans Layout</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Pricing Plans</h1>
    <div class="pricing-container">
      <div class="pricing-card pro-plan">
        <h2>Pro</h2>
        <p class="price">$19/month</p>
        <ul class="features">
          <li><span class="check">вњ“</span> Access to our core features</li>
          <li><span class="check">вњ“</span> Email support</li>
          <li><span class="check">вњ“</span> Advanced analytics</li>
          <li><span class="check">вњ“</span> Custom integrations</li>
          <li><span class="cross">вњ—</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card basic-plan">
        <h2>Basic</h2>
        <p class="price">$9/month</p>
        <ul class="features">
          <li><span class="check">вњ“</span> Access to our core features</li>
          <li><span class="check">вњ“</span> Email support</li>
          <li><span class="cross">вњ—</span> Advanced analytics</li>
          <li><span class="cross">вњ—</span> Custom integrations</li>
          <li><span class="cross">вњ—</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>

      <div class="pricing-card premium-plan">
        <h2>Premium</h2>
        <p class="price">$29/month</p>
        <ul class="features">
          <li><span class="check">вњ“</span> Access to our core features</li>
          <li><span class="check">вњ“</span> Email support</li>
          <li><span class="check">вњ“</span> Advanced analytics</li>
          <li><span class="check">вњ“</span> Custom integrations</li>
          <li><span class="check">вњ“</span> Priority support</li>
        </ul>
        <a href="#">Choose Plan</a>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  align-items: center;
  justify-content: center;
  flex-direction: column;
  min-height: 100vh;
}

h1 {
  margin-bottom: 20px;
}

h2,
p {
  margin: 0;
}

.pricing-container {
  display: flex;
  flex-wrap: wrap;
  align-content: center;
  border: 2px solid #000;
  gap: 20px;
  width: 80%;
  max-width: 900px;
  height: 70vh;
  padding: 0 20px;
}

.pricing-card {
  background: rgb(247, 247, 247);
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
  padding: 20px;
  flex: 0 0 200px;
  text-align: center;
  transition: transform 0.3s ease;
  display: flex;
  flex-direction: column;
  justify-content: space-between;
  border: 2px solid black;
}

.pricing-card:hover {
  transform: scale(1.05);
}

.price {
  font-size: 1.5rem;
  font-weight: bold;
  color: #3e3e3e;
  margin: 10px;
  border-bottom: 1px solid #433d28;
}

.features {
  list-style-type: none;
  padding: 0;
  text-align: left;
}

.features li {
  padding: 5px 0;
  font-size: 0.95rem;
}

.check {
  color: #219c54;
  font-weight: bold;
  margin-right: 8px;
}

.cross {
  color: crimson;
  font-weight: bold;
  margin-right: 8px;
}

a {
  text-decoration: none;
  display: inline-block;
  margin-top: 15px;
  background: #f1be32;
  color: #000;
  font-weight: bold;
  border: 3px solid #eebe3a;
  padding: 5px;
  transition: background 0.3s;
}

a:hover {
  background: #f7c436;
}

.basic-plan {
  order: 0;
}
.pro-plan {
  order: 1;
  flex-grow: 2;
  background: #f0f8ff;
}
.premium-plan {
  order: 2;
}

@media screen and (max-width: 1000px) {
  .pricing-container {
    height: auto;
    padding-bottom: 20px;
  }

  h1 {
    display: block;
    font-size: 1.5rem;
    text-align: center;
  }

  .pricing-card {
    flex: 0 0 175px;
  }
}

@media screen and (max-width: 650px) {
  .pricing-container {
    flex-direction: column;
    height: auto;
  }

  .pricing-card {
    flex: 1 1 auto;
    width: 100%;
  }
}
```
