---
id: a97fd23d9b809dac9921074f
title: Implementacija funkcije za sabiranje sa opcionalnim argumentima
challengeType: 26
dashedName: build-an-optional-arguments-sum-function
---

# --description--

U ovom laboratorijumu ćete izgraditi funkciju koja prima do dva argumenta i sabira ih, ali ako postoji samo jedan argument, vraća funkciju koja čeka na drugi broj da bi izvršila zbir.

**Cilj:** Ispunite priče korisnika ispod i osigurajte da svi testovi prođu kako biste završili laboratorijum.

**Priče korisnika:**

1. Kreirajte funkciju `addTogether`.
2. Ako funkcija primi dva argumenta, funkcija bi trebalo da vrati zbir oba argumenta.
3. Ako je dostavljen samo jedan argument, `addTogether` bi trebalo da vrati funkciju.
   - Kada se vraćena funkcija pozove sa jednim argumentom, ona bi trebalo da vrati zbir.

   ```js
   const sumTwoAnd = addTogether(2);
   sumTwoAnd(3); // 5
   ```

4. Ako nijedan od argumenata nije validan broj, vratite `undefined`.

# --hints--

You should have the `addTogether` function.

```js
assert.isFunction(addTogether);
```

`addTogether(2, 3)` should return 5.

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)` should return 53.4.

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)` should return `undefined`.

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)` should return `undefined`.

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")` should return `undefined`.

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)` should return a function.

```js
assert.isFunction(addTogether(5));
```

`addTogether(5)(7)` should return 12.

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])` should return `undefined`.

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")` should return `undefined`.

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
