---
id: 579e2a2c335b9d72dd32e05c
title: Implementirajte algoritam Slice i Splice
challengeType: 26
dashedName: implement-the-slice-and-splice-algorithm
---

# --description--

U ovom laboratorijumu, biće vam potrebno da kreirate algoritam za spajanje dva niza.

**Cilj:** Ispunite priče korisnika ispod i učinite da svi testovi prođu kako biste završili laboratorijum.

**Priče korisnika:**

1. Kreirajte funkciju ``frankenSplice`` koja prima dva niza i indeks.
1. Kopirajte svaki element prvog niza u drugi niz, po redu, počevši od datog indeksa, i vratite rezultujući niz.
1. Ulazni nizovi bi trebalo da ostanu isti nakon izvršavanja funkcije.

# --hints--

`frankenSplice([1, 2, 3], [4, 5], 1)` should return `[4, 1, 2, 3, 5]`.

```js
assert.deepEqual(frankenSplice([1, 2, 3], [4, 5], 1), [4, 1, 2, 3, 5]);
```

`frankenSplice([1, 2], ["a", "b"], 1)` should return `["a", 1, 2, "b"]`.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
assert.deepEqual(frankenSplice(testArr1, testArr2, 1), ['a', 1, 2, 'b']);
```

`frankenSplice(["claw", "tentacle"], ["head", "shoulders", "knees", "toes"], 2)` should return `["head", "shoulders", "claw", "tentacle", "knees", "toes"]`.

```js
assert.deepEqual(
  frankenSplice(
    ['claw', 'tentacle'],
    ['head', 'shoulders', 'knees', 'toes'],
    2
  ),
  ['head', 'shoulders', 'claw', 'tentacle', 'knees', 'toes']
);
```

All elements from the first array should be added to the second array in their original order. `frankenSplice([1, 2, 3, 4], [], 0)` should return `[1, 2, 3, 4]`.

```js
assert.deepEqual(frankenSplice([1, 2, 3, 4], [], 0), [1, 2, 3, 4]);
```

The first array should remain the same after the function runs.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr1, [1, 2]);
```

The second array should remain the same after the function runs.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr2, ['a', 'b']);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function frankenSplice(arr1, arr2, n) {
  let result = arr2.slice();
  for (let i = 0; i < arr1.length; i++) {
    result.splice(n + i, 0, arr1[i]);
  }
  return result;
}
```
