---
id: 5f33071498eb2472b87ddee4
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovom projektu, ćeš naučiti osnove CSS (Cascading Style Sheets) za kreiranje menija kafe. CSS je jezik koji se koristi za stilizovanje dokumenta HTML. On definiše kako bi elementi HTML trebalo da budu prikazani na ekranu.

Kako ste naučili u poslednjim koracima Cat Photo App, postoji osnovna struktura potrebna za početak izgradnje vaše veb stranice. Svaki dokument od HTML bi trebalo da ima deklaraciju od `DOCTYPE` i element od `html`. `DOCTYPE` objašnjava koju verziju pregledača sadrži taj dokument za HTML. A element od `html` povezuje na koreni element koji uključuje sve ostale elemente.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

Povećanje etikete za `<!DOCTYPE html>`, na polje `html` koje ima svojstvo polja `lang` sa vrednošću od `en`.

# --hints--

Trebalo bi da imate deklaraciju `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Trebalo bi da imate etiketu pakovanja za `<html lang="en">` na engleskom jeziku.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Trebalo bi da imaš oznaku za zatvaranje `</html>`.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
