---
id: 5f3477aefa51bfc29327200b
title: Korak 14
challengeType: 0
dashedName: step-14
---

# --description--

Tri elementa su modifikovana pisanjem `CSS` unutar tagova ``style``. Ovo radi, ali pošto će biti više stilova, bolje je sve stilove staviti u odvojeni fajl i povezati ga.

Različan fajl od `styles.css` je kreiran za vas. Možete prebacivati između fajlova koristeći kartice na vrhu editora.

Počni ponovo da napišeš stilove koje si kreirao u fajlu `styles.css`. Osiguraj da ne uključuješ otvarajuće i zatvarajuće oznake za `style`.

# --hints--

Tvoj fajl `styles.css` bi trebalo da ima biranje tipa `h1, h2, p`.

```js
assert(code.slice(code.indexOf('</html>')).replace(/[\s\n]*/g, "").match(__helpers.concatRegex(__helpers.permutateRegex(['h1', 'h2', 'p'], { elementsSeparator: ',' }), /{/)));
```Tvoj selektor mora postaviti atribut `text-align` na `center`.

```js
assert(code.slice(code.indexOf('</html>')).match(/text-align:\s*center;?/));
```

Tvoj kod ne bi smega imati oznake `<style>`.

```js
assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
