---
id: 5d822fd413a79914d39e98d1
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Dajte elementu `.background-buildings` atribut `width` i `height` sa `100%` kako bi ga učinili punom širinom i visinom svog roditelja, `body`.

# --hints--

You should use the `background-buildings` class to select the correct element.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

Your `.background-buildings` element should have a `width` of `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.width, '100%');
```

Your `.background-buildings` element should have a `height` of `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.height, '100%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings"></div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--


--fcc-editable-region--
    
```

