---
id: 5d822fd413a79914d39e98cd
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

U CSS-u, možete ciljati sve sa zvezdastim znakom (asterisk). Dodajte granicu svome sve koristeći selektor `*`, i dajući mu `border` od `1px solid black`. Ovo je trik koji pomaže u vizualizaciji gde su elementi i njihova veličina. To ćete kasnije ukloniti.

# --hints--

You should use the `*` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

You should use the `border` property to style all the elements.

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

All elements should have a `1px solid black` border.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
