---
id: 5d822fd413a79914d39e98cf
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Možete videti `body` (to je najunutrašnja kutija na vašoj stranici); kutija oko nje je element `html`. Neka vaš `body` popuni ceo vidokrug (viewport) tako što mu date `height` od `100vh`. Uklonite podrazumevani `margin` iz `body` postavljanjem `margin` na `0`. Na kraju, postavite svojstvo `overflow` na `hidden` da sakrijete sve trake za skrolovanje koje se pojavljuju kada nešto prelazi vidokrug.

# --hints--

You should use the `body` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Your `body` should have a `height` of `100vh`.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

Your `body` should have a `margin` of `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

Your `body` should have the `overflow` property set to `hidden`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--
    
```

