---
id: 56533eb9ac21ba0edf2244b3
title: Napravite konverter iz Celzijusa u Farenhajta
challengeType: 26
dashedName: build-a-celsius-to-fahrenheit-converter
---

# --description--

U ovom laboratorijumu ćete napisati funkciju koja konvertuje temperaturu iz Celzijusa u Farenhajta. Formula za konverziju iz Celzijusa u Farenhajt je:

```js
fahrenheit = celsius * (9/5) + 32
```

**Cilj:** Ispunite priče korisnika ispod i prođite sve testove kako biste završili laboratorijum.

**Priče korisnika:**

1. Treba da kreirate funkciju pod nazivom `convertCtoF`.
2. `convertCtoF` treba da primi jedan numerički argument, koji je temperatura u Celzijusima.
3. `convertCtoF` treba da vrati broj.

# --hints--

You should create a function named `convertCtoF`.

```js
assert.isFunction(convertCtoF);
```

`convertCtoF` should take a single parameter.

```js
assert.lengthOf(convertCtoF, 1);
```

`convertCtoF(0)` should return a number.

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` should return a value of `-22`.

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` should return a value of `14`.

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` should return a value of `32`.

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` should return a value of `68`.

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` should return a value of `86`.

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
