---
id: 66c07238b01053abaf812065
title: Izgradite kalkulator faktoriela
challengeType: 26
dashedName: build-a-factorial-calculator
---

# --description--

Faktorijel je proizvod celog broja i svih celih brojeva ispod njega. Na primer, faktorijel od `5` je `5 * 4 * 3 * 2 * 1 = 120`. U ovom laboratorijumu, kreiraćete kalkulator za faktorijele koji uzima broj od korisnika i izračunava faktorijel tog broja.

**Cilj:** Ispunite priče korisnika ispod i prođite sve testove kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da deklarišete varijablu `num` i dodelite joj broj po izboru. Dodeljeni broj treba da bude između `1` i `20`, uključujući oba broja.
1. Kreirajte funkciju pod nazivom `factorialCalculator` koja uzima broj kao argument i vraća faktorijel tog broja.
1. Unutar funkcije, deklarišite varijablu `result` i dodelite joj vrednost `1`. Koristeći petlju (loop), prođite kroz sve brojeve od `1` do ulaznog broja (uključujući oba) i za svaki broj, pomnožite varijablu `result` sa trenutnim brojem i dodelite rezultat varijabli `result`. Možete izabrati da koristite petlju tipa `for`, petlju tipa `while` ili petlju tipa `do...while` ovde.
1. Trebalo bi da pozovete funkciju `factorialCalculator` sa `num` kao argumentom i dodelite rezultat varijabli `factorial`.
1. Trebalo bi da sačuvate konačni izlaz u formatu `Factorial of [num] is [factorial]` i dodelite ga varijabli `resultMsg`.
1. Trebalo bi da ispišete vrednost `resultMsg` na konzolu.

# --hints--

You should have a `num` variable.

```js
assert.isDefined(num);
```

You should assign a value to the `num` variable

```js
assert.isNotNull(num);
```

The value of `num` should be between `1` and `20`.

```js
assert.isAtLeast(num, 1);
assert.isAtMost(num, 20);
```

The `num` value should be a number.

```js
assert.isNumber(num);
```

You should have a function named `factorialCalculator`.

```js
assert.isFunction(factorialCalculator);
```

The `factorialCalculator` function should take a number as an argument.

```js
assert.match(factorialCalculator.toString(), /\s*factorialCalculator\(\s*\w+\s*\)/);
```

The factorial of `5` should return `120`.

```js
assert.strictEqual(factorialCalculator(5), 120);
```

The factorial of `7` should return `5040`.

```js
assert.strictEqual(factorialCalculator(7), 5040);
```

You should call your `factorialCalculator` function with the `num` variable as the argument.

```js
assert.match(__helpers.removeJSComments(code), /factorialCalculator\(\s*num\s*\)\s*;?\s*$/m);
```

You should define a `factorial` variable and assign the result of the `factorialCalculator` function to it.

```js
assert.isDefined(factorial);
assert.strictEqual(factorial, factorialCalculator(num));
```

Your `factorialCalculator` should produce the correct result.

```js
assert.strictEqual(factorialCalculator(6), 720);
assert.strictEqual(factorialCalculator(9), 362880);
assert.strictEqual(factorialCalculator(11), 39916800);
```

Your `resultMsg` should contain a string in a format `Factorial of [num] is [factorial]`.

```js
assert.strictEqual(resultMsg, `Factorial of ${num} is ${factorialCalculator(num)}`);
assert.notMatch(__helpers.removeJSComments(code), /Factorial of \d{1,2} is \d+/);
```

You should output the value of `resultMsg` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(resultMsg\)\s*;?\s?$/m);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const num = 5;

function factorialCalculator(n) {
  let result = 1;

  for (let i = 2; i <= n; i++) {
    result *= i;
  }

  return result;
}

const factorial = factorialCalculator(num);
const resultMsg = `Factorial of ${num} is ${factorial}`;
console.log(resultMsg);
```
