---
id: 66a9b8f14b963916a3baa732
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Da biste dodali natpis za `table`, možete koristiti element <dfn>table caption</dfn>.

Evo primera korišćenjem elementa `caption`:

```html
<table>
  <caption>Football Scores</caption>
</table>
```

Unutar vašeg `table` elementa, ugnježđujte element `caption` sa tekstom `Calculus Final Exam Grades`.

# --hints--

You should have an opening `caption` tag.

```js
assert.match(code, /<caption>/i);
```

You should have a closing `caption` tag.

```js
assert.match(code, /<\/caption>/i);
```

Your `caption` element should be inside your `table` element.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.parentElement.tagName, 'TABLE');
```

Your `caption` should have the text `Calculus Final Exam Grades`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.innerText.trim(), 'Calculus Final Exam Grades');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
    --fcc-editable-region--
      
    --fcc-editable-region--      
    </table>
  </body>
</html>
```
