---
id: 66a9bcf00f13a418368a272e
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Element zaglavlja tabele sastoji se od elementa <dfn>red u tabeli</dfn>, `tr`, koji sadrži elemente <dfn>zaglavlje ćelije tabele</dfn>, `th`.

Evo primera korišćenjem elemenata `tr` i `th` za sportsku tabelu:

```html
<table>
  <caption>Football Scores</caption>
  <thead>
    <tr>
      <th>Team</th>
      <th>Wins</th>
      <th>Losses</th>
    </tr>
  </thead>
</table>
```

Unutar vašeg `thead` elementa, dodajte element `tr`.

Unutar vašeg `tr` elementa, dodajte tri `th` elementa. 

Prvi `th` element treba da sadrži tekst `Last Name`. Drugi `th` element treba da sadrži tekst `First Name`. Treći `th` element treba da sadrži tekst `Grade`.

# --hints--

You should have an opening `tr` tag.

```js
assert.match(code, /<tr>/i);
```

You should have a closing `tr` tag.

```js
assert.match(code, /<\/tr>/i);
```

Your `tr` element should be inside your `thead` element. 

```js
const trEl = document.querySelector('tr');
assert.strictEqual(trEl?.parentElement?.tagName, 'THEAD');
```

Your first `th` element should have the text of `Last Name`.

```js
const thEl = document.querySelectorAll('th')[0];
assert.strictEqual(thEl?.textContent.trim(), 'Last Name');
```

Your second `th` element should have the text of `First Name`.

```js
const thEl = document.querySelectorAll('th')[1];
assert.strictEqual(thEl?.textContent.trim(), 'First Name');
```

Your third `th` element should have the text of `Grade`.

```js
const thEl = document.querySelectorAll('th')[2];
assert.strictEqual(thEl?.textContent.trim(), 'Grade');
```

All of your `th` elements should be inside your `tr` element.

```js
assert.match(code, /<tr>.*<th>.*<\/th>.*<th>.*<\/th>.*<th>.*<\/th>.*<\/tr>/is);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
      --fcc-editable-region--
        
      --fcc-editable-region--
      </thead>
    </table>
  </body>
</html>
```
