---
id: 66a9b48a211a73155621d0df
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovoj radionici, veštaćete se sa radom sa HTML tabelama kreiranjem tabele završnog ispita za grupu studenata.

U prethodnim časovima, naučili ste kako da radite sa elementom `table` za predstavljanje tabele podataka.

Unutar vašeg elementa `body`, ugnježđujte element `table`.

# --hints--

You should have an opening `table` tag.

```js
assert.match(code, /<table>/i);
```

You should have a closing `table` tag.

```js
assert.match(code, /<\/table>/i);
```

Your `table` element should be inside your `body` element.

```js
const tableEl = document.querySelector('table');
assert.strictEqual(tableEl?.parentElement?.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
