---
id: 66a9c14e3b34c719e34bae20
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Da biste dodali podatke o učeniku u tabelu, moraćete da koristite elemente za red tabele i podatke tabele.

Element `<dfn>table data</dfn>`, ``td``, koristi se za kreiranje ćelije u tabeli.

Evo primera korišćenja elementa ``td`` za tabelu sportista:

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

Unutar vašeg elementa tela tabele, dodajte element red tabele, ``tr``, sa tri elementa podataka tabele, ``td``. 

Prvi element podataka tabele treba da sadrži prezime `Davis`. 

Drugi element podataka tabele treba da sadrži ime `Alex`. 

Treći element podataka tabele treba da sadrži razred/ocenu `54`.

# --hints--

You should have a `tr` element inside your `tbody` element.

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

You should have a `td` element with the text `Davis` inside the `tr` element.

```js
assert.strictEqual(document.querySelector('tbody tr td')?.textContent.trim(), 'Davis');
```

You should have a `td` element with the text `Alex` inside the `tr` element.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1]?.textContent.trim(), 'Alex');
```

You should have a `td` element with the text `54` inside the `tr` element.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2]?.textContent.trim(), '54');
```

You should have all three of the `td` elements inside the `tr` element.

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
