---
id: 68f29dbf9125476b4808d6fb
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Sada je vreme da stilizujemo kontejner `settings-card`. Kreirajte selektor klase za `settings-card`.

Ovo će biti mesto gde će se nalaziti sav vaš formatiranje za kontejner. Postavite `max-width` na `250px` da biste definisali ukupnu veličinu vašeg kontejnera.

Sledeći korak, postavite vaš `padding` na `20px`, tako da vaš sadržaj ima razmak između sebe i ivice kontejnera.

Nakon toga, kreirajte zaobljen rub postavljanjem vašeg `border-radius` na `10px`.

Zatim postavite `box-shadow` sa vrednostima `0 2px 6px rgba(0,0,0,0.2)`. Ovo će stvoriti suptilan „podignut“ izgled koji će dati dubinu kontejneru.

# --hints--

You should have a class selector in your css for `settings-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".settings-card"));
```

You should add `max-width` property to the `.settings-card` class selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth);
```

You should have a `max-width` property with a value of `250px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth, "250px");
```

You should add a `padding` property to the `.settings-card` class selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding);
```

You should have a `padding` property with a value of `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding, "20px");
```

You should add a `border-radius` property to your `.settings-card` class selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius);
```

You should have a `border-radius` property with a value of `10px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius, "10px");
```

You should add a `box-shadow` property to your `.settings-card` class selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow);
```

You should have a `box-shadow` with the value of `0 2px 6px rgba(0,0,0,0.2)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow, "rgba(0, 0, 0, 0.2) 0px 2px 6px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--
```
