---
id: 69008059d43d6751385883c1
title: Ispravljanje grešaka operatera za inkrementaciju i dekrementaciju u bug-ovoj aplikaciji
challengeType: 26
dashedName: debug-increment-and-decrement-operator-errors-in-a-buggy-app
---

# --description--

Vaš prijatelj vam je tražio da mu pomognete sa debagovanjem njegovog koda. On vežba inkrement i dekrement operatore, ali dobija neočekivane rezultate u konzoli.

Vaš zadatak je pomoći mu da identifikuje greške i ažurira kod tako da izlaz konzole odgovara očekivanim rezultatima.

Možete se osvrnuti na prethodne lekcije o inkrement i dekrement operatorima, kao i na prefiksnu i postfiksnu notaciju.

**User Stories:**

1. Trebalo bi da imate varijablu pod nazivom `updatedCounter` koja ispravno proizvodi vrednost `11` korišćenjem inkrementatora. Ažurirajte izraz da koristi ispravnu prefiksnu ili postfiksnu notaciju umesto hardkodiranja vrednosti.
2. Trebalo bi da imate varijablu pod nazivom `finalScore` koja ispravno proizvodi vrednost `8` korišćenjem inkrementatora. Ažurirajte izraz da koristi ispravnu prefiksnu ili postfiksnu notaciju umesto hardkodiranja vrednosti.
3. Trebalo bi da imate varijablu pod nazivom `updatedCoins` koja ispravno proizvodi vrednost `2` korišćenjem dekrementatora. Ažurirajte izraz da koristi ispravnu prefiksnu ili postfiksnu notaciju umesto hardkodiranja vrednosti.
4. Trebalo bi da imate varijablu pod nazivom `newHealth` koja ispravno proizvodi vrednost `7` korišćenjem dekrementatora. Ažurirajte izraz da koristi ispravnu prefiksnu ili postfiksnu notaciju umesto hardkodiranja vrednosti.

# --hints--

You should have a variable named `updatedCounter` that uses prefix increment notation (`++counter`).

```js
assert.match(code, /updatedCounter\s*=\s*\+\+counter\s*;/);
assert.strictEqual(typeof updatedCounter, 'number');
```

`updatedCounter` should equal to `11`.

```js
assert.strictEqual(updatedCounter, 11);
```

You should have a variable named `finalScore` that uses postfix increment notation (`score++`).

```js
assert.match(code, /finalScore\s*=\s*score\+\+\s*;/);
assert.strictEqual(typeof finalScore, 'number');
```

`finalScore` should equal to `8`.

```js
assert.strictEqual(finalScore, 8);
```

You should have a variable named `updatedCoins` that uses prefix decrement notation (`--coins`).

```js
assert.match(code, /updatedCoins\s*=\s*--coins\s*;/);
assert.strictEqual(typeof updatedCoins, 'number');
```

`updatedCoins` should equal to `2`.

```js
assert.strictEqual(updatedCoins, 2);
```

You should have a variable named `newHealth` that uses postfix decrement notation (`health--`).

```js
assert.match(code, /newHealth\s*=\s*health--\s*;/);
assert.strictEqual(typeof newHealth, 'number');
```

`newHealth` should equal to `7`.

```js
assert.strictEqual(newHealth, 7);
```

# --seed--

## --seed-contents--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = counter++;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = ++score;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = coins--;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = --health;
console.log(`Health is still ${newHealth}`);
```

# --solutions--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = ++counter;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = score++;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = --coins;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = health--;
console.log(`Health is still ${newHealth}`);
```
