id: 587d78ae367417b2b2512afe
title: Koristite skraćeni flex svojstvo
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cbpW2tE'
forumTopicId: 301112
dashedName: use-the-flex-shorthand-property
---

# --description--

Postoji skraćeni način za postavljanje više flex svojstava zajedno. Svojstva `flex-grow`, `flex-shrink` i `flex-basis` mogu se postaviti zajedno korišćenjem svojstva `flex`.

Na primer, `flex: 1 0 10px;` će postaviti element da ima `flex-grow: 1;`, `flex-shrink: 0;` i `flex-basis: 10px;`.

Standardna postavka svojstava je `flex: 0 1 auto;`.

# --instructions--

Dodajte CSS svojstvo `flex` na `#box-1` i `#box-2`. Postavite vrednost za `#box-1` tako da mu `flex-grow` bude `2`, `flex-shrink` bude `2`, a `flex-basis` bude `150px`. Postavite vrednost za `#box-2` tako da mu `flex-grow` bude `1`, `flex-shrink` bude `1`, a `flex-basis` bude `150px`.

Ove vrednosti će uzrokovati da se `#box-1` proširi kako bi popunio dodatni prostor dvostruko u odnosu na `#box-2` kada je kontejner veći od 300px i smanjiće ga dvostruko u odnosu na `#box-2` kada je kontejner manji od 300px. 300px je ukupna širina `flex-basis` vrednosti oba boksa.

# --hints--

Element `#box-1` bi trebalo da ima svojstvo `flex` postavljeno na vrednost `2 2 150px`.

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];

assert.equal(flexGrow, '2');
assert.equal(flexShrink, '2');
assert.equal(flexBasis, '150px');
```

Element `#box-2` bi trebalo da ima svojstvo `flex` postavljeno na vrednost `1 1 150px`.

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];

assert.equal(flexGrow, '1');
assert.equal(flexShrink, '1');
assert.equal(flexBasis, '150px');
```

Vaš kod bi trebalo da koristi svojstvo `flex` za `#box-1` i `#box-2`.

```js
assert.lengthOf(code.match(/flex:\s*?\d\s+?\d\s+?150px;/g), 2);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    flex: 2 2 150px;
    height: 200px;
  }

  #box-2 {
    background-color: orangered;
    flex: 1 1 150px;
    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```