id: 587d78ad367417b2b2512afa
title: Koristite svojstvo flex-wrap za prelamanje redova ili kolona
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cQv9ZtG'
forumTopicId: 301114
dashedName: use-the-flex-wrap-property-to-wrap-a-row-or-column
---

# --description--

CSS flexbox ima mogućnost razdvajanja fleksibilnog kontejnera u više redova (ili kolona). Po podrazumevanju, fleksibilni kontejner će postaviti sve fleksibilne elemente zajedno. Na primer, ceo red će biti na jednoj liniji.

Međutim, korišćenje svojstva `flex-wrap` govori CSS-u da prebaci elemente. To znači da višak elementi prelaze u novi red ili kolonu. Tačka gde se prelamanje dešava zavisi od veličine elemenata i veličine kontejnera.

CSS takođe ima opcije za smer prelamanja:

<ul><li><code>nowrap</code>: ovo je podrazumevana postavka, i ne prelaže elemente.</li><li><code>wrap</code>: prelaže elemente na više linija odozgo nadole ako je red, i o lijevo desno ako je kolona.</li><li><code>wrap-reverse</code>: prelaže elemente na više linija od dole ka gore ako je red, i od desna ka lijevo ako je kolona.</li></ul>

# --instructions--

Trenutni raspored ima previše kutija u jednom redu. Dodajte CSS svojstvo `flex-wrap` elementu `#box-container` i postavite mu vrednost `wrap`.

# --hints--

Element `#box-container` treba da ima svojstvo `flex-wrap` sa vrednošću `wrap`.

```js
const boxContainer = document.querySelector('#box-container');
const flexWrap = window.getComputedStyle(boxContainer)['flex-wrap'];
assert.strictEqual(flexWrap, 'wrap');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;
    flex-wrap: wrap;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```