id: 587d78af367417b2b2512b00
title: Koristi svojstvo align-self
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvzfv'
forumTopicId: 301107
dashedName: use-the-align-self-property
---

# --description--

Svojstvo za fleksibilne elemente je `align-self`. Ovo svojstvo vam omogućava da prilagodite poravnanje svakog pojedinačnog elementa, umesto da postavljate sve zajedno. Ovo je važno jer druge uobičajene metode za postavljanje koje koriste CSS svojstva `float`, `clear` i `vertical-align` ne rade sa fleksibilnim elementima.

`align-self` prihvata vrednosti iste kao i `align-items` i nadjačava bilo koju vrednost postavljenu svojstvom `align-items`.

# --instructions--

Dodaj CSS svojstvo `align-self` za `#box-1` i `#box-2`. Postavi vrednost `center` za `#box-1` i vrednost `flex-end` za `#box-2`.

# --hints--

Element `#box-1` treba da ima svojstvo `align-self` sa vrednošću `center`.

```js
const boxOne = document.querySelector('#box-1');
const alignment = window.getComputedStyle(boxOne)['align-self'];
assert.strictEqual(alignment, 'center');
```

Element `#box-2` treba da ima svojstvo `align-self` sa vrednošću `flex-end`.

```js
const boxTwo = document.querySelector('#box-2');
const alignment = window.getComputedStyle(boxTwo)['align-self'];
assert.strictEqual(alignment, 'flex-end');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    align-self: center;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    align-self: flex-end;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```