id: 587d78ac367417b2b2512af4
title: Koristite svojstvo flex-direction za kreiranje kolone
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cZmWeA4'
forumTopicId: 301109
dashedName: use-the-flex-direction-property-to-make-a-column

# --description--

Poslednja dva izazova koristila su svojstvo `flex-direction` postavljeno na `row`. Ovo svojstvo takođe može kreirati kolonu uređujući potomke fleksibilnog kontejnera vertikalno.

# --instructions--

Dodajte CSS svojstvo `flex-direction` elementu `#box-container`, i postavite mu vrednost `column`.

# --hints--

Element `#box-container` bi trebalo da ima svojstvo `flex-direction` postavljeno na `column`.

```js
const boxContainer = document.querySelector('#box-container');
const flexDirection = window.getComputedStyle(boxContainer)['flex-direction'];
assert.strictEqual(flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
    flex-direction: column;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```