id: 587d78ab367417b2b2512af2
title: Koristite svojstvo flex-direction za kreiranje redova
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cBEkbfJ'
forumTopicId: 301110
dashedName: use-the-flex-direction-property-to-make-a-row

# --description--

Dodavanje `display: flex` elementu ga pretvara u fleksibilni kontejner. To omogućava raspoređivanje bilo koje dete tog elementa u red ili kolonu. To se radi dodavanjem svojstva `flex-direction` na roditeljski element i postavljanjem vrednosti na `row` ili `column`. Kreiranje reda poravnava decu horizontalno, a kreiranje kolone vertikalno.

Ostale opcije za `flex-direction` su `row-reverse` i `column-reverse`.

**Napomena:** Podrazumevana vrednost svojstva `flex-direction` je `row`.

# --instructions--

Dodajte CSS svojstvo `flex-direction` elementu `#box-container` i postavite mu vrednost `row-reverse`.

# --hints--

Element `#box-container` treba da ima svojstvo `flex-direction` postavljeno na `row-reverse`.

```js
const boxContainer = document.querySelector('#box-container');
const flexDirection = window.getComputedStyle(boxContainer)['flex-direction'];
assert.strictEqual(flexDirection, 'row-reverse');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
    flex-direction: row-reverse;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```