id: 587d78ae367417b2b2512aff
title: Koristite svojstvo 'order' za ponovno raspoređivanje elemenata
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvNAG'
forumTopicId: 301116
dashedName: use-the-order-property-to-rearrange-items
---

# --description--

Svojstvo `order` se koristi za određivanje CSS redosleda u kojem su fleksibilni elementi prikazani unutar flex kontejnera. Podrazumevano, elementi će biti prikazani u istom redosu kakav je bio u HTML izvoru. Ovo svojstvo prima broj kao vrednost, a mogu se koristiti i negativni brojevi.

# --instructions--

Dodajte CSS svojstvo `order` elementima `#box-1` i `#box-2`. Dodelite `#box-1` vrednost `2`, a `#box-2` vrednost `1`.

# --hints--

Fleksibilni element `#box-1` bi trebalo da ima svojstvo `order` sa vrednošću `2`.

```js
const boxOne = document.querySelector('#box-1');
const order = window.getComputedStyle(boxOne)['order'];
assert.strictEqual(order, '2');
```

Fleksibilni element `#box-2` bi trebalo da ima svojstvo `order` sa vrednošću `1`.

```js
const boxTwo = document.querySelector('#box-2');
const order = window.getComputedStyle(boxTwo)['order'];
assert.strictEqual(order, '1');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    order: 2;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    order: 1;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```