id: 587d78ad367417b2b2512afb
title: Koristite svojstvo flex-shrink za smanjenje veličine elemenata
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cd3PBfr'
forumTopicId: 301113
dashedName: use-the-flex-shrink-property-to-shrink-items

# --description--

Do sada, sva svojstva u izazovima se tiču fleksibilnog kontejnera (roditelja fleksibilnih elemenata). Međutim, postoje i nekoliko ključnih svojstava za fleksibilne elemente.

Prvo je svojstvo `flex-shrink`. Kada se koristi, omogućava elementu da smanji veličinu ako je fleksibilni kontejner previše mali. Elementi smanjuju svoju veličinu kada je širina roditeljskog kontejnera manja od ukupne širine svih fleksibilnih elemenata unutar njega.

Svojstvo `flex-shrink` uzima broj kao vrednost. Što je veći broj, to će element više smanjiti u odnosu na druge elemente unutar kontejnera. Na primer, ako jedan element ima vrednost `flex-shrink` od `1`, a drugi ima vrednost `flex-shrink` od `3`, onaj sa vrednošću `3` će se smanjiti tri puta više nego drugi.

# --instructions--

Dodajte CSS svojstvo `flex-shrink` za `#box-1` i `#box-2`. Dajte `#box-1` vrednost `1`, a `#box-2` vrednost `2`.

# --hints--

Element `#box-1` bi trebalo da ima svojstvo `flex-shrink` sa vrednošću `1`.

```js
const boxOne = document.querySelector('#box-1');
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
assert.equal(flexShrink, '1');
```

Element `#box-2` bi trebalo da ima svojstvo `flex-shrink` sa vrednošću `2`.

```js
const boxTwo = document.querySelector('#box-2');
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
assert.equal(flexShrink, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;
    flex-shrink: 1;
  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;
    flex-shrink: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```