id: 587d78ae367417b2b2512afc
title: Koristite svojstvo flex-grow za proširivanje elemenata
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c2p78cg'
forumTopicId: 301111
dashedName: use-the-flex-grow-property-to-expand-items
---

# --description--

Suprotnost sa svojstvom `flex-shrink` je svojstvo `flex-grow`. Zapamtite da `flex-shrink` kontroliše veličinu elemenata kada kontejner skuplja. Svojstvo `flex-grow` kontroliše veličinu elemenata kada roditeljski kontejner raste.

Koristeći primer sa prethodnog izazova, ako jedan element ima vrednost `flex-grow` od `1`, a drugi ima vrednost `flex-grow` od `3`, onaj sa vrednošću `3` će biti tri puta veći od drugog.

# --instructions--

Dodajte CSS svojstvo `flex-grow` za `#box-1` i `#box-2`. Dajte `#box-1` vrednost `1`, a `#box-2` vrednost `2`.

# --hints--

Element `#box-1` treba da ima svojstvo `flex-grow` sa vrednošću `1`.

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
assert.equal(flexGrow, '1');
```

Element `#box-2` treba da ima svojstvo `flex-grow` sa vrednošću `2`.

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
assert.equal(flexGrow, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-grow: 1;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-grow: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```