id: 587d78ab367417b2b2512af0
title: 'Koristite display: flex za pozicioniranje dve kutije'
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cgz3QS7'
forumTopicId: 301105
dashedName: use-display-flex-to-position-two-boxes
---

# --description--

Ovaj deo koristi alternativni format izazova da pokaže kako koristiti CSS za pozicioniranje elemenata na način koji se prilagođava uređaju. Prvo, izazov će definisati teoriju, a zatim praktičan izazov koristeći jednostavan deo tvita koristiće koncept flexbox-a.

Postavljanje CSS svojstva `display: flex;` na element omogućava vam da koristite druga svojstva fleksibilnog rasporeda za izgradnju prilagodljivog dizajna stranice.

# --instructions--

Dodajte CSS svojstvo `display` na `#box-container` i postavite mu vrednost na `flex`.

# --hints--

`#box-container` treba da ima svojstvo `display` sa vrednošću `flex`.

```js
const boxContainer = document.querySelector('#box-container');
const displayStyle = window.getComputedStyle(boxContainer)['display'];
assert.strictEqual(displayStyle, 'flex');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    height: 500px;

  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    height: 500px;
    display: flex;
  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```