id: 587d78ae367417b2b2512afd
title: Koristite svojstvo flex-basis za postavljanje početne veličine elementa
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c3d9nCa'
forumTopicId: 301108
dashedName: use-the-flex-basis-property-to-set-the-initial-size-of-an-item
---

# --description--

Svojstvo `flex-basis` određuje početnu veličinu elementa pre nego što CSS izvrši prilagođavanje za `flex-shrink` ili `flex-grow`.

Dimenzije koje koristi svojstvo `flex-basis` iste su kao i druge jedinice veličine (`px`, `em`, `%`, itd.). Vrednost `auto` postavlja veličinu elemenata u zavisnosti od sadržaja.

# --instructions--

Postavite početnu veličinu kutija koristeći `flex-basis`. Dodajte CSS svojstvo `flex-basis` za `#box-1` i `#box-2`. Dajte `#box-1` vrednost `10em`, a `#box-2` vrednost `20em`.

# --hints--

Element `#box-1` bi trebalo da ima svojstvo `flex-basis`.

```js
const boxOne = document.querySelector('#box-1');
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

Element `#box-1` bi trebalo da ima vrednost `flex-basis` od `10em`.

```js
assert.match(code, /#box-1\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?10em;/g);
```

Element `#box-2` bi trebalo da ima svojstvo `flex-basis`.

```js
const boxTwo = document.querySelector('#box-2');
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

Element `#box-2` bi trebalo da ima vrednost `flex-basis` od `20em`.

```js
assert.match(code, /#box-2\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?20em;/g);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-basis: 10em;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-basis: 20em;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```