---
id: 6853e58d52ec64349a5760ca
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovom radionici, implementiraćete algoritam najkraćeg puta. Napišaćete Python funkciju koja računa najkraći put između čvorova u grafu i takođe vraća pređeni put.

Na primer, ako je dat graf gde su gradovi povezani putevima različitih udaljenosti, algoritam će pronaći najkraću rutu iz jednog grada u drugi. Ako želite da putujete iz Grada A do Grada D, algoritam bi mogao da utvrdi da je putovanje A ⇨ B ⇨ C ⇨ D (ukupno: 15km) kraće od direktnog puta A ⇨ D (20km).

Da biste počeli, definišite varijablu pod nazivom `INF` i dodelite joj vrednost `float('inf')`, koja predstavlja beskonačnost. Kasnije ćete je koristiti za označavanje beskonačne udaljenosti između dva čvora.

# --hints--

You should define the variable `INF` at the top of your file.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("INF")`));
  }
});
```

You should assign `float('inf')` to the variable `INF`.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("INF").is_equivalent("INF = float('inf')")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
