---
id: 686253f1b051998ad4904e3e
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Sada ćete kreirati glavnu funkciju koja prima tri parametra: matrica susednosti, početni čvor i opcionalni ciljni čvor.

Kreirajte funkciju pod nazivom `shortest_path` koja prima tri parametra: `matrix`, `start_node` i `target_node`.

Dodelite `None` kao podrazumevanu vrednost za `target_node`.

Parametar `target_node` je podrazumevano podešen na `None`, što ukazuje da ako nije specificiran ciljni čvor, funkcija bi trebalo da izračuna najkraće staze od početnog čvora do svih ostalih čvorova u grafu.

Dodajte `pass` izjavu unutar tela funkcije za sada.

# --hints--

You should create a function named `shortest_path`.

```js
({
    test: () => {
      assert(runPython(`_Node(_code).has_function("shortest_path")`));
    }
});
```

The function should take in three parameters: `matrix`, `start_node`, and `target_node` with `target_node` having a default value of `None`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_args("matrix, start_node, target_node=None")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--

--fcc-editable-region--
```
