---
id: 68760464eb9c2e79cc913fac
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

U ovom koraku, dodati ćete petlju koja će se izvršiti jednom za svaki čvor u grafu. Ova petlja će omogućiti algoritmu ažuriranje udaljenosti i putanja tokom više prolaza.

Kreirajte petlju ``for`` koja se izvršava ``n`` puta. Koristite ``_`` kao varijablu petlje jer ne morate koristiti vrednost iteracije.  

Unutar petlje, potrebno je pripremiti se za odabir sledećeg čvora za obradu kreiranjem dve varijable: 

- jedna koja će držati najmanju udaljenost pronađenu do sada u trenutnoj iteraciji  
- i druga za skladištenje indeksa čvora koji ima ovu najmanju udaljenost.

Kreirajte varijable ``min_distance`` i ``current``, i postavite ih na ``INF`` i ``-1``, respektivno.

# --hints--

You should have a `for` loop within the `shortest_path` function.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0]`
        )
      );
    }
});
```

Your `for` loop should iterate over `range(n)`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_iter().is_equivalent("range(n)")`
        )
      );
    }
});
```

Your for loop should use `_` as the iteration variable.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_vars().is_equivalent("_")`
        )
      );
    }
});
```

You should create a variable named `min_distance` and set it to `INF` inside the for loop.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("min_distance = INF")`
        )
      );
    }
});
```

You should create a variable named `current` and set it to `-1` inside the for loop.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("current = -1")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
--fcc-editable-region--

--fcc-editable-region--
```
