---
id: 6876089f7a500b7c782e5b83
title: Korak 12
challengeType: 20
dashedName: step-12
---

# --description--

Morate odlučiti da li je trenutni čvor bolji izbor od onog koji ste već pronašli (ako postoji). Da biste to uradili, dodaćete uslovnu izjavu unutar petlje.

Uslov bi trebalo da uradi dve stvari:

- Proveriti da li čvor još nije poseta.
- Uporediti poznatu udaljenost sa trenutnim `min_distance`.

Unutar unutrašnje petlje `for`, dodajte izjavu `if` koja proverava da li `node_no` nije poseta **i** da li je `distances[node_no]` manje od `min_distance`.

Dodajte `pass` kao zamenu (placeholder) unutar uslovnog bloka.

# --hints--

You should add an `if` statement inside your inner `for` loop.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0]`
        )
      );
    }
});
```

Your `if` statement should check if the node is unvisited (`visited[node_no]` is falsy) and `distances[node_no]` is less than `min_distance`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("not visited[node_no] and distances[node_no] < min_distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

--fcc-editable-region--
    for _ in range(n):  
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            pass

--fcc-editable-region--
```
