---
id: 68e52994bc7ea2e3dec530a1
title: Korak 23
challengeType: 20
dashedName: step-23
---

# --description--

Sada vam je potreban način za upravljanje više strategija popusta i pronalaženje najbolje cene.

Kreirajte klasu pod nazivom `DiscountEngine`. Ova klasa će upravljati svim strategijama popusta i izračunati najbolju cenu za proizvod.

Dajte joj metod `__init__` koji prima `self` i `strategies` kao parametre, gde `strategies` ima tipsku indikaciju (type hint) od `list[DiscountStrategy]`. Korišćenjem `list[DiscountStrategy]` specificira se da je parametar lista koja sadrži objekte `DiscountStrategy`. Tip povratne vrednosti metoda treba da bude `None`. Unutar metode `__init__`, sačuvajte `strategies` kao atribut instance.

# --hints--

You should create a class named `DiscountEngine`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountEngine")`)
})
```

Your `DiscountEngine` class should have an `__init__` method.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").has_function("__init__")`)
})
```

Your `__init__` method should have parameters `self` and `strategies: list[DiscountStrategy]`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_args("self, strategies: list[DiscountStrategy]")`)
})
```

Your `__init__` method should have a return type of `None`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_returns("None")`)
})
```

You should assign `strategies` to `self.strategies` inside your `__init__` method.

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_stmt("self.strategies = strategies")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
