---
id: 68e529226b0defdff26c4a8b
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Sada ćete početi da gradite interfejs sistema popusta. Python pruža način za definisanje apstraktnih klas koje služe kao nacrti za druge klase. To se radi korišćenjem klase `ABC` (Abstract Base Class) iz modula `abc`.

Na vrh vašeg koda dodajte sledeće deklarisanje:

```py
from abc import ABC, abstractmethod
```

Ovo uvozi klasu `ABC` i dekorator `abstractmethod`, koje ćete koristiti za definisanje metoda koje mora implementirati svaka klasa koja nasleđuje vašu apstraktnu klasu.

# --hints--

You should have the mentioned import statement at the top of your code.

```js
({
  test: () => runPython(`assert _Node(_code).find_imports()[0].is_equivalent("from abc import ABC, abstractmethod")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

product = Product('Wireless Mouse', 50.0)
print(product)
```
