---
id: 68e529226b0defdff26c4a89
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Trenutno, ako kreirate instancu `Product` i odštapišete je, videćete podrazumevanu reprezentaciju objekta koja nije baš informativna. Na primer, može izgledati nešto kao `<__main__.Product object at 0x7f9b8c2d1d30>`.

Da biste konvertovali podrazumevanu reprezentaciju objekta u čitljiv niz znakova, kreirajte metod `__str__` u svojoj klasi `Product` koji vraća string sa formatom `name - $price` (gde se `name` i `price` treba zameniti vrednostima odgovarajućih atributa).

# --hints--

Your `Product` class should have a `__str__` method.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__str__")`)
})
```

Your `__str__` method should return a string in the format `name - $price`.

```js
({
  test: () => runPython(`
    test_product = Product("Test Item", 25.5)
    assert str(test_product) == "Test Item - $25.5"
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float):
        self.name = name
        self.price = price
--fcc-editable-region--
    
--fcc-editable-region--
```
