---
id: 68e524c94e2e4ab7b4bfc7e2
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovom radionici, kreći ćete sa izgradnjom kalkulatora popusta koji može primeniti različite strategije popusta na proizvode. Sistem će odrediti najbolju cenu za klijenta na osnovu više pravila popusta.

Počnite definisanjem klase pod nazivom `Product`. Dajte svojoj klasi metod `__init__` sa sledećim parametrima: `self`, `name` i `price`. Unutar metode `__init__`, dodelite parametre `name` i `price` instancijskim atributima istog imena.

# --hints--

You should create a class named `Product`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Product")`)
})
```

Your `Product` class should have an `__init__` method.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__init__")`)
})
```

Your `__init__` method should have the following parameters: `self`, `name`, and `price`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name, price")`)
})
```

You should assign `name` to `self.name` inside your `__init__` method.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.name = name")`)
})
```

You should assign `price` to `self.price` inside your `__init__` method.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.price = price")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
