---
id: 68e5293bd00d2fe134f58982
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Kao što je diskutovano u prethodnoj lekciji, apstraktne metode su metode koje mora implementirati svaka klasa koja nasleđuje od apstraktne bazne klase. One služe kao kontrakt, osiguravajući da sve strategije popusta imaju određene neophodne metode.

Unutar vaše `DiscountStrategy` klase, definišite metodu nazvanu `is_applicable`. Koristite dekorator `@abstractmethod` iznad definicije metode. Metoda treba da ima sledeće parametre: `self`, `product` i `user_tier`. Koristite statement `pass` u telu metode.

Dodajte tipove podataka (type hints) kako biste učinili potpis metode jasnim: `product` bi trebalo da bude tipa `Product`, a `user_tier` bi trebalo da bude tipa `str`.

# --hints--

Your `DiscountStrategy` class should have a method named `is_applicable`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("is_applicable")`)
})
```

Your `is_applicable` method should have the `@abstractmethod` decorator.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_decorators("abstractmethod")`)
})
```

Your `is_applicable` method should have parameters `self`, `product: Product`, and `user_tier: str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

--fcc-editable-region--
class DiscountStrategy(ABC):
    pass
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
