---
id: 68e5293bd00d2fe134f58983
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Sada dodajte drugi apstraktni metod pod nazivom `apply_discount` u vašu klasu `DiscountStrategy`. Koristite dekorator `@abstractmethod` i date mu parametre `self` i `product`.

Dodajte tipove: `product` bi trebalo da bude tipa `Product`, a metoda treba da vrati `float`. Koristite `pass` u telu metode.

# --hints--

Your `DiscountStrategy` class should have a method named `apply_discount`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("apply_discount")`)
})
```

Your `apply_discount` method should have the `@abstractmethod` decorator.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_decorators("abstractmethod")`)
})
```

Your `apply_discount` method should have parameters `self` and `product: Product`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_args("self, product:Product")`)
})
```

Your `apply_discount` method should have a return type hint of `-> float`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_returns("float")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
