---
id: 68e9207ddcba053c5cb24413
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

U Pythonu, tipovi naznaka (type hints) su opcioni signali koji drugim programerima govore kakav je očekivan tip podataka varijable ili funkcije. Evo primera koji dodaje tip naznaka `str` parametru `name`.

```py
def greet(name: str):
    return 'Hello, ' + name
```

Za parametre metode `__init__`, ažurirajte parametar `name` da ima tip naznaka (type hint) `str`, a parametar `price` da ima tip naznaka (type hint) `float`.

# --hints--

You should add the type hint `str` to the `name` parameter and the type hint `float` to the `price` parameter in the `__init__` method of the `Product` class.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name: str, price: float")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Product:
    def __init__(self, name, price):
        self.name = name
        self.price = price
--fcc-editable-region--
```
