---
id: 68f119ad905b29ff59d858c8
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Kao i u prethodnom koraku, u vašoj metodi `__str__`, dodajte tip hint za povratnu vrednost od `str` jer ova metoda vraća string reprezentaciju objekta.

# --hints--

Your `__str__` method should have a return type hint of `str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__str__").has_returns("str")`)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price
--fcc-editable-region--
    def __str__(self):
        return f'{self.name} - ${self.price}'
--fcc-editable-region--
```
