---
id: 66b102ee0bdbad34a9f85ba0
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, nastavićete da učite o stringovima gradeći Chatbota za nastavnika.

Da biste počeli, dodajte izjavu `console`, sa porukom `"Hi there!"`.

# --hints--

You should have a `console` statement.

```js
assert.match(code, /console\.log(.*)/);
```

Your console statement should output the message `"Hi there!"`.

```js
assert.match(code, /console\.log\((['"])(Hi\s+there!)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
