---
id: 66b59b12a745e10011158f7b
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Sledeći korak je kreiranje još nekoliko promenljivih koje će se koristiti u budućim porukama bota.

Kreirajte promenljivu nazvanu `subject` i dodelite joj vrednost niza (string) `"JavaScript"`.

Zatim kreirajte promenljivu nazvanu `topic` i dodelite joj vrednost niza (string) `"strings"`.

# --hints--

You should have a variable called `subject`.

```js
assert.isNotNull(subject);
```

You should assign a string to your `subject` variable.

```js
assert.isString(subject);
```

You should assign the string `"JavaScript"` to your `subject` variable.

```js
assert.equal(subject, "JavaScript");
```

You should have a variable called `topic`.

```js
assert.isNotNull(topic);
```

You should assign a string to your `topic` variable.

```js
assert.isString(topic);
```

You should assign the string `"strings"` to your `topic` variable.

```js
assert.equal(topic, "strings");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

--fcc-editable-region--

--fcc-editable-region--
```
