---
id: 66b59829dba144ff1351220f
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Sada je vreme da kreirate pozdrav koristeći varijablu `botName`.

U prethodnim lekcijama, naučili ste kako da konkatenišete nizove koristeći *template literal*-e na ovaj način:

```js
const name = "John";

// "Hello, John!"
`Hello, ${name}!`;
```

Počnite kreiranjem varijable nazvane `greeting`.

Zatim, koristeći sintaksu *template literal*-a, dodelite niz koji kaže `My name is`, praćen varijablom `botName`, i završava tačkom (`.`).

Na kraju, logujte varijablu `greeting` u konzolu.

# --hints--

You should have a variable called `greeting`.

```js
assert.isNotNull(greeting);
```

You should assign a string to your `greeting` variable.

```js
assert.isString(greeting);
```

You should use template literals to concatenate the string `My name is` with the `botName` variable followed by a period (`.`).

```js
assert.equal(greeting, "My name is teacherBot.");
```

You should have a `console` statement that logs the `greeting` variable to the console.

```js
assert.match(code, /console\.log\(.*greeting.*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
