---
id: 66b6efddeca35833cd6f0b03
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

U prethodnim lekcijama, naučili ste kako da pristupate karakterima u nizu na ovaj način:

```js
const firstName = "Jessica";
// returns "J"
firstName[0];
```  

Zapamtite da brojevi indeksa počinju na `0`, pa će prvi slovo u nizu uvek biti na indeksu `0`.

Počnite dodavanjem još jedne `console` naredbe. 

Unutar `console` naredbe, ispišite prvo slovo varijable `subject` koristeći notaciju uglastih zagrada i tačan broj indeksa.

# --hints--

Your `console` statement should output the first letter of the `subject` variable using bracket notation and index `0`.

```js
assert.match(code, /console\.log\(subject\[0\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
