---
id: 66b6e39031393e30f2c48d0f
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Da biste dobili dužinu niza, možete koristiti svojstvo <dfn>length</dfn>. Ovo svojstvo vraća broj znakova u nizu.

Evo primera:

```js
const greeting = "Hello, world!";

// Output: 13
console.log(greeting.length);
```

Počnite korišćenjem svojstva `length` da biste dobili dužinu niza `subject` i logovali tu vrednost u konzolu.

# --hints--


You should use the `length` property to get the length of the `subject` string.

```js
assert.match(code, /subject\.length/);
```

You should log the length of the `subject` string to the console.

```js
assert.match(code, /console\.log\(\s*subject\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

--fcc-editable-region--

--fcc-editable-region--
```
