---
id: 68da6138d70960c197ec5ef1
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Sada je potrebno da identifikujete kada ste izgradili kompletnu kombinaciju zagradi. Kompletnu kombinaciju ima dužina jednaka dvostrukom broja parova (jer svaki par doprinosi jednu otvarajuću i jednu zatvarajuću zagradu).

Unutar vaše `while` petlje, dodajte `if` naredbu da proverite da li je `len(current) == 2 * pairs`. Kada je ovaj uslov tačan, dodajte (append) `current` u listu `result`.

# --hints--

You should have an `if` statement inside your `while` loop.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0]
`) })
```

Your `if` statement should check if `len(current) == 2 * pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("len(current) == 2 * pairs")
`) })
```

You should append `current` to the `result` list inside your new `if` statement.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("result.append(current)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
--fcc-editable-region--
        
--fcc-editable-region--

    return result

print(gen_parentheses(1))
```
