---
id: 68da6138d70960c197ec5ef3
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Sada morate da obradite dodavanje zatvarajućih zagradi. Ključno pravilo je da možete dodati zatvarajuću zagradu samo ako održava ravnotežu, što znači da mora biti više korišćenih otvarajućih zagradi nego zatvarajućih.

I dalje unutar bloka `else`, dodajte drugu `if` instancu da proverite da li je `closes_used < opens_used`.

Ako je ovaj uslov tačan, dodajte još jedan novi tupel u `queue`: `(current + ')', opens_used, closes_used + 1)`. Ovo predstavlja stanje nakon dodavanja zatvarajuće zagrade.

# --hints--

You should have a second `if` statement inside your `else` block.

```js
({ test: () => runPython(`
import ast
assert isinstance(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1], ast.If)
`) })
```

Your second nested `if` statement should check if `closes_used < opens_used`.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1].test).is_equivalent("closes_used < opens_used")
`) })
```

You should append `(current + ')', opens_used, closes_used + 1)` to the queue when the condition is true.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1]).has_call("queue.append((current + ')', opens_used, closes_used + 1))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
        else:
            if opens_used < pairs:
                queue.append((current + '(', opens_used + 1, closes_used))
--fcc-editable-region--
            
--fcc-editable-region--

    return result

print(gen_parentheses(1))
```
