---
id: 68da604b8b6012a33e2344fb
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovom radionici, implementiraćete funkciju koja generiše sve validne kombinacije zagradi koristeći pristup širinu prvenog pretrage (BFS). Na primer, validne kombinacije za par zagradi su `(())` i `()()`.

Počnite kreiranjem funkcije pod nazivom `gen_parentheses` sa jednim parametrom `pairs`. Za sada, vratite prazan spisak iz funkcije.

# --hints--

You should define a function named `gen_parentheses`.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("gen_parentheses")
`) })
```

Your `gen_parentheses` function should have a single parameter named `pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_args("pairs")
`) })
```

Your function should return an empty list.

```js
({ test: () => runPython(`
foo = _Node(_code).find_function("gen_parentheses")
assert foo.has_return("[]") or foo.has_return("list()")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
