---
id: 68da6138d70960c197ec5eec
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Pre nego što implementirate glavni algoritam, morate validirati ulaz. Parametar `pairs` bi trebalo da bude ceo broj, jer predstavlja broj parova zagradi koje treba generisati.

Dodajte izjavu `if` na početak vaše funkcije da proverite da li `pairs` nije ceo broj. Koristite funkciju `isinstance()` za to.

Ako je uslov tačan, vratite string `The number of pairs should be an integer`.

# --hints--

You should have an `if` statement in your `gen_parentheses` function.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0]
`) })
```

Your `if` statement should check if `pairs` is not an integer using `not isinstance(pairs, int)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(pairs, int)")
`) })
```

You should return the string `The number of pairs should be an integer` when the input is not an integer.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be an integer'
args = [1.1, "1", [1, 2], {}]
assert  all(gen_parentheses(arg) == msg for arg in args)
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
--fcc-editable-region--
    
--fcc-editable-region--
    
    return []
```
