---
id: 68da6138d70960c197ec5eef
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Za pristup pretrazi u širinu (breadth-first search), koristićeš red (queue) za praćenje različitih stanja dok gradiš kombinacije zagradi. Svako stanje će biti predstavljeno kao tupel koji sadrži tri elementa:

- Trenutni niz (string) koji se gradi
- Broj otvarajućih zagradi korišćen do sada
- Broj zatvarajućih zagradi korišćen do sada

Kreiraj varijablu pod nazivom `queue` i inicijalizuj je listom koja sadrži jedan tupel: `('', 0, 0)`. Ovo predstavlja početno stanje sa praznim nizom i nula zagradi korišćene.

# --hints--

You should declare a variable named `queue` in your `gen_parentheses` function.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("queue")
`) })
```

You should initialize `queue` with a list containing the tuple `('', 0, 0)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("queue").is_equivalent("queue = [('', 0, 0)]")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
--fcc-editable-region--
    
--fcc-editable-region--
    result = []
    return result
```
