---
id: 6998b99a8e195cd0c1f211df
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

Na ovom radionici, napravićete jednostavan tracker divljine pomoću objekata JavaScript-a.

Evo primera objekta:

```js
const animal = {}; // empty object
```

Ovo kreira prazan objekat.

Sada kreirajte varijablu pod imenom `tiger` i dodelite joj prazan objekat.

# --hints--

You should create a variable named `tiger`.

```js
assert.exists(tiger);
```

The `tiger` variable should be assigned an empty object.

```js
assert.deepEqual(tiger, {});
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
