---
id: 6999faf886634b72be7a819b
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Za ovaj tracker divljih životinja, ažuriraćete objekat sa novim svojstvima.

Evo primera kreiranja objekta sa svojstvom:

```js
const animal = {
  name: "Lion" // property: value
};
```

Sada ažurirajte objekat `tiger` tako da uključuje svojstvo nazvano `species`.

Postavite mu vrednost `"Tiger"`.

# --hints--

`tiger` should be an object.

```js
assert.isObject(tiger);
```

`tiger` should have a `species` property.

```js
assert.property(tiger, "species");
```

`tiger.species` should be equal to `"Tiger"`.

```js
assert.propertyVal(tiger, "species", "Tiger");
```

# --seed--

## --seed-contents--

```js
const tiger = {
--fcc-editable-region--
  
--fcc-editable-region--
};
```
