---
id: 69999ff7f7862fdfb885ea5d
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Sada ažurirajte vaš objekat ``tiger`` dodavanjem nove osobine nazvane ``age``.

Postavite svojstvo ``age`` na vrednost ``5``.

# --hints--

`tiger` should have an `age` property.

```js
assert.property(tiger, "age");
```

`tiger.age` should be equal to `5`.

```js
assert.propertyVal(tiger, "age", 5);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
--fcc-editable-region--
  
--fcc-editable-region--
};
```
