---
id: 6999ac171928136168654774
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

U ovom koraku, kreiraćete funkciju za pristup svojstvu `species` objekta.

Evo malog primera pristupa svojstvu iz objekta:

```js
const dog = {
  species: "Dog"
};
const getAnimalSpecies = (pet) => {
  return pet.species; // access species using dot notation
};
console.log(getAnimalSpecies(dog)); // Dog
```

U ovom primeru, funkcija prima objekat kao parametar i vraća njegovo svojstvo `species`.

Sada kreirajte funkciju nazvanu `getSpecies`.

Funkcija bi trebalo da primi jedan parametar pod imenom `animal`.

Vratite svojstvo `species` od `animal` koristeći dot notaciju.

Nakon kreiranja funkcije, koristite `console.log` da pozovete `getSpecies(tiger)` kako biste videli rezultat u konzoli.

# --hints--

You should create a function named `getSpecies`.

```js
assert.isFunction(getSpecies);
```

The `getSpecies` function should have a single parameter, `animal`.

```js
const regex = __helpers.functionRegex('getSpecies', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

You should log `getSpecies(tiger)` to the console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getSpecies\s*\(\s*tiger\s*\)\s*\)/);
```

`getSpecies` should access the species property using dot notation.

```js
assert.match(code, /animal\.species/);
```

Calling `getSpecies(tiger)` should return `"Tiger"`.

```js
assert.strictEqual(getSpecies(tiger), "Tiger");
```

`getSpecies` should use the function parameter and work with any object.

```js
const lion = { species: "Lion" };
assert.strictEqual(getSpecies(lion), "Lion");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
