---
id: 6999b0198d18d402c36cc5fe
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

U ovom koraku, kreiraćete funkciju koja uklanja svojstvo iz objekta.

Evo primera uklanjanja svojstva korišćenjem `delete`:

```js
const bird = {
  species: "Parrot",
  canFly: true
};

const removeFlight = (pet) => {
  delete pet.canFly; // remove property using delete keyword
  return pet; // return the updated object
};

console.log(removeFlight(bird));
// { species: "Parrot" }
```

U ovom primeru, svojstvo `canFly` se uklanja iz objekta `bird`.

Sada kreirajte funkciju nazvanu `removeEndangeredStatus`. Funkcija treba da primi jedan parametar pod imenom `animal`.

Unutar funkcije, uklonite svojstvo `isEndangered` iz objekta `animal` koristeći ključnu reč `delete`. Vratite ažurirani objekat `animal`.

Nakon kreiranja funkcije, koristite `console.log` da pozovete `removeEndangeredStatus(tiger)`, kako biste videli ažurirani objekat u konzoli.

# --hints--

You should create a function named `removeEndangeredStatus`.

```js
assert.isFunction(removeEndangeredStatus);
```

The `removeEndangeredStatus` function should have a single parameter, `animal`.

```js
const regex = __helpers.functionRegex('removeEndangeredStatus', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

The `removeEndangeredStatus` function should return the updated `animal` object.

```js
const testAnimal = { species: "Cat", isEndangered: true };
const result = removeEndangeredStatus(testAnimal);
assert.strictEqual(result, testAnimal);
```

`removeEndangeredStatus` should use the `delete` keyword to remove the `isEndangered` property.

```js
assert.match(code, /delete\s+animal\.isEndangered/);
```

You should log `removeEndangeredStatus(tiger)` to the console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*removeEndangeredStatus\s*\(\s*tiger\s*\)\s*\)/);
```

Calling `removeEndangeredStatus(tiger)` should remove the `isEndangered` property.

```js
const updatedTiger = removeEndangeredStatus(tiger);

assert.notProperty(updatedTiger, "isEndangered");
```

`removeEndangeredStatus` should use the function parameters and work with any object.

```js
const lion = {
  species: "Lion",
  age: 4,
  isEndangered: true
};

const updatedLion = removeEndangeredStatus(lion);

assert.notProperty(updatedLion, "isEndangered");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

--fcc-editable-region--

--fcc-editable-region--
```
