---
id: 6999b0d79ea938edcb032237
title: Korak 12
challengeType: 1
dashedName: step-12
---

# --description--

U ovom poslednjem koraku radionice za praćenje divljine, kreiraćete funkciju koja pristupa svojstvu objekta korišćenjem notacije sa uglastim zagradama.

Evo primera pristupa svojstvu korišćenjem notacije sa uglastim zagradama:

```js
const cat = {
  species: "Cat",
  age: 3
};

const property = "species";

console.log(cat[property]); // "Cat"
```

U ovom primeru, notacija sa uglastim zagradama omogućava vam pristup svojstvu korišćenjem varijable.

Sada kreirajte funkciju nazvanu `getProperty`.

Funkcija bi trebalo da primi dva parametra: `animal` i `propertyName`.

Vratite vrednost svojstva korišćenjem notacije sa uglastim zagradama.

Nakon kreiranja funkcije, koristite `console.log` da pozovete `getProperty(tiger, "species")` i `getProperty(elephant, "age")` kako biste videli vraćene vrednosti u konzoli.

Time je radionica za praćenje divljine završena!

# --hints--

You should create a function named `getProperty`.

```js
assert.isFunction(getProperty);
```

The `getProperty` function should have two parameters: `animal` and `propertyName`.

```js
const regex = __helpers.functionRegex('getProperty', ['animal', 'propertyName']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getProperty` should use bracket notation to access the property.

```js
assert.match(code, /animal\s*\[\s*propertyName\s*\]/);
```

You should log `getProperty(tiger, "species")` to the console.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*tiger\s*,\s*["']species["']\s*\)\s*\)/
);
```

You should log `getProperty(elephant, "age")` to the console.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*elephant\s*,\s*["']age["']\s*\)\s*\)/
);
```

Calling `getProperty(tiger, "species")` should return `"Tiger"`.

```js
assert.strictEqual(getProperty(tiger, "species"), "Tiger");
```

Calling `getProperty(elephant, "age")` should return `12`.

```js
assert.strictEqual(getProperty(elephant, "age"), 12);
```

`getProperty` should use the function parameters and work with any object.

```js
const lion = {
  species: "Lion",
  weight: 200
};

assert.strictEqual(getProperty(lion, "weight"), 200);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

const getProperty = (animal, propertyName) => {
  return animal[propertyName];
};

console.log(getProperty(tiger, "species"));
console.log(getProperty(elephant, "age"));
```
