---
id: 6999a0c1698b93d5d41156ed
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Sada dodajte još jednu svojstvo objektu `tiger` nazvanu `isEndangered`.

Postavite svojstvo `isEndangered` na `true`.

# --hints--

`tiger` should have an `isEndangered` property.

```js
assert.property(tiger, "isEndangered");
```

`tiger.isEndangered` should be equal to `true`.

```js
assert.propertyVal(tiger, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
--fcc-editable-region--
  
--fcc-editable-region--
};
```
