---
id: 699a00de9f564bb0effaa14d
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--




U ovom koraku, kreiraćete funkciju za pristup svojstvu ``age`` objekta.

Evo primera:

```js
const cat = {
  age: 3
};

const getAnimalAge = (pet) => {
  return pet.age; // access age using dot notation
}

console.log(getAnimalAge(cat)); // 3
```

U ovom primeru, funkcija prima objekat kao parametar i vraća njegovo svojstvo ``age``.

Sada kreirajte funkciju nazvanu ``getAge``.

Funkcija treba da primi jedan parametar pod nazivom ``animal``.

Vratite svojstvo ``age`` objekta ``animal`` koristeći notaciju tačke.

Nakon kreiranja funkcije, koristite ``console.log`` za pozivanje ``getAge(tiger)`` kako biste videli rezultat u konzoli.

# --hints--

You should create a function named `getAge`.

```js
assert.isFunction(getAge);
```

The `getAge` function should have a single parameter, `animal`.

```js
const regex = __helpers.functionRegex('getAge', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getAge` should access the `age` property using dot notation.

```js
assert.match(code, /animal\.age/);
```

You should log `getAge(tiger)` to the console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getAge\s*\(\s*tiger\s*\)\s*\)/);
```

Calling `getAge(tiger)` should return `5`.

```js
assert.strictEqual(getAge(tiger), 5);
```

`getAge` should use the function parameter and work with any object.

```js
const lion = { age: 8 };
assert.strictEqual(getAge(lion), 8);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
