---
id: 669fdc11c9b0263fe0814a7a
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

U svakom od elemenata `li`, biće vam potreban anchor element.

Za prvi anchor element, tekst bi trebalo da bude `About`, a vrednost atributa `href` bi trebalo da bude `"#about"`. Hash simbol ispred `about` predstavlja ime `id`, koje će biti dodato kasnije u projektu.

Za drugi anchor element, tekst bi trebalo da bude `Posts`, a vrednost atributa `href` bi trebalo da bude `"#posts"`.

Za treći anchor element, tekst bi trebalo da bude `Contact`, a vrednost atributa `href` bi trebalo da bude `"#contact"`.

# --hints--

Your first anchor element should have the text of `About` inside your first `li` element.

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "About");
```

Your first anchor element should have an `href` attribute set to `"#about"`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

Your second anchor element should have the text of `Posts` inside your second `li` element.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement?.innerText.trim(), "Posts");
```

Your second anchor element should have an `href` attribute set to `"#posts"`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

Your third anchor element should have the text of `Contact` inside your third `li` element.

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "Contact");
```

Your third anchor element should have an `href` attribute set to `"#contact"`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
--fcc-editable-region--
          <li></li>
          <li></li>
          <li></li>
--fcc-editable-region--
        </ul>
      </nav>
    </header>
  </body>
</html>
```
