---
id: 669fcb5106eeb13ab7afcd09
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Na ovom radionici ćete vežbati rad sa semantičnim HTML-om kreiranjem blog stranice posvećene gospodinu Whiskersu mačku.

Prvi deo koji ćete napraviti je stranica <dfn>header</dfn>.

Element `header` se koristi za predstavljanje uvodnog sadržaja, kao što su navigacija stranice i druge uvodne informacije.

Evo primera korišćenjem elementa `header`:

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

Unutar elementa `body`, dodajte element `header`.

# --hints--

You should have an opening `header` tag.

```js
assert.match(code, /<header>/i);
```

You should have a closing `header` tag.

```js
assert.match(code, /<\/header>/i);
```

Your `header` element should be inside the `body` element.

```js
assert.exists(document.querySelector('body > header'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </body>
</html>
```
