---
id: 669fdfc9a5e3da42d2376609
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Sada kada ste dodali sekciju o nama, pokušajte da kliknete na link `About` kako biste videli da se stranica skroluje do te sekcije.

Sledeća sekcija na blog stranici biće lista postova o gospodinu Whiskersu.

Dodajte još jedan `section` element sa `id` podešenim na `"posts"`.

Unutar `section` elementa, dodajte `h2` element sa tekstom od `Posts`.

# --hints--

Your code should have an opening `section` tag.

```js
assert.match(code, /<section.*>/i);
```

Your code should have a closing `section` tag.

```js
assert.match(code, /<\/section>/i);
```

Your `section` element should have an `id` set to `"posts"`.

```js
assert.match(code, /<section.*id="posts".*>/i);
```

Your `section` element should be within your `main` element.

```js
const sectionElement = document.getElementById('posts');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

Your `h2` element should be inside the `section` element.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

Your `h2` element should have the text of `Posts`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.innerText.trim(), 'Posts');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </main>
  </body>
</html>
```
