---
id: 67d83df6f82eda3868dd2a84
title: Kreiraj RPG lika
challengeType: 27
dashedName: build-an-rpg-character
---

# --description--

U ovom laboratorijumu ćete vežbati osnove Pythona praveći malu aplikaciju koja kreira lik za RPG avanturu.

**Cilj:** Ispunite priče korisnika ispod i učinite da svi testovi prođu kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da imate funkciju pod nazivom `create_character`.
1. Funkcija treba da prihvati, po redu, ime lika praćeno sa tri statistike: snaga, inteligencija i karizma.
1. Ime lika treba da se validira:
   - Ako ime lika nije string, funkcija bi trebalo da vrati `The character name should be a string`.
   - Ako je ime lika prazan string, funkcija bi trebalo da vrati `The character should have a name`.
   - Ako je ime lika duže od 10 karaktera, funkcija bi trebalo da vrati `The character name is too long`.
   - Ako ime lika sadrži razmake, funkcija bi trebalo da vrati `The character name should not contain spaces`.
1. Statistike takođe treba da se validiraju:
   - Ako jedna ili više statistika nisu celobrojne, funkcija bi trebalo da vrati `All stats should be integers`.
   - Ako je jedna ili više statistika manje od 1, funkcija bi trebalo da vrati `All stats should be no less than 1`.
   - Ako je jedna ili više statistika veće od 4, funkcija bi trebalo da vrati `All stats should be no more than 4`.
   - Ako je zbir svih statistika drugačiji od 7, funkcija bi trebalo da vrati `The character should start with 7 points`.
1. Ako sve vrednosti prođu verifikaciju, funkcija bi trebalo da vrati string sa četiri reda:
   - prvi red treba da sadrži ime lika
   - redovi 2-4 treba da počinju sa skraćenicom statistike, `STR`, `INT` ili `CHA` (u ovom redu), zatim razmak, i zatim broj punih tačaka (`●`) jednak vrednosti statistike, i broj praznih tačaka (`○`) kako bi se dostiglo 10. Primer: ako je vrednost snage 3, mora postojati 3 pune tačke praćene sa 7 praznih tačaka. Tačke su date u editoru.

Evo stringa koji bi trebalo da vrati `create_character('ren', 4, 2, 1)`:

   ```md
   ren
   STR ●●●●○○○○○○
   INT ●●○○○○○○○○
   CHA ●○○○○○○○○○
   ```

NAPOMENA: iako su `str` i `int` uobičajene skraćenice za statistike, zapamtite da su to rezervisane ključne reči u Pythonu i ne bi trebalo da se koriste kao imena promenljivih.

# --hints--

You should have a function named `create_character`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_function('create_character')
`
    )
})
```

When `create_character` is called with a first argument that is **not** a string it should return `The character name should be a string`.

```js
({
    test: () => runPython(
        `
wrong_types = [3, 5, {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}]
for wrong_type in wrong_types:
    assert create_character(wrong_type, 1, 2, 4) == 'The character name should be a string'
`
    )
})
```

When `create_character` is called with a first argument that is a string it **should not** return `The character name should be a string`.

```js
({
    test: () => runPython(
        `assert create_character('ren', 1, 2, 4) != 'The character name should be a string'`
    )
})
```

When `create_character` is called with a first argument that is an empty string, it should return `The character should have a name`.

```js
({
    test: () => runPython(
        `
assert create_character('', 4, 2, 1) == 'The character should have a name'
`
    )
})
```

When `create_character` is called with a first argument that is **not** an empty string, it **should not** return `The character should have a name`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should have a name'
`
    )
})
```

When `create_character` is called with a first argument that is longer than 10 characters it should return `The character name is too long`.

```js
({
    test: () => runPython(
        `
for length in range(11, 27):
    assert create_character('abcdefghijklmnopqrstuvwxyz'[:length], 4, 2, 1) == 'The character name is too long'
`
    )
})
```

The `create_character` function **should not** say that the character is too long when it's not longer than 10 characters.

```js
({
    test: () => runPython(
        `
for length in range(0, 11):
  assert create_character('aaabbcccdd'[:length], 4, 2, 1) != 'The character name is too long'
`
    )
})
```

When `create_character` is called with a first argument that contains a space it should return `The character name should not contain spaces`.

```js
({
    test: () => runPython(
        `
assert create_character('cha cha', 4, 1, 2) == 'The character name should not contain spaces'
`
    )
})
```

When `create_character` is called with a first argument that *does not* contain a space it *should not* return `The character name should not contain spaces`.

```js
({
    test: () => runPython(
        `
assert create_character('chacha', 4, 2, 1) != 'The character name should not contain spaces'
`
    )
})
```

When `create_character` is called with a second, third or fourth argument that is **not** an integer it should return `All stats should be integers`.

```js
({
    test: () => runPython(
        `
wrong_types = ['3', '5', {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}, 'str', 'friend']
for wrong_type in wrong_types:
    assert create_character('friend', wrong_type, 2, 1) == 'All stats should be integers'
    assert create_character('friend', 2, wrong_type, 1) == 'All stats should be integers'
    assert create_character('friend', 2, 1, wrong_type) == 'All stats should be integers'
`
    )
})
```

When `create_character` is called with a second, third and fourth argument that are all integers it **should not** return `All stats should be integers`.

```js
({
    test: () => runPython(
        `
assert create_character('friend', 4, 2, 1) != 'All stats should be integers'
`
    )
})
```

When `create_character` is called with a second, third or fourth argument that is lower than `1` it should return `All stats should be no less than 1`.

```js
({
    test: () => runPython(
        `
expected = 'All stats should be no less than 1'

assert create_character('ren', 0, 4, 3) == expected
assert create_character('ren', 4, 0, 3) == expected
assert create_character('ren', 4, 3, 0) == expected
assert create_character('ren', -1, 4, 4) == expected
assert create_character('ren', 4, -1, 4) == expected
assert create_character('ren', 4, 4, -1) == expected
`
    )
})
```

When `create_character` is called with a second, third and fourth argument that are all no less than `1` it **should not** return `All stats should be no less than 1`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no less than 1'
`
    )
})
```

When `create_character` is called with a second, third or fourth argument that is higher than `4` it should return `All stats should be no more than 4`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 1, 5, 1) == 'All stats should be no more than 4'
assert create_character('ren', 5, 1, 1) == 'All stats should be no more than 4'
assert create_character('ren', 1, 1, 5) == 'All stats should be no more than 4'
`
    )
})
```

When `create_character` is called with a second, third and fourth argument that are all no more than `4` it **should not** return `All stats should be no more than 4`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no more than 4'
`
    )
})
```

When `create_character` is called with a second, third or fourth argument that **do not** sum to `7` it should return `The character should start with 7 points`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 4, 4) == 'The character should start with 7 points'
assert create_character('ren', 1, 1, 1) == 'The character should start with 7 points'
assert create_character('ren', 1, 2, 3) == 'The character should start with 7 points'
`
    )
})
```

When `create_character` is called with a second, third and fourth argument that sum to `7` it **should not** return `The character should start with 7 points`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should start with 7 points'
`
    )
})
```

`create_character('ren', 4, 2, 1)` should return `ren\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) == f'ren\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})
```

When `create_character` is called with valid values it should output the character stats as required.

```js
({
    test: () => runPython(
        `
assert create_character('Bob', 1, 2, 4) == f'Bob\\nSTR ●○○○○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●●●●○○○○○○'
assert create_character('aaabbbcccd', 4, 2, 1) == f'aaabbbcccd\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})

```

# --seed--

## --seed-contents--

```py
full_dot = '●'
empty_dot = '○'


```

# --solutions--

```py
full_dot = '●'
empty_dot = '○'

def create_character(character_name, strength, intelligence, charisma):
    if type(character_name) != str:
        return 'The character name should be a string'
    if len(character_name) == 0:
        return 'The character should have a name'
    if len(character_name) > 10:
        return 'The character name is too long'
    if ' ' in character_name:
        return 'The character name should not contain spaces'

    if type(strength) != int or type(intelligence) != int or type(charisma) != int:
       return 'All stats should be integers'
    if strength < 1 or intelligence < 1 or charisma < 1:
        return 'All stats should be no less than 1'
    if strength > 4 or intelligence > 4 or charisma > 4:
        return 'All stats should be no more than 4'
    if strength + intelligence + charisma != 7:
        return 'The character should start with 7 points'

    character_sheet = f'{character_name}\n'
    character_sheet += f'STR {full_dot * strength}{empty_dot * (10 - strength)}\n'
    character_sheet += f'INT {full_dot * intelligence}{empty_dot * (10 - intelligence)}\n'
    character_sheet += f'CHA {full_dot * charisma}{empty_dot * (10 - charisma)}'
    return character_sheet
```
