---
id: 69a92f0b9cc04bb0d5327bb5
title: Implementujte Stak (Stack)
challengeType: 26
dashedName: implement-a-stack
---

# --description--

U ovom laboratorijskom zadatku, implementiraćete strukturu podataka stak (stack) koristeći funkcije. Stak je Last-In-First-Out (LIFO) struktura podataka gde se elementi dodaju i uklanjaju sa vrha.

**Cilj:** Ispunite priče korisnika ispod i učinite da svi testovi prođu kako biste završili laboratorijski zadatak.

**Priče korisnika:**

1. Trebalo bi da imate funkciju `initStack` koja vraća objekat sa svojstvom `collection` postavljeno na prazan niz.
1. Trebalo bi da imate funkciju `push` koja dodaje element na vrh steka.
1. Trebalo bi da imate funkciju `pop` koja uklanja i vraća gornji element steka, ili `undefined` ako ga nema.
1. Trebalo bi da imate funkciju `peek` koja vraća gornji element steka bez uklanjanja, ili `undefined` ako ga nema.
1. Trebalo bi da imate funkciju `isEmpty` koja vraća `true` ako stak ne sadrži elemente, i `false` u suprotnom slučaju.
1. Trebalo bi da imate funkciju `clear` koja uklanja sve elemente iz steka.

**Napomena:** Većina testova zavisi od funkcija `initStack` i `push`. Implementirajte ih prvo, jer testovi za `pop`, `peek`, `isEmpty` i `clear` zahtevaju dodavanje elemenata u stek.

# --hints--

You should have a `push` function.

```js
assert.isFunction(push);
```

`push` function should add an element to the top of the stack.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

You should have a `pop` function.

```js
assert.isFunction(pop);
```

`pop` function should remove and return the top element of the stack.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
assert.strictEqual(pop(stack), 'third');
assert.strictEqual(pop(stack), 'second');
```

`pop` function should return `undefined` if the stack is empty.

```js
const stack = initStack();
assert.isUndefined(pop(stack));
```

`pop` function should return falsy values correctly.

```js
const stack = initStack();
push(stack, '');
assert.strictEqual(pop(stack), '');
assert.lengthOf(stack.collection, 0);
```

You should have a `peek` function.

```js
assert.isFunction(peek);
```

`peek` function should return the top element of the stack without removing it.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.strictEqual(peek(stack), 'second');
assert.strictEqual(peek(stack), 'second');
```

`peek` function should return `undefined` if the stack is empty.

```js
const stack = initStack();
assert.isUndefined(peek(stack));
```

`peek` function should return falsy values correctly.

```js
const stack = initStack();
push(stack, 0);
assert.strictEqual(peek(stack), 0);
```

You should have an `isEmpty` function.

```js
assert.isFunction(isEmpty);
```

`isEmpty` function should return `true` for an empty stack.

```js
const stack = initStack();
assert.isTrue(isEmpty(stack));
```

`isEmpty` function should return `false` for a non-empty stack.

```js
const stack = initStack();
push(stack, 'element');
assert.isFalse(isEmpty(stack));
```

`isEmpty` function should return `false` when the top element is falsy. 

```js
const stack = initStack();
push(stack, false);
assert.isFalse(isEmpty(stack));
```

You should have a `clear` function.

```js
assert.isFunction(clear);
```

`clear` function should remove all elements from the stack.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
clear(stack);
assert.lengthOf(stack.collection, 0);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function initStack() {
  return {
    collection: []
  };
}

function print(stack) {
  console.log(stack.collection);
}

function push(stack, element) {
  stack.collection.push(element);
}

function pop(stack) {
  return stack.collection.pop();
}

function peek(stack) {
  return stack.collection[stack.collection.length - 1];
}

function isEmpty(stack) {
  return stack.collection.length === 0;
}

function clear(stack) {
  stack.collection = [];
}
```
