id: 5e44412c903586ffb414c94c
title: Izgradi projekat za formatiranje aritmetike
challengeType: 23
saveSubmissionToDB: true
forumTopicId: 462359
dashedName: build-an-arithmetic-formatter-project
---

# --description--

Učenici osnovne škole često uređuju aritmetičke probleme vertikalno kako bi bilo lakše da ih reše. Na primer, "235 + 52" postaje:

```py
  235
+  52
-----
```

Napravi funkciju `arithmetic_arranger` koja prima niz stringova koji predstavljaju aritmetičke probleme, i vraća te probleme formatirane vertikalno i u kolone. Funkcija treba da prihvati opcionalni drugi argument. Kada je drugi argument postavljen na `True`, rešenja treba prikazati.

## Primer

Poziv funkcije:

```py
arithmetic_arranger(["32 + 698", "3801 - 2", "45 + 43", "123 + 49"])
```

Rezultat:

```py
   32      3801      45      123
+ 698    -    2    + 43    +  49
-----    ------    ----    -----
```

Poziv funkcije:

```py
arithmetic_arranger(["32 + 8", "1 - 3801", "9999 + 9999", "523 - 49"], True)
```

Rezultat:

```py
  32         1      9999      523
+  8    - 3801    + 9999    -  49
----    ------    ------    -----
  40     -3800     19998      474
```

## Pravila

Funkcija treba da vrati tačno formatirane probleme, inače će vratiti **string** koji opisuje smislenu grešku za korisnika.

- Uslovi koji vraćaju grešku:
  - Ako je broj problema prevelik. Granica je **pet problema**, više toga će vratiti:
    `'Error: Too many problems.'`
  - Dopušteni operatori za funkciju su **plus** i **minus**. Množenje i deljenje vraćaće grešku. Ostali operateri nisu navedeni ovde i ne treba testirati ih. Greška koja se vraća je:
    `"Error: Operator must be '+' or '-'."`
  - Svaki broj (operand) mora sadržati samo cifre. Inače, funkcija će vratiti:
    `'Error: Numbers must only contain digits.'`
  - Svaki operand (tj. broj na bilo kojoj strani operatera) sme da ima maksimalno četiri cifre. Inače, vraćaće se niz grešaka:
    `'Error: Numbers cannot be more than four digits.'`
- Ako korisnik unese probleme sa ispravnom strukturom, formatirani prikaz treba da uzme u obzir ova pravila:
  - Mora postojati jedan razmak između operatera i najdužeg od dva operanda; operater će biti na istom nivou kao drugi operand, svi operandi treba da budu poravnati po desnoj strani (prvi će biti gore, a drugi dole).
  - Brojevi treba da budu poravnati po desnoj strani.
  - Mora postojati četiri razmaka između svakog problema.
  - Ispod svakog problema mora postojati red crtica. Ove linije moraju se protezati na punu dužinu celog problema, svaka za sebe. (Primer gore pokazuje kako bi trebalo da izgleda.)

Napomena: Otvorite konzolu pretraživača sa F12 da vidite detaljne rezultate testiranja.

# --hints--

`arithmetic_arranger(["3801 - 2", "123 + 49"])` treba da vrati `  3801      123\n-    2    +  49\n------    -----`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3801 - 2", "123 + 49"]), '  3801      123\\n-    2    +  49\\n------    -----')`);
  }
})
```

`arithmetic_arranger(["1 + 2", "1 - 9380"])` treba da vrati `  1         1\n+ 2    - 9380\n---    ------`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["1 + 2", "1 - 9380"]), '  1         1\\n+ 2    - 9380\\n---    ------')`);
  }
})
```

`arithmetic_arranger(["3 + 855", "3801 - 2", "45 + 43", "123 + 49"])` treba da vrati `    3      3801      45      123\n+ 855    -    2    + 43    +  49\n-----    ------    ----    -----`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 + 855", "3801 - 2", "45 + 43", "123 + 49"]), '    3      3801      45      123\\n+ 855    -    2    + 43    +  49\\n-----    ------    ----    -----')`);
  }
})
```

`arithmetic_arranger(["11 + 4", "3801 - 2999", "1 + 2", "123 + 49", "1 - 9380"])` treba da vrati `  11      3801      1      123         1\n+  4    - 2999    + 2    +  49    - 9380\n----    ------    ---    -----    ------`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["11 + 4", "3801 - 2999", "1 + 2", "123 + 49", "1 - 9380"]), '  11      3801      1      123         1\\n+  4    - 2999    + 2    +  49    - 9380\\n----    ------    ---    -----    ------')`);
  }
})
```

`arithmetic_arranger(["44 + 815", "909 - 2", "45 + 43", "123 + 49", "888 + 40", "653 + 87"])` treba da vrati `'Error: Too many problems.'`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["44 + 815", "909 - 2", "45 + 43", "123 + 49", "888 + 40", "653 + 87"]), 'Error: Too many problems.')`);
  }
})
```

`arithmetic_arranger(["3 / 855", "3801 - 2", "45 + 43", "123 + 49"])` treba da vrati `"Error: Operator must be '+' or '-'."`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 / 855", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Operator must be '+' or '-'.")`);
  }
})
```

`arithmetic_arranger(["24 + 85215", "3801 - 2", "45 + 43", "123 + 49"])` treba da vrati `'Error: Numbers cannot be more than four digits.'`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["24 + 85215", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Numbers cannot be more than four digits.")`);
  }
})
```

`arithmetic_arranger(["98 + 3g5", "3801 - 2", "45 + 43", "123 + 49"])` treba da vrati `'Error: Numbers must only contain digits.'`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["98 + 3g5", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Numbers must only contain digits.")`);
  }
})
```

`arithmetic_arranger(["3 + 855", "988 + 40"], True)` treba da vrati `    3      988\n+ 855    +  40\n-----    -----\n  858     1028`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 + 855", "988 + 40"], True), "    3      988\\n+ 855    +  40\\n-----    -----\\n  858     1028")`);
  }
})
```

`arithmetic_arranger(["32 - 698", "1 - 3801", "45 + 43", "123 + 49", "988 + 40"], True)` treba da vrati `   32         1      45      123      988\n- 698    - 3801    + 43    +  49    +  40\n-----    ------    ----    -----    -----\n -666     -3800      88      172     1028`.

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["32 - 698", "1 - 3801", "45 + 43", "123 + 49", "988 + 40"], True), "   32         1      45      123      988\\n- 698    - 3801    + 43    +  49    +  40\\n-----    ------    ----    -----    -----\\n -666     -3800      88      172     1028")`);
  }
})
```