---
id: 69860157f7fce39f9df399a3
title: Implementirati algoritam za N kraljica
challengeType: 26
dashedName: implement-the-n-queens-algorithm
---

# --description--

Problem N Kraljica traži da postavite N kraljica na NxN šahovnicu tako da se nijedna od dve kraljice ne napada međusobno (da dve ne dele red, kolonu ili dijagonalu).

Na primer, ako postoji šahovnica 4x4, jedna validna raspored je:

```md
[1, 3, 0, 2]
```

To znači da je u redu 0 kraljica postavljena u kolonu 1; u redu 1, kraljica je postavljena u kolonu 3; u redu 2, kraljica je postavljena u kolonu 0; a u redu 3, kraljica je postavljena u kolonu 2.

Vizuelno, ovaj raspored izgleda ovako:

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

Gde `Q` predstavlja kraljicu, a `.` praznu polju.

U ovom laboratorijumu, implementiraćete rešivač problema N Kraljica koristeći pristup pretrazi u dubinu (depth-first search).


**Cilj**: Ispunite priče korisnika ispod i prođite sve testove kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da imate funkciju pod nazivom `dfsNQueens`.
2. Funkcija treba da prihvati tačno jedan argument: ceo broj `n`.
3. Ako je `n` manje od `1`, funkcija bi trebalo da vrati prazan niz (`[]`).
4. Funkcija treba da vrati listu rešenja; svako rešenje je sam po sebi lista dužine `n`, gde element na indeksu `i` predstavlja indeks kolone (0-baziran) kraljice u redu `i`.

# --hints--

You should have a function named `dfsNQueens` that takes one argument.

```js
assert.isFunction(dfsNQueens);
assert.lengthOf(dfsNQueens, 1);
```

If `n` is less than `1`, the function should return an empty array.

```js
assert.sameDeepOrderedMembers(dfsNQueens(0), []);
assert.sameDeepOrderedMembers(dfsNQueens(-1), []);
assert.sameDeepOrderedMembers(dfsNQueens(-5), []);
```

The function should return an array of solutions, where each solution is an array of length `n`.

```js
const output = dfsNQueens(4);
assert.isArray(output);
output.forEach(solution => {
    assert.isArray(solution);
    assert.lengthOf(solution, 4);
})
```

`dfsNQueens(1)` should return `[[0]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(1), [[0]]);
```

`dfsNQueens(2)` should return `[]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(2), []);
```

`dfsNQueens(3)` should return `[]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(3), []);
```

`dfsNQueens(4)` should return `[[1, 3, 0, 2], [2, 0, 3, 1]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(4), [[1, 3, 0, 2], [2, 0, 3, 1]]);
```

`dfsNQueens(5)` should return `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(5), [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]);
```

`dfsNQueens(5).length` should be `10`.

```js
assert.lengthOf(dfsNQueens(5), 10);
```

`dfsNQueens(8).length` should be `92`.

```js
assert.lengthOf(dfsNQueens(8), 92);
```

`dfsNQueens` should return the correct result.

```js
assert.lengthOf(dfsNQueens(7), 40);
assert.sameDeepOrderedMembers(dfsNQueens(6), [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfsNQueens(n) {
  if (n < 1) return [];

  const results = [];

  function isSafe(queens, row, col) {
    for (let r = 0; r < queens.length; r++) {
      const c = queens[r];
      if (c === col || Math.abs(row - r) === Math.abs(col - c)) {
        return false;
      }
    }
    return true;
  }

  function dfs(queens) {
    const row = queens.length;
    if (row === n) {
      results.push(queens.slice());
      return;
    }

    for (let col = 0; col < n; col++) {
      if (isSafe(queens, row, col)) {
        queens.push(col);
        dfs(queens);
        queens.pop();
      }
    }
  }

  dfs([]);
  return results;
}
```
