---
id: ab6137d4e35944e21037b769
title: Napravite konverter za naslovni slučaj
challengeType: 26
dashedName: build-a-title-case-converter
---

# --description--

U ovom laboratorijumu ćete kreirati funkciju koja konvertuje niz karaktera u format naslova (title case). Format naslova znači da je prvo slovo svake reči veliko, a ostatak reči malim slovima.

`"Web Development Is Awesome"` je primer niza karaktera u formatu naslova.

**Cilj**: Ispunite priče korisnika ispod i osigurajte da svi testovi prođu kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da imate funkciju `titleCase` koja prima niz karaktera kao argument.
1. Funkcija `titleCase` bi trebalo da vrati niz karaktera sa prvim slovom svake reči kapitalizovanim, a ostatkom reči malim slovima.
1. `titleCase("I like to code")` bi trebalo da vrati `"I Like To Code"`.
1. `titleCase("javaScript is fun")` bi trebalo da vrati `"Javascript Is Fun"`.

# --hints--

You should create a function named `titleCase`.

```js
assert.isFunction(titleCase);
```

`titleCase` should take a single argument.

```js
assert.lengthOf(titleCase, 1);
```


`titleCase("I'm a little tea pot")` should return a string.

```js
assert.isString(titleCase("I'm a little tea pot"));
```

`titleCase("I'm a little tea pot")` should return the string `I'm A Little Tea Pot`.

```js
assert.strictEqual(titleCase("I'm a little tea pot"), "I'm A Little Tea Pot");
```

`titleCase("sHoRt AnD sToUt")` should return the string `Short And Stout`.

```js
assert.strictEqual(titleCase('sHoRt AnD sToUt'), 'Short And Stout');
```

`titleCase("HERE IS MY HANDLE HERE IS MY SPOUT")` should return the string `Here Is My Handle Here Is My Spout`.

```js
assert.strictEqual(
  titleCase('HERE IS MY HANDLE HERE IS MY SPOUT'),
  'Here Is My Handle Here Is My Spout'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function titleCase(str) {
  return str
    .split(' ')
    .map(word => word.charAt(0).toUpperCase() + word.substring(1).toLowerCase())
    .join(' ');
}
```
