---
id: 6823ac607bfdbc46331b2559
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML je skraćenica za HyperText Markup Language. To je kod koji definiše strukturu i sadržaj web stranice. Ovo je vaš uresnik koda, gde ćete pisati HTML.

Pronađite liniju 1 u uresniku i ukucajte ovaj tekst:

```md
Welcome to freeCodeCamp
```

Kada završite, kliknite na dugme "proverite svoj kod" da vidite da li je tačan.

# --hints--

You should have the text `Welcome to freeCodeCamp` in your editor. Double-check for spelling.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*freecodecamp/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
