---
id: 6823d6244511f252c8300eed
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Kada trebate dodati pasus na web stranicu, možete koristiti element `p` ovako:

```html
<p>This is a paragraph element.</p>
```

Pretvorite `Learn the skills to become a full-stack developer` u pasusni element.

# --hints--

Your `p` element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector("p"));
```

Your `p` element should have a closing tag. Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/p\s*\>/);
```

Your `p` element's text should be `Learn the skills to become a full-stack developer`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*Learn\s*the\s*skills\s*to\s*become\s*a\s*full-stack\s*developer\s*\<\/p\s*\>/i);
```

Your `p` element should be below your `h2` element. 

```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
--fcc-editable-region--
Learn the skills to become a full-stack developer
--fcc-editable-region--
```
