---
id: 6823e637c1c0ed56f781b4fc
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Na kraju, JavaScript čini vašu web stranicu interaktivnom — omogućava vam da kažete stranici šta da uradi kada neko klikne na dugme, pošalje formular ili mnogo drugih stvari.

Za poslednji korak radionice, dodajte još jedan element `h3` i `p` na stranicu koji opisuje JavaScript.

Prvo, dodajte element `h3` sa tekstom:

```md
Introduction to JavaScript
```

Zatim, ispod tog elementa `h3`, dodajte element `p` sa sledećim tekstom:

```md
JavaScript adds interactivity to a webpage
```

# --hints--

You should add a third `h3` element to the page.

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

Your `h3` element's text should be `Introduction to JavaScript`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*JavaScript\s*\<\/h3\s*\>/i);
```

You should add a fourth `p` element to the page.

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

Your `p` element should have the text `JavaScript adds interactivity to a webpage`

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*JavaScript\s*adds\s*interactivity\s*to\s*a\s*web\s*page\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

<h3>Introduction to JavaScript</h3>
<p>JavaScript adds interactivity to a webpage</p>
```
