---
id: 6823d9ac8bdc3853df65a1ff
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

U HTML-u postoji šest elemenata za naslove: `h1` do `h6`. Koriste se za prikaz važnosti sekcija na vašoj veb stranici, pri čemu je `h1` najvažniji, a `h6` najmanje.

Ispod vašeg `p` elementa, dodajte naslov `h3` sa tekstom:

```md
Introduction to HTML
```

# --hints--

Your `h3` element should have an opening `<h3>` tag.

```js
assert.exists(document.querySelector("h3"));
```

Your `h3` element should have a closing `</h3>` tag.

```js
assert.match(code, /<\/h3\s*\>/);
```

Your `h3` element's text should be `Introduction to HTML`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```
