---
id: 682ba2318000b62f179bdf04
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

HTML je sastavljen od elemenata. Prvi koji ćete koristiti je element `h1`:

```html
<h1>Welcome to freeCodeCamp</h1>
```

Počinje otvarajućim tagom (`<h1>`), završava zatvarajućim tagom (`</h1>`), i ima tekst koji će prikazati između tagova.

Pretvorite svoj tekst `Welcome to freeCodeCamp` u element `h1` dodavanjem otvarajućeg taga ispred njega i zatvarajućeg taga iza njega.

# --hints--

Your `h1` element should have an opening `<h1>` tag.

```js
assert.exists(document.querySelector("h1"));
```

Your `h1` element should have a closing `</h1>` tag.

```js
assert.match(code, /<\/h1\s*\>/);
```

Your `h1` element should look like this: `<h1>Welcome to freeCodeCamp</h1>`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*freeCodeCamp\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to freeCodeCamp
--fcc-editable-region--
```
