---
id: 682cd206883fc7b25eb539c5
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Primećujte da HTML koji pišete u editoru se prikazuje u pregledniku. U ovoj radionici, pisati ćete HTML za delimičnu stranicu kurikuluma.

Ispod vašeg elementa `h1`, unesite sledeće na praznu liniju:

```md
Full-Stack Curriculum
```

# --hints--

You should have the text `Full-Stack Curriculum` in your editor. Double-check for spelling.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full-stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
--fcc-editable-region--

--fcc-editable-region--
```
