---
id: 66a207974c806a19d6607073
title: Napravi program za biblioteku knjiga
challengeType: 25
dashedName: build-a-book-inventory-app
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**Cilj:** Popunite korisniÄŤke priÄŤe ispod i proÄ‘ite sve testove kako biste zavrĹˇili ovaj laboratorijski rad.

**PriÄŤe korisnika:**

1. Unapaswa kuwa na kipengele cha `h1` chenye maandishi ya `Book Inventory`.
1. Unapaswa kuwa na kipengele cha `table` chenye safu zilizo na majina ya `Title`, `Author`, `Category`, `Status`, na `Rate`.
1. Jedwali lako linapaswa kuwa na angalau mistari minne, mstari wa kwanza kwa ajili ya vichwa vya safu na mingine ikiwa na taarifa za vitabu.
1. Kila mstari wa jedwali ndani ya mwili wa jedwali unapaswa kuwa na darasa la vitu la `read`, `to-read`, au `in-progress`.
1. Vipengele vya `td` vya safu ya `Status` vinapaswa kuwa na kipengele cha `span` chenye sifa ya kipengele cha `class` cha `status` kinachozunguka maandishi ya `Read`, `To Read`, au `In Progress`, kulingana na darasa la mstari huo.
1. Vipengele vya `td` vya safu ya `Rate` vinapaswa kuwa na kipengele cha `span` chenye sifa ya kipengele cha `class` cha `rate` kinachozunguka vipengele vitatu vya `span` vilivyo tupu.
1. Vipengele vya `.rate` vilivyowekwa ndani ya mistari ya `.read` vinapaswa kuwa na darasa la ziada lenye thamani ya `one`, `two`, au `three`, kulingana na kiwango binafsi. Thamani hii inapaswa kuja baada ya `rate`.
1. Unapaswa kuunda wachaguaji wa sifa za kipengele tatu ili kulenga vipengele vyenye darasa la `read`, `to-read`, na `in-progress`, na kuweka sifa yao ya `background-image` kutumia `linear-gradient` unayochagua.
1. Unapaswa kuweka sifa ya `display` ya kila kipengele cha `span` kuwa `inline-block`.
1. Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` ambavyo ni mzao wa vipengele vya `tr` vyenye darasa la `to-read` na kuweka sifa zao za `border` na `background-image`.
1. Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` ambavyo ni mzao wa vipengele vya `tr` vyenye darasa la `read` na kuweka sifa zao za `border` na `background-image`.
1. Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` ambavyo ni mzao wa vipengele vya `tr` vyenye darasa la `in-progress` na kuweka sifa zao za `border` na `background-image`.
1. Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` vyenye darasa la `status` na vipengele vya `span` vyenye thamani ya darasa kuanzia na `rate` na kuweka sifa zao za `height`, `width`, na `padding`.
1. Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vya `span` ambavyo ni watoto wa moja kwa moja wa vipengele vya `span` vyenye thamani ya `class` kuanzia na `rate` na kuweka sifa zao za `border`, `border-radius`, `margin`, `height`, `width`, na `background-color`.
1. Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga mzazi wa kwanza wa vipengele vya `span` ambavyo vina neno `one` kama sehemu ya thamani yao ya `class` na kuweka sifa yake ya `background-image` kutumia `linear-gradient`.
1. Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga wazazi wawili wa kwanza wa vipengele vya `span` ambavyo vina neno `two` kama sehemu ya thamani yao ya `class` na kuweka sifa zao za `background-image` kutumia `linear-gradient`.
1. Unapaswa kutumia kichaguzi cha sifa ya kipengele kulenga vipengele vitatu vya `span` ambavyo ni mzao wa vipengele vya `span` ambavyo vina neno `three` kama sehemu ya thamani yao ya `class` na kuweka sifa zao za `background-image` kutumia `linear-gradient`.

# --hints--

Trebalo bi da imaĹˇ element sa tekstom `h1` koji sadrĹľi tekst `Book Inventory`.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'Book Inventory');
```

Trebalo bi da ima samo jedan element `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Trebalo bi da imaĹˇ element `table`.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

Trebalo bi da ima jedan element `thead` i jedan element `tbody` unutar `table`.

```js
assert.lengthOf(document.querySelectorAll("table > thead"), 1);
assert.lengthOf(document.querySelectorAll("table > tbody"), 1);
```

Unutar `thead` treba da ima jedan element od `tr` sa pet komponenti od `th`.

```js
assert.lengthOf(document.querySelectorAll("table > thead > tr"), 1);
assert.lengthOf(document.querySelectorAll("table > thead > tr > th"), 5);
```

Tvoja prva stranica bi trebalo da ima tekst `Title` kao naslov sekcije.

```js
assert.equal(document.querySelectorAll('th')[0]?.textContent.trim(), 'Title');
```

VaĹˇ drugi red mora da sadrĹľi tekst `Author` kao naslov sekcije.

```js
assert.equal(document.querySelectorAll('th')[1]?.textContent.trim(), 'Author');
```

VaĹˇ treÄ‡i red bi trebalo da sadrĹľi tekst `Category` kao zaglavlje sekcije.

```js
assert.equal(document.querySelectorAll('th')[2]?.textContent.trim(), 'Category');
```

VaĹˇ ÄŤetvrti red treba da ima tekst `Status` kao naslov sekcije.

```js
assert.equal(document.querySelectorAll('th')[3]?.textContent.trim(), 'Status');
```

VaĹˇ peti red bi trebalo da sadrĹľi tekst `Rate` kao naslov sekcije.

```js
assert.equal(document.querySelectorAll('th')[4]?.textContent.trim(), 'Rate');
```

Tabela bi trebalo da ima najmanje ÄŤetiri reda.

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
```

Svaki red mora imati 5 kolona uvek.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isNotEmpty(rows);
for (const row of rows) {
  assert.lengthOf(row.children, 5);
}
```

Svaki red tabele, osim reda zaglavlja sekcije, treba da ima klasu `read`, `to-read` ili `in-progress`.

```js
const rows = document.querySelectorAll('tr');
assert.isNotEmpty(rows);
for (let i = 1; i < rows.length; i++) {
  const classList = [...rows[i].classList];
  const currentProgress = classList[0]; 
  assert.oneOf(currentProgress, ['read', 'to-read', "in-progress"]); 
}
```

Elementi od ``td`` u redu ``Status`` treba da imaju atribut/element ``span``.

```js
const statusColumnData = document.querySelectorAll('td:nth-child(4)');
assert.isNotEmpty(statusColumnData);
for (let e of statusColumnData) {
  assert.equal(e?.children[0]?.tagName, 'SPAN');
}
```

Svaki element redova `span` mora imati klasu `status` iz reda `Status`.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
assert.isNotEmpty(statusSpans);
for (let e of statusSpans) {
    assert.isTrue(e.classList.contains('status'));
}
```

Svaki element sa `.status` mora imati tekst `Read`, `To Read` ili `In Progress`, u zavisnosti od klase njegovog reda.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
const rows = Array.from(document.querySelectorAll('tr')).slice(1);
assert.isNotEmpty(statusSpans);
for (let i = 0; i < rows.length; i++) {
  switch (statusSpans[i]?.innerText.trim()) {
        case 'Read':
            assert.isTrue(rows[i].classList.contains('read'));
            break;
        case 'To Read':
            assert.isTrue(rows[i].classList.contains('to-read'));
            break;
        case 'In Progress':
            assert.isTrue(rows[i].classList.contains('in-progress'));
            break;
        default:
            assert.fail(); 
    }
}
```

Elementi za `td` iz reda `Rate` treba da imaju element/atribut `span`.

```js
const rateColumnData = document.querySelectorAll('tr td:last-child');
assert.isNotEmpty(rateColumnData);
for (let e of rateColumnData) {
    assert.equal(e.children[0]?.tagName, 'SPAN')
}
```

Svaki element `span` koji je direktni potomak elementa `td` iz redova `Rate` mora imati klasu `rate` kao prvu klasu.

```js
const rateSpans = document.querySelectorAll('tr td:last-child > span:first-child');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.classList[0], 'rate');
}
```

Svaki element `.rate` mora imati tri prazna elementa `span`.

```js
const rateSpans = document.getElementsByClassName('rate');
assert.isNotEmpty(rateSpans);
for (let e of rateSpans) {
    assert.equal(e.children.length, 3);
    for (let child of e.children) {
        assert.equal(child.tagName, 'SPAN');
        assert.equal(child.innerText.length, 0);
    }
}
```

Elementi ``.rate`` postavljeni unutar linija ``.read`` treba da imaju dodatnu klasu nakon klase ``rate`` sa vrednoĹˇÄ‡u ``one``, ``two``, ili ``three``, zavisno od pojedinaÄŤnih zahteva.

```js
const readBooksRates = document.querySelectorAll('.read .rate');
assert.isNotEmpty(readBooksRates);
for (let e of readBooksRates) {
    assert.oneOf(e.classList[1], ['one', 'two', 'three']);
}
```

Trebalo bi da imate selektor atributa elementa koji cilja linije sa klasom `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"]'));
```

Treba da koristite selektor atributa elementa za ciljanje redova sa klasom `read` i postavite njihov atribut `background-image` na kombinaciju linija boje koje birate.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="read"]')?.backgroundImage, 'linear-gradient(');
```

Trebalo bi da imate selektor atributa elementa koji cilja linije sa klasom `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]'));
```

Treba da koristiĹˇ selektor atributa elementa za ciljanje linija sa klasom `to-read` i postavljanje njihovog atributa `background-image` na kombinaciju boje linije koju biraĹˇ.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]')?.backgroundImage, 'linear-gradient(');
```

Trebalo bi da imaĹˇ selektor atributa elementa koji cilja linije sa klasom `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]'));
```

Treba da koristite selektor atributa elementa da ciljate na linije sa klasom `in-progress` i postavite njihov atribut `background-image` na kombinaciju boje linije koju birate.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]')?.backgroundImage, 'linear-gradient(');
```

Trebalo bi da postavite atribut `display` svakom elementu `span` na `inline-block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('span')?.getPropVal('display'), 'inline-block');
```

Trebalo bi da imate selektor atributa elementa za ciljanje elemenata `span` sa klasom `status` koji su potomci elemenata `tr` sa klasom `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]'));
```

Treba koristiti selektor atributa elementa za ciljanje elemenata `span` koji imaju klasu `status`, koji su potomci elemenata `tr` koji imaju klasu `to-read` i postavljanje njihovog atributa `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.border);
```

Trebalo bi da koristite selektor svojstva elementa za ciljanje elemenata `span` sa klasom `status`, koji su potomci elemenata `tr` sa klasom `to-read` i postavljanjem njihovog svojstva `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.backgroundImage);
```

Treba vam selektor atributa elementa koji cilja elemente ``span`` koji imaju klasu ``status``, a koji su potomci elemenata ``tr`` koji imaju klasu ``read``.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]'));
```

Trebalo bi da se koristi selektor atributa za ciljanje elemenata `span` koji imaju klasu `status`, a koji su potomci elemenata `tr` koji imaju klasu `read` i postaviti njihov atribut `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.border);
```

Treba koristiti selektor atributa elementa za ciljanje elemenata `span` sa klasom `status`, koji su potomci elemenata `tr` sa klasom `read` i postaviti njihov atribut `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.backgroundImage);
```

Treba vam selektor atributa elemenata koji cilja na elemente `span` sa klasom `status`, a koji su potomci elemenata `tr` sa klasom `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]'));
```

Treba koristiti selector atributa za ciljanje elemenata `span` sa klasom `status` koji su potomci elemenata `tr` sa klasom `in-progress` i postaviti njihov atribut `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.border);
```

Treba koristiti selektor atributa elementa da cilja na elemente `span` sa klasom `status` koji su potomci elemenata `tr` sa klasom `in-progress` i postaviti njihov atribut `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.backgroundImage);
```

Treba vam selektor atributa elementa da biste ciljali elemente `span` koji imaju klasu `status` i elemente `span` sa vrednoĹˇÄ‡u klase koja poÄŤinje od `rate`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.exists(selector1 || selector2);
```

Treba koristiti selektor atributa elementa za ciljanje elemenata `span` sa klasom `status` i elemenata `span` sa vrednoĹˇÄ‡u klase koja poÄŤinje sa `rate`, te postavljanjem njihovog atributa `height`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.height || selector2?.height);
```

Trebalo bi da koristite selektor atributa elemenata za ciljanje elemenata `span` koji imaju klasu `status` i elemente `span` koji imaju vrednost klase koja poÄŤinje od `rate`, i postavljanjem njihovog atributa `width`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.width || selector2?.width);
```

Trebalo bi da koristite selektor atributa elementa za ciljanje elemenata `span` koji imaju klasu `status` i elemenata `span` koji imaju vrednost klase poÄŤevĹˇi od `rate`, i postaviti njihov atribut `padding`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.padding || selector2?.padding);
```

Trebalo bi da imaĹˇ selektor atributa elementa koji cilja elemente `span` koji su direktni potomci elemenata `span` sa vrednoĹˇÄ‡u `class`, poÄŤevĹˇi od `rate`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span'));
```

Treba da koristite selektor atributa elementa za ciljanje elemenata `span` koji su direktni potomci elemenata `span` sa vrednoĹˇÄ‡u `class` poÄŤevĹˇi od `rate` i postavljanjem njihovog atributa `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border'));
```

Trebalo bi da se koristi selektor atributa elementa da cilja na elemente `span` koji su direktni potomci elemenata `span` sa vrednoĹˇÄ‡u od `class` poÄŤevĹˇi od `rate` i postavljanjem njihovog atributa `border-radius`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border-radius'));
```

Treba koristiti selektor atributa elementa za ciljanje elemenata `span` koji su direktni potomci elemenata `span` sa vrednoĹˇÄ‡u `class` poÄŤevĹˇi od `rate` i postavljanjem njihovog atributa `margin`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('margin'));
```

Trebalo bi da koristite selektor atributa elementa za ciljanje elemenata `span` koji su direktni potomci elemenata `span` sa vrednoĹˇÄ‡u `class` poÄŤevĹˇi od `rate` i postavljanjem njihovog atributa `height`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('height'));
```

Trebalo bi da koristiĹˇ selektor atributa elementa za ciljanje elemenata `span` koji su direktni potomci elemenata `span` sa vrednoĹˇÄ‡u `class` poÄŤevĹˇi od `rate` i postave njihov atribut `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('width'));
```

Trebalo bi da koristite selektor atributa elementa za ciljanje elemenata `span` koji su direktni potomci elemenata `span` sa vrednoĹˇÄ‡u `class`, poÄŤevĹˇi od `rate`, i postavite njihov atribut `background-color`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('background-color'));
```

Morate imati selektor svojstva elementa koji cilja prvog roditelja elemenata sa `span`, a koji sadrĹľe reÄŤ `one` kao deo njihove vrednosti od `class`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Trebalo bi da koristite selektor atributa elementa za ciljanje prvog roditelja elemenata sa `span` koji sadrĹľe reÄŤ `one` kao deo njihove vrednosti `class` i postavite njegov atribut `background-image` koristeÄ‡i `linear-gradient`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image').includes('linear-gradient('));
```

Trebalo bi da ima selektor svojstva elementa koji cilja na prvih dvoje roditelja elemenata `span`, koji sadrĹľe reÄŤ `two` kao deo njihove vrednosti `class`.

```js

const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Treba koristiti selektor atributa elementa da cilja na dva prvog roditelja elemenata sa `span` koji sadrĹľe reÄŤ `two` kao deo njihove vrednosti `class`, i postaviti njihove atribute `background-image` koristeÄ‡i `linear-gradient`.

```js
const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-of-type(2)',
  'span[class~="two"] > :nth-of-type(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]

assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.backgroundImage.includes('linear-gradient(')) 
```

Treba da ima selektor atributa elementa koji cilja elemente sa `span`, koji su potomci elemenata sa `span`, koji sadrĹľe reÄŤ `three` kao deo njihove vrednosti `class`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.exists(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

Trebalo bi da koristite selektor atributa elementa za ciljanje elemenata od `span` koji su potomci elemenata od `span` koji sadrĹľe reÄŤ `three` kao deo njihove vrednosti od `class` i postaviti njihove atribute od `background-image` koristeÄ‡i `linear-gradient`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.include(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image'), 'linear-gradient(');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main>
    <h1>Book Inventory</h1>
    <div class="table-container">
      <table>
        <thead>
          <tr>
            <th>Title</th>
            <th>Author</th>
            <th>Category</th>
            <th>Status</th>
            <th>Rate</th>
          </tr>
        </thead>
        <tbody>
          <tr class="read">
            <td>The Three Musketeers</td>
            <td>Alexandre Dumas</td>
            <td>Historical Novel</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate three">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="to-read">
            <td>The Plague</td>
            <td>Albert Camus</td>
            <td>Philosophical Novel</td>
            <td><span class="status">To Read</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>The Metamorphosis</td>
            <td>Franz Kafka</td>
            <td>Novella</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate one">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Dead Souls</td>
            <td>Nikolai Gogol</td>
            <td>Picaresque</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="in-progress">
            <td>Lord of the Flies</td>
            <td>William Golding</td>
            <td>Allegorical Novel</td>
            <td><span class="status">In Progress</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Do Androids Dream of Electric Sheep?</td>
            <td>Philip K. Dick</td>
            <td>Science Fiction</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
        </tbody>
      </table>
    </div>
  </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    font-family: Arial, sans-serif;
}

.table-container {
  overflow-x: auto;
}

table {
    border-collapse: collapse;
    width: 100vw;
}

thead {
    background-image: linear-gradient(hsl(197, 92%, 77%, 0.60), hsl(197, 92%, 50%, 0.60));
}

th {
    padding: 0.4em;
}

td {
    text-align: center;
    padding: 0.3em;
}

tr[class="read"] {
    background-image: linear-gradient(hsl(120, 100%, 85%, 0.6), hsl(120, 100%, 65%, 0.6));
}

tr[class="to-read"] {
    background-image: linear-gradient(hsl(120, 20%, 95%, 0.6), hsl(120, 20%, 75%, 0.6));
}

tr[class="in-progress"] {
    background-image: linear-gradient(hsl(40, 100%, 85%, 0.60), hsl(40, 100%, 65%, 0.6));
}

span {
    display: inline-block;
}

span[class="status"] {
    border-radius: 20%/60%;
}

span[class="status"],
span[class^="rate"] {
    height: 1.85em;
    width: 7em;
    padding: 0.3em;
}

tr[class="to-read"] span[class="status"] {
    border: 0.1em solid hsl(5, 100%, 45%);
    background-image: linear-gradient(hsl(5, 100%, 75%), hsl(5, 100%, 50%));
}

tr[class="read"] span[class="status"] {
    border: 0.1em solid hsl(120, 75%, 45%);
    background-image: linear-gradient(hsl(120, 75%, 75%), hsl(120, 75%, 50%));
}

tr[class="in-progress"] span[class="status"] {
    border: 0.1em solid hsl(40, 90%, 40%);
    background-image: linear-gradient(hsl(40, 90%, 75%), hsl(40, 90%, 50%));
}

span[class^="rate"] > span {
    border: 0.1em solid hsl(0, 0%, 50%);
    border-radius: 50%;
    background-color: hsl(0, 15%, 95%);
    height: 100%;
    width: 20%;
    margin: 0.1em;
}

span[class~="three"] span {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="two"] :nth-child(1),
span[class~="two"] :nth-child(2) {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="one"] :first-child {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}
```
