---
id: 691f7bbf76229bb97c60827d
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovom radionici, veštaćete se korišćenjem petlji `for...of` kreiranjem funkcije koja broji koliko puta je niz prisutan u nizu nizova.

Pre nego što napišete tu funkciju, kreiraćete jednostavniju koja loguje svaki karakter niza na konzolu.

Počnite definisanjem prazne funkcije nazvane `printCharacters` sa parametrom `str`.

# --hints--

You should create a `printCharacters` function.

```js
assert.isFunction(printCharacters);
```

Your `printCharacters` function should have one parameter called `str`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
