---
id: 6920c9cda1c7f056f75e4b1a
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Sada je vreme da napišete svoju prvu `for...of` petlju. U prethodnim lekcijama, naučili ste da možete koristiti `for...of` petlju za iteraciju po vrednostima iz iterabilnog objekta (npr. nizovi znakova ili nizovi).

Evo primera za `for...of` petlju:

```js
for (const num of [1, 2, 3]) {
  // code block to be executed
}
```

Dodajte `for...of` petlju sa praznim blokom koda unutar vaše funkcije. Trebalo bi da iterira po svakom znaku u argumentu `str`.

# --hints--

Your `printCharacters` function should have a `for...of` loop inside its code block.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functionprintCharacters\(str\)\{for\(\S+of\S+\)\{\S*?\}\}/);
```

Your `for...of` loop should iterate over the `str` argument of your function.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function printCharacters\(str\) \{\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+str/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
