---
id: 6921d688b8d58581d34459b9
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

U narednim koracima, kreiraćete funkciju koja broji koliko se niz znakova pojavljuje u nizu nizova znakova.

Da biste počeli, definišite praznu funkciju pod nazivom `getMatchedWordCount` sa parametrima `sentence` i `match` u tom redosledu.

# --hints--

You should create a `getMatchedWordCount` function.

```js
assert.isFunction(getMatchedWordCount);
```

Your `getMatchedWordCount` function should have two parameters: `sentence` and `match`, in that order.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
