---
id: 6921dae10700e300f26994e3
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Da biste videli izlaz vaše `getMatchedWordCount` funkcije u njenom trenutnom stanju, dodajte `console.log` izjavu ispod nje. Unutar `console.log()`, pozovite funkciju sa argumentom `["I", "really", "really", "really", "like", "to", "code"]` za parametar `sentence` i `"really"` za parametar `match`.

# --hints--

You should call the function `getMatchedWordCount` inside a `console.log` statement.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
assert.match(normalizedCode, /console\.log\(getMatchedWordCount\(\S*?\)\)/);
```

You should call the function `getMatchedWordCount` with the argument `["I", "really", "really", "really", "like", "to", "code"]` for the `sentence` parameter and `"really"` for the `match` parameter.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/)[0];
assert.match(normalizedConsoleLogCode, /\[('|"|`)I\1,('|"|`)really\2,('|"|`)really\3,('|"|`)really\4,('|"|`)like\5,('|"|`)to\6,('|"|`)code\7\],('|"|`)really\8/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
