---
id: acda2fb1324d9b0fa741e6b5
title: Izgradite alat za potvrđivanje završetka
challengeType: 26
dashedName: implement-a-string-ending-checker-function
---

# --description--

U ovom laboratorijumu, implementiraćete funkciju koja proverava da li niz završava sa zadanim ciljnim nizom.

**Cilj:** Ispunite priče korisnika ispod i učinite da svi testovi prođu kako biste zavili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da kreirate funkciju pod nazivom `confirmEnding` koja prima dva parametra: niz koji se proverava i niz sa kojim se proverava.
2. Funkcija bi trebalo da vrati `true` ako prvi niz završava sa drugim nizom, i `false` u suprotnom slučaju.
3. Ne smete da koristite metodu `.endsWith()`; umesto toga, koristite jednu od JavaScript metoda za podstring (substring) kako biste ovo postigli.

# --hints--

You should create a function named `confirmEnding`.

```js
assert.isFunction(confirmEnding);
```

`confirmEnding` should take 2 parameters.

```js
assert.lengthOf(confirmEnding, 2);
```

`confirmEnding("Bastian", "n")` should return `true`.

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation", "on")` should return `true`.

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor", "n")` should return `false`.

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` should return `false`.

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name", "name")` should return `true`.

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame", "same")` should return `true`.

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")` should return `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame", "game")` should return `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing", "mountain")` should return `false`.

```js
assert.isFalse(
  confirmEnding(
    "If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing",
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")` should return `true`.

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

Your code should not use the built-in method `.endsWith()` to solve the lab.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
