---
id: 66fbcdfed6848e48947dbbc9
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovom radionici, kreiraćete praćenje recepata koristeći JavaScript objekte.

Počnite kreiranjem praznog niza pod nazivom `recipes`. Ovo je niz u koji ćete kasnije dodati objekte recepata.

# --hints--

You should create an array named `recipes`.

```js
assert.isArray(recipes);
```

Your `recipes` array should be empty.

```js
assert.isEmpty(recipes);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
