---
id: 66fbcf750a62784cf11f562d
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Kreiraj objekat pod imenom `recipe1`. Unutar objekta `recipe1`, kreiraj svojstvo `name` sa vrednošću `"Spaghetti Carbonara"`.

Takođe unutar objekta `recipe1`, kreiraj svojstvo `ingredients` sa nizom kao vrednošću. Niz bi trebalo da sadrži `spaghetti`, `Parmesan cheese`, `pancetta` i `black pepper`.

# --hints--

You should create an object named `recipe1`.

```js
assert.isObject(recipe1);
```

Your `recipe1` object should have a `name` property.

```js
assert.property(recipe1, 'name');
```

Your `name` property should be set to `Spaghetti Carbonara`.

```js
assert.equal(recipe1.name, 'Spaghetti Carbonara');
```

Your `recipe1` object should have an `ingredients` property.

```js
assert.property(recipe1, 'ingredients');
```

Your `ingredients` property should have an array value.

```js
assert.isArray(recipe1.ingredients);
```

The array value of your `ingredients` property should have `spaghetti`, `Parmesan cheese`, `pancetta`, and `black pepper` in it.

```js
assert.deepEqual(recipe1.ingredients, ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];
--fcc-editable-region--

--fcc-editable-region--
```
