---
id: 670e4f45f7116c0f216a5177
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Pre nego što nastavite, trebalo bi da vežbate kako pristupiti svojstvima objekta.

Možete koristiti ili tačku (.`.`) ili zagradnu (.`[]`) notaciju za ovo. Evo primera:

```js
const person = {
  name: "John",
  age: 30,
  job: "Software Engineer"
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

Pristupite svojstvu `name` objekta `recipe1`, i dodelite ga promenljivoj `recipe1Name`.

Zatim, pristupite svojstvu `cookingTime` objekta `recipe2` i dodelite ga promenljivoj `recipe2CookingTime`.

Konačno, pristupite svojstvu `ingredients` objekta `recipe3` i dodelite ga promenljivoj `recipe3Ingredients`.

Uverite se da su sve varijable koje ste kreirali logovane u konzolu.

# --hints--

You should create a `recipe1Name` variable.

```js
assert.isDefined(recipe1Name);
```

You should assign the value of the `name` property of `recipe1` to your `recipe1name` variable.

```js
assert.strictEqual(recipe1Name, recipe1.name);
```

You should log `recipe1Name` to the console.

```js
assert.match(code, /console\.log\(\s*recipe1Name\s*\)/);
```

You should create a `recipe2CookingTime` variable.

```js
assert.isDefined(recipe2CookingTime);
```

You should assign the value of the `cookingTime` property of `recipe2` to your `recipe2CookingTime` variable.

```js
assert.strictEqual(recipe2CookingTime, recipe2.cookingTime);
```

You should log `recipe2CookingTime` to the console.

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/);
```

You should create a `recipe3Ingredients` variable.

```js
assert.isDefined(recipe3Ingredients);
```

You should assign the value of the `ingredients` property of `recipe3` to your `recipe3Ingredients` variable.

```js
assert.strictEqual(recipe3Ingredients, recipe3.ingredients);
```

You should log `recipe3Ingredients` to the console.

```js
assert.match(code, /console\.log\(recipe3Ingredients\)/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
