---
id: 60a3e3396c7b40068ad6998b
title: Korak 32
challengeType: 0
dashedName: step-32
---

# --description--

Korisno je da vaše margine budu usmerene u jednom pravcu.

U ovom slučaju, donja margina elementa `.one` gura `.two` dole za 20 piksela.

U selektoru `.two`, koristite skraćenu svojstvo (shorthand property) `margin` da postavite gornju marginu na `0`, horizontalnu marginu na `auto` i donju marginu na `20px`. Ovo će ukloniti njegovu gornju marginu, centrirati ga horizontalno i postaviti njegovu donju marginu na 20 piksela.

# --hints--

You should set the `margin` property to `0 auto 20px`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '0px auto 20px');
assert(hasMargin);
```

Your `.two` element should have a `margin` value of `0 auto 20px`.

```js
const twoMargin = new __helpers.CSSHelp(document).getStyle('.two')?.getPropertyValue('margin');
assert(twoMargin === '0px auto 20px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
