---
id: 60a3e3396c7b40068ad6997e
title: Korak 20
challengeType: 0
dashedName: step-20
---

# --description--

Sada je `.one` centriran horizontalno, ali njegova gornja margina se proteže preko platna i na ivicu okvira, pomerajući celo platno nadole za 20 piksela.

Dodajte `padding` od `1px` u element `.canvas` kako biste elementu `.one` dali nešto čvrsto na šta se može osloniti/odmaknuti.

# --hints--

You should set the `padding` property to `1px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '1px');
assert(hasPadding);
```

Your `.canvas` element should have a `padding` value of `1px`.

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('padding');
assert(canvasPadding === '1px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--
  
--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
