---
id: 60a3e3396c7b40068ad6996e
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Margina je površina izvan kutije, i može se koristiti za kontrolu razmaka između drugih kutija ili elemenata.

Ovde donji element ima veću gornju marginu, što ga pomera niže na stranici.

Sada kada ste brzo pregledali CSS box model, vreme je da počnete sa slikom Rothka.

Uklonite element `img`.

# --hints--

You should not have an `img` element in your code.

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
