---
id: 60a3e3396c7b40068ad6996c
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

U prethodnim lekcijama o CSS box modelu, naučili ste da je svaki HTML element tretiran kao kutija sa četiri površine.

Zamislite da primite kutiju od vašeg omiljenog onlajn prodavnice — sadržaj je predmet u kutiji, ili u vašem slučaju, zaglavlje, paragraf ili element slike.

Promenite atribut `src` u `<img>` sa `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` na `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`.

# --hints--

The `img` element should have a `src` of `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```

