---
id: 60a3e3396c7b40068ad69979
title: Korak 15
challengeType: 0
dashedName: step-15
---

# --description--

Dodajte novi `div` element unutar vašeg `.canvas` elementa.

Dajte novom `div` atribut `class` sa vrednošću `one`. Ovo će biti vaš prvi pravougaonik.

# --hints--

You should create a new `div` element.

```js
assert(document.querySelectorAll('div').length === 3);
```

You should nest the new `div` element within your `.canvas` element.

```js
assert(document.querySelector('.canvas').children[0].tagName === 'DIV');
```

Your new `div` should have a `class` attribute with a value `one`.

```js
assert(document.querySelector('.canvas').children[0].className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
