---
id: 60a3e3396c7b40068ad6996d
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Sadržaj je okružen prostorom koji se naziva padding, slično tome kako omotnica od sveže odvaja predmet od kutije oko njega.

Razmislite o granici kao o kartonskoj kutiji u koju je vaš predmet poslat.

Promenite atribut `src` u `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`.

# --hints--

The `img` element should have a `src` of `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
