---
id: 60a3e3396c7b40068ad69971
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Vreme je za CSS.

Iako vaš `<div>` nema tekst, tretiraće ga i dalje kao kutiju sa sadržajem.
Napišite CSS pravilo koje koristi klasni selektor `.canvas` i postavi njegovu `width` na 500 piksela.
Evo CSS pravila koje postavlja širinu klase `card` na 300 piksela:

```css
.card {
  width: 300px;
}
```

# --hints--

Your code should have a `.canvas` selector.

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

You should set the `width` property to `500px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

Your `.canvas` element should have a `width` of `500px`.

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
