---
id: 66b205e6eacba4c4e54ea434
title: Izgradite maskeru za e-mail adrese
challengeType: 26
dashedName: build-an-email-masker
---

# --description--

U ovom laboratorijumu, maskiraćete deo korisničkog imena adrese e-pošte zvezdicama. Maskiranje je termin koji se koristi za skrivanje ili zamenu osetljivih informacija zvezdicama ili drugim karakterima.

Na primer, ako je adresa e-pošte bila `myEmail@email.com`, maskirana adresa e-pošte biće `m*****l@email.com`.

**Cilj:** Ispunite sledeće korisničke priče i prođite sve testove kako biste završili laboratorijum.

**Korisničke priče:**

1. Kreirajte funkciju pod nazivom `maskEmail` koja prima `email` kao argument.
2. Unutar funkcije, trebalo bi da maskirate `email` i dodate naziv domena na njega. Zapamtite da možete koristiti metode poput `slice`, `repeat`, `indexOf` ili čak `replace` kako biste vam pomogli.
3. Izvan funkcije, deklarišite varijablu pod nazivom `email` da sačuvate adresu e-pošte koju želite da maskirate.
4. Pozovite funkciju `maskEmail` sa varijablom `email` i ispišite rezultat na konzolu. 
5. `maskEmail("apple.pie@example.com")` bi trebalo da vrati `"a*******e@example.com"`.
6. `maskEmail("freecodecamp@example.com")` bi trebalo da vrati `"f**********p@example.com"`.
7. `maskEmail("info@test.dev")` bi trebalo da vrati `"i**o@test.dev"`.
8. `maskEmail("user@domain.org")` bi trebalo da vrati `"u**r@domain.org"`.

# --hints--

You should define a function named `maskEmail`.

```js
assert.isFunction(maskEmail);
```

The `maskEmail` function should take a string, `email` as an argument.

```js
assert.match(maskEmail.toString(), /\s*maskEmail\(\s*\w+\s*\)/);
```

Outside the function, you should have an `email` variable.

```js
assert.isDefined(email);
```

You should assign a valid email address to your `email` variable.

```js
assert.match(email, /^[a-zA-Z0-9._%-]+\+?[a-zA-Z0-9._%-]+@[a-zA-Z0-9.-]+\.[a-zA-Z]{2,}$/);
```

`maskEmail("apple.pie@example.com")` should return `"a*******e@example.com"`.

```js
assert.strictEqual(maskEmail("apple.pie@example.com"),"a*******e@example.com");
```

`maskEmail("freecodecamp@example.com")` should return `"f**********p@example.com"`.

```js
assert.strictEqual(maskEmail("freecodecamp@example.com"), "f**********p@example.com");
```

`maskEmail("info@test.dev")` should return `"i**o@test.dev"`.

```js
assert.strictEqual(maskEmail("info@test.dev"), "i**o@test.dev");
```

`maskEmail("user@domain.org")` should return `"u**r@domain.org"`.

```js
assert.strictEqual(maskEmail("user@domain.org"), "u**r@domain.org");
```

Your `maskEmail` should produce the correct result.

```js
assert.strictEqual(maskEmail("johnappleseed@email.com"), "j***********d@email.com");
assert.strictEqual(maskEmail("john.doe@example.com"), "j******e@example.com");
assert.strictEqual(maskEmail("jane.smith@sample.org"), "j********h@sample.org");
assert.strictEqual(maskEmail("michael.jordan@testmail.net"), "m************n@testmail.net");
assert.strictEqual(maskEmail("mike+nospam@testmail.net"), "m*********m@testmail.net");
```

You should log the output of calling `maskEmail` with `email` as argument.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(maskEmail\(email\)\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function maskEmail(email) {
  const atIndex = email.indexOf("@");
  const userName = email.slice(0, atIndex);
  const domain = email.slice(atIndex);

  const maskedName =
    userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1] + domain;

  return maskedName;
}

const email = "apple.pie@example.com";
console.log(maskEmail(email));
```
