---
id: 68e42de6396f5c49d90d1951
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

U prethodnoj lekciji, naučili ste o metodi `replace()` koja se koristi za lociranje podstringa u stringu i zamenu ga drugom vrednošću.

Evo primera:

```js
const text = "I love JavaScript!";
const newText = text.replace("JavaScript", "coding");
// Result: "I love coding!"
```

Zapamtite da su stringovi nemenjivi (immutable), što znači da u tom primeru nije modifikovan originalni `text`. Umesto toga se kreira novi string. 

Kreirajte novu varijablu nazvanu `replacedString` i dodelite joj `originalString.replace("cats", "dogs")`.

# --hints--

You should have a variable called `replacedString`.

```js
assert.exists(replacedString);
```

Your `replacedString` variable should be a string.

```js
assert.isString(replacedString);
```

You should assign `originalString.replace("cats", "dogs")` to the `replacedString` variable.

```js
assert.equal(replacedString, "I love dogs.");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--

--fcc-editable-region--
```
