---
id: 68e43e884ef09d4d2ab05631
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Sada je vreme da vidimo rezultat korišćenjem metode `replace()`.

Počnite dodavanjem `console.log()` sa vrednošću `"After using the replace() method:"`.

Zatim ispod tog `console.log()`, dodajte još jedan `console.log()` sa varijablom `replacedString`.

Pogledajte konzolu i trebalo bi da vidite da nova rečenica kaže `"I love dogs."` umesto originalnog `"I love cats."`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should log `"After using the replace() method:"` to the console.

```js
assert.equal(getLogs()[2], "After using the replace() method:");
```

You should log the `replacedString` variable to the console. 

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingReplacedStr = codeWithoutComments.match(/console\.log\(\s*replacedString\s*\)/g)

assert.isAtLeast(loggingReplacedStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
--fcc-editable-region--

--fcc-editable-region--
```
