---
id: 68e44c305f1f935d949d76bb
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

Kao što je pomenuto u prethodnom koraku, postoji dodatni razmak nakon poslednjeg `"love"` u rečenici `"I love love love  learning."`. Da biste uklonili taj dodatni razmak, možete koristiti metod `trimEnd()`.

Na kraju metode `"love ".repeat(3)`, lančujte metodu `trimEnd()`. Metode možete lančovati ovako:

```js
.firstMethod().secondMethod()
```

Sada kada proverite konzolu, trebalo bi da vidite da je uklonjen dodatni razmak.

I sa tom poslednjom promenom, vaš radionica je gotova!

# --hints--

You should chain the `trimEnd()` method to the `"love ".repeat(3)` method.

```js
assert.equal(repeatedLove, "love ".repeat(3).trimEnd());
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--
const repeatedLove = "love ".repeat(3)
--fcc-editable-region--
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```

# --solutions--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3).trimEnd();
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```
