---
id: 68e42b62d7093345c6975f22
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Ispod vaše varijable `originalString`, dodajte `console.log()` sa stringom `"Original string:"`.

Zatim ispod tog `console.log()`, dodajte još jedno `console.log()` sa varijablom `originalString`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should log the string `"Original string:"` to the console.

```js
assert.equal(getLogs()[0], "Original string:");
```

You should log the `originalString` variable to the console. 

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingOriginalStr = codeWithoutComments.match(/console\.log\(\s*originalString\s*\)/g)

assert.isAtLeast(loggingOriginalStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
--fcc-editable-region--

--fcc-editable-region--
```
