---
id: 68e41084bb275539f72de9fe
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovom radionici, vežvaćete rad sa metodama `replace()`, `replaceAll()` i `repeat()` gradeći aplikaciju za transformisanje stringova.

Za početak, kreirajte varijablu nazvanu `originalString` i dodelite joj vrednost `"I love cats."`.

# --hints--

You should have a variable called `originalString`.

```js
assert.exists(originalString);
```

Your `originalString` variable should be a string.

```js
assert.isString(originalString);
```

You should assign the string `"I love cats."` to the `originalString` variable.

```js
assert.equal(originalString, "I love cats.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
