---
id: 68e4417b7317835208145323
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Sledeći korak je kreiranje nove varijable pod nazivom `dogsOnlySentence` i dodeljivanje joj `exampleSentence.replace("cats", "dogs")`.

Zatim ispod te varijable, logujte na konzolu `"Replacing all occurrences of cats with dogs:"`. Ispod toga, `console.log()`, dodajte još jedno `console.log()` za varijablu `dogsOnlySentence`.

Sada bi trebalo da vidite tekst `"I love dogs and cats are so much fun!"` logovan na konzolu. Ali to nije željeni rezultat. U sledećem koraku, naučićete kako da ga popravite tako da tekst čita `"I love dogs and dogs are so much fun!"`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should have a `dogsOnlySentence` variable.

```js
assert.exists(dogsOnlySentence);
```

Your `dogsOnlySentence` variable should be a string.

```js
assert.isString(dogsOnlySentence);
```

You should assign `exampleSentence.replace("cats", "dogs")` to your `dogsOnlySentence` variable.

```js
assert.equal(dogsOnlySentence, "I love dogs and cats are so much fun!");
```

You should log `"Replacing all occurrences of cats with dogs:"` to the console.

```js
assert.equal(getLogs()[6], "Replacing all occurrences of cats with dogs:");
```

You should log the `dogsOnlySentence` variable to the console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingDogsOnlySentence = codeWithoutComments.match(/console\.log\(\s*dogsOnlySentence\s*\)/g)

assert.isAtLeast(loggingDogsOnlySentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--

--fcc-editable-region--
```
