---
id: 68e448ec4ac4205968c0a805
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Bilo bi lepo da reč `"love"` bude ponovljena tri puta u rečenici `"I love learning!"`. 

Kao što se sećate iz prethodne lekcije, možete ponoviti niz određeni broj puta koristeći metod `repeat()`. Evo primera:

```js
const word = "Hello!";
const repeatedWord = word.repeat(3);

console.log(repeatedWord);  // "Hello!Hello!Hello!"
```

Pošto su nizovi nemenjivi (immutable), ovaj metod neće modifikovati originalni niz. Vraćaće novi niz sa ponovljenim sadržajem. 

Kreirajte varijablu nazvanu `repeatedLove` i dodelite joj `"love ".repeat(3)`. Zatim logujte u konzolu varijablu `repeatedLove` kako biste videli rezultat. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should have a variable called `repeatedLove`.

```js
assert.exists(repeatedLove);
```

Your `repeatedLove` variable should be a string.

```js
assert.isString(repeatedLove);
```

You should assign `"love ".repeat(3)` to the `repeatedLove` variable.

```js
assert.equal(repeatedLove, "love ".repeat(3));
```

You should log the `repeatedLove` variable to the console.

```js
assert.equal(getLogs()[10], repeatedLove);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--

--fcc-editable-region--
```
