---
id: 587d8259367417b2b2512c85
title: Implementiraj algoritam Selektivne sortiranja
challengeType: 26
dashedName: implement-the-selection-sort-algorithm
---

# --description--

U ovom laboratorijumu ćete implementirati algoritam selekcionog sortiranja (selection sort). Selektivno sortiranje radi tako što bira minimalnu vrednost u listi i zamjenjuje je sa prvom vrednošću u listi. Zatim počinje od druge pozicije, bira najmanju vrednost u preostalom popisu i zamjenjuje je sa drugim elementom. Nastavlja iterirati kroz listu i menjati elemente sve dok ne stigne do kraja liste. Tada je lista sortirana. Selektivno sortiranje ima kvadratnu vremensku složenost u svim slučajevima.

**Cilj:** Ispunite korisničke priče ispod i prođite sve testove kako biste završili laboratorijum.

**Korisničke priče:**

1. Trebalo bi da imate funkciju `selectionSort`.
1. Funkcija `selectionSort` treba da primi niz i vrati niz sa istim elementima, ali sortiranim od najmanjeg do najvećeg.

# --hints--

`selectionSort` should be a function.

```js
assert.isFunction(selectionSort);
```

`selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` should return an array that is unchanged except for order.

```js
assert.sameMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`selectionSort` should return a sorted array (least to greatest).

```js
assert.sameOrderedMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
)
```

`selectionSort` should not use the built-in `.sort()` method.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
