---
id: 6140d4bc9db3c81c51a09ab7
title: Korak 15
challengeType: 0
dashedName: step-15
---

# --description--

Nastavljajući obrazac, izaberite sledeće `.cabin` elemente i primenite na njih specifična pravila:

- Drugi `.cabin` treba da ima svojstvo `right` postavljeno na `17%` i svojstvo `top` postavljeno na `93.5%`.
- Treći `.cabin` treba da ima svojstvo `right` postavljeno na `67%` i svojstvo `top` postavljeno na `93.5%`.
- Četvrti `.cabin` treba da ima svojstvo `left` postavljeno na `-8.5%` i svojstvo `top` postavljeno na `50%`.
- Peti `.cabin` treba da ima svojstvo `left` postavljeno na `17%` i svojstvo `top` postavljeno na `7%`.
- Šesti `.cabin` treba da ima svojstvo `right` postavljeno na `17%` i svojstvo `top` postavljeno na `7%`.

# --hints--

You should have a `.cabin:nth-of-type(2)` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)'));
```

Your `.cabin:nth-of-type(2)` selector should have a `right` property set to `17%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.right, '17%');
```

Your `.cabin:nth-of-type(2)` selector should have a `top` property set to `93.5%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.top, '93.5%');
```

You should have a `.cabin:nth-of-type(3)` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)'));
```

Your `.cabin:nth-of-type(3)` selector should have a `right` property set to `67%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.right,'67%');
```

Your `.cabin:nth-of-type(3)` selector should have a `top` property set to `93.5%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.top, '93.5%');
```

You should have a `.cabin:nth-of-type(4)` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)'));
```

Your `.cabin:nth-of-type(4)` selector should have a `left` property set to `-8.5%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.left, '-8.5%');
```

Your `.cabin:nth-of-type(4)` selector should have a `top` property set to `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.top, '50%');
```

You should have a `.cabin:nth-of-type(5)` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)'));
```

Your `.cabin:nth-of-type(5)` selector should have a `left` property set to `17%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.left, '17%');
```

Your `.cabin:nth-of-type(5)` selector should have a `top` property set to `7%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.top, '7%');
```

You should have a `.cabin:nth-of-type(6)` selector.

```js
const def = (s) => new __helpers.CSSHelp(document).getStyle(s)
assert.exists(def('.cabin:nth-of-type(6)') || def('.cabin:last-of-type'));
```

Your `.cabin:nth-of-type(6)` selector should have a `right` property set to `17%`.

```js
const right = (s) => new __helpers.CSSHelp(document).getStyle(s)?.right
assert.equal(right('.cabin:nth-of-type(6)') || right('.cabin:last-of-type'),'17%');
```

Your `.cabin:nth-of-type(6)` selector should have a `top` property set to `7%`.

```js
const top = (s) => new __helpers.CSSHelp(document).getStyle(s)?.top
assert.equal(top('.cabin:nth-of-type(6)') || top('.cabin:last-of-type'),'7%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
--fcc-editable-region--

--fcc-editable-region--
```
