---
id: 646c5e727e56e743c9aed4a1
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

CSS pozicioniranje vam omogućava da odredite kako želite da se element pozicionira u pretraživaču. Ima svojstvo `position` koje možete postaviti na `static`, `absolute`, `relative`, `sticky` ili `fixed`.

Nakon što postavite svojstvo `position` elementa, možete pomeriti element postavljanjem vrednosti u pikselima ili procentu za jedno ili više svojstava `top`, `right`, `left` ili `bottom`.

`static` je podrazumevano pozicioniranje za sve elemente. Ako ga dodelite nekom elementu, nećete moći da ga pomerite pomoću `top`, `right`, `left` ili `bottom`.

Dajte `.cat-head` svojstvo `position` sa vrednošću `static`, a zatim postavite svojstva `top` i `left` na `100px`.

# --hints--

Your `.cat-head` selector should have a `position` property set to `static`. Make sure you add a semicolon.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'static')
```

Your `.cat-head` selector should have a `top` property set to `100px`. Make sure you add a semicolon.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px')
```

Your `.cat-head` selector should have a `left` property set to `100px`. Make sure you add a semicolon.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {


  
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
