---
id: 6476fb7cbaafa36d65e9cf35
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

`fixed` je vrednost svojstva `position` koja vam omogućava da učinite element fiksiranim na stranici bez obzira gde korisnik skroluje po stranici.

Morate napraviti još nekoliko oznaka (markupa) da biste videli kako pozicioniranje `fixed` funkcioniše. U vašem HTML-u, kreirajte `div` element sa klasom `box`.

# --hints--

You should have a `div` element.

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

Your `div` element should have the class `box`.

```js
assert.equal(document.querySelectorAll('div')[1]?.getAttribute('class'), 'box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  top: 300px;
  left: 400px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
