---
id: 6476f7a4827bcc61682f2347
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Kada koristite vrednost `absolute` za svojstvo `position`, element se izvlači iz normalnog toka dokumenta, a zatim je njegova pozicija određena svojstvima `top`, `right`, `bottom` i `left`.

Postavite svojstvo `position` vašeg `.cat-head` elementa na `absolute`, a zatim postavite svojstva `top` i `left` na bilo koju pozitivnu piksel vrednost.

<!-- **Napomena**: Možete eksperimentisati sa svojstvima `top`, `left`, `bottom` i `right` ovde, ali test će proći samo za `top` od `300px`, i levo od `400px`. -->

# --hints--

Your `.cat-head` selector should have a `position` property set to `absolute`. Make sure you add a semicolon.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'absolute');
```

Your `.cat-head` selector should have a `top` property set to any number in `px`. Make sure you add a semicolon.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/);
```

Your `.cat-head` selector should have a `left` property set to any number in `px`. Make sure you add a semicolon.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
