---
id: 6476f5c17f99146071ee884c
title: Korak 8
challengeType: 0
dashedName: step-8
---

# --description--

Mogli ste videti da se ništa ne dešava. Element `.cat-head` nije se pomerio uprkos postavljanju `top` i `left` na vrednost `100px` za svaki od njih. Ali to nije slučaj sa pozicioniranjem `relative`.

Kada koristite vrednost `relative`, element je i dalje pozicioniran prema normalnom toku dokumenta, ali se aktiviraju vrednosti `top`, `left`, `bottom` i `right`.

Umesto `static`, dajte vašem `.cat-head` poziciju od `relative`, i ostavite obe svojstvo, `top` i `left`, kako su.

# --hints--

Your `.cat-head` selector should have a `position` property set to `relative`. Make sure you add a semicolon.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'relative');
```

Your `.cat-head` selector should have a `top` property set to `100px`. Make sure you add a semicolon.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px');
```

Your `.cat-head` selector should have a `left` property set to `100px`. Make sure you add a semicolon.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
