---
id: 6477062778c85972eb648030
title: Korak 14
challengeType: 0
dashedName: step-14
---

# --description--

Sada treba da centrirate glavu mačke.

Dajte elementu `.cat-head` svojstvo `position` postavljeno na `absolute`. Postavite vrednost `0` za svojstva `right`, `left`, `top`, `bottom`, zatim postavite njegovo svojstvo `margin` sa svih strana na `auto`. To je jedan način da se centrir element vertikalno i horizontalno korišćenjem CSS pozicioniranja.

# --hints--

Your `.cat-head` selector should have a `position` property set to `absolute`. Make sure you add a semicolon.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'absolute');
```

Your `.cat-head` selector should have a `top` property set to `0`. Make sure you add a semicolon.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '0px');
```

Your `.cat-head` selector should have a `left` property set to `0`. Make sure you add a semicolon.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '0px');
```

Your `.cat-head` selector should have a `right` property set to `0`. Make sure you add a semicolon.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right, '0px');
```

Your `.cat-head` selector should have a `bottom` property set to `0`. Make sure you add a semicolon.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom, '0px');
```

Your `.cat-head` selector should have a `margin` property set to `auto`. Make sure you add a semicolon.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
