---
id: 64770351e8586671ec0911f0
title: Korak 13
challengeType: 0
dashedName: step-13
---

# --description--

Pozicioniranje ``sticky`` je hibrid pozicioniranja ``relative`` i ``fixed``. Omogućava elementu da se **zalepi** za određenu poziciju unutar svog kontejnera ili prozora, na osnovu položaja skrolovanja.

Promenite vrednost svojstva ``position`` za ``.cat-head`` na ``sticky``, postavite ``top`` na ``0``, a zatim uklonite ``left`` i njegovu vrednost.

**Napomena**: Da biste videli kako radi ``sticky``, morate postaviti nekoliko tekstova pre i posle vašeg elementa ``.cat-head``. Ako skrolujete dole nakon toga, videćete da se ``.cat-head`` zaglavi na vrhu i ostaje tamo.

# --hints--

Your `.cat-head` selector should have a `position` property set to `sticky`. Make sure you add a semicolon.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'sticky');
```

Your `.cat-head` selector should have a `top` property set to `0`. Make sure you add a semicolon.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '0px');
```

You should not have the `left` property and its value in your code.

```js
assert.notMatch(code, /left:\s*100px;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
