---
id: 647d855ac12cd436059acd39
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Koristite univerzalni selektor da dodate `box-sizing: border-box;` u vaš CSS. Ovo osigurava da elementi uključuju padding i border u svoju navedenu širinu i visinu.

# --hints--

You should have the universal selector (`*`).

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Your universal selector should have a `box-sizing` property set to `border-box`.

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
