---
id: 68507cadbf36aa089a84ce2d
title: Implementirati Luhnov algoritam
challengeType: 27
dashedName: implement-the-luhn-algorithm
---

# --description--

Luhnov algoritam, poznat i kao algoritmi "modulus 10" ili "mod 10", jednostavna je formula za proveru zbirke koja se koristi za validaciju različitih identifikacionih brojeva, poput brojeva kreditnih kartica. Ovo su koraci za validaciju broja korišćenjem Luhnovog algoritma:

- Počevši od desne strane i isključujući najdesni cifru (provernu cifru), udvojte vrednost svake druge cifre.
- Ako je rezultat udvostručavanja cifre veći od `9`, saberi cifre kako bi se dobila jedna cifra. Alternativno, možete oduzeti `9` od rezultata.
- Uzmite zbir svih cifara uključujući provernu cifru.
- Ako je zbir svih cifara deljiv sa `10`, onda je broj validan; inače nije validan.

Na primer, razmotrite broj `453914881`. Koraci za njegovu validaciju korišćenjem Luhnovog algoritma su:

 ```md 
 Account number      4   5   3   9   1   4   8   8   1 
 Double every other  4  10   3  18   1   8   8  16   1 
 Sum 2-char digits   4   1   3   9   1   8   8   7   1 
 ``` 

- Zatim saberi sve brojeve, `4 + 1 + 3 + 9 + 1 + 8 + 8 + 7 + 1 = 42`.
- Pošto `42` nije deljiv sa `10`, broj je nevalidan.

U ovom laboratorijumu, kreiraćete validator kreditne kartice koristeći Luhnov algoritam.

**Cilj:** Ispunite priče korisnika ispod i učinite da svi testovi prođu kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da definišete funkciju pod nazivom `verify_card_number` koja prima niz cifara (koji predstavlja broj kartice) i proverava da li je on validan prema Luhnovom algoritmu.

2. Unutar funkcije `verify_card_number`:

   - Trebalo bi da obradite sve crtice ili razmake koji mogu biti prisutni u broju kartice koji mu je prosleđen.
   - Vratite `VALID!` ako je broj kartice validan; inače, vratite `INVALID!`.

Kada završite projekat, trebalo bi da vidite sledeće poruke u zavisnosti od unosa:

| Card Number         | Message  |
| ------------------- | -------- |
| 453914889           | VALID!   |
| 4111-1111-1111-1111 | VALID!   |
| 1234 5678 9012 3456 | INVALID! |

# --hints--

You should have a function named `verify_card_number`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('verify_card_number')
    `)
})
```

`verify_card_number('453914889')` should return `VALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914889') == 'VALID!'
    `)
})
```

`verify_card_number('4111-1111-1111-1111')` should return `VALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('4111-1111-1111-1111') == 'VALID!'
    `)
})
```

`verify_card_number('453914881')` should return `INVALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914881') == 'INVALID!'
    `)
})
```

`verify_card_number('1234 5678 9012 3456')` should return `INVALID!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('1234 5678 9012 3456') == 'INVALID!'
    `)
})
```

`verify_card_number` should return `VALID!` when called with a valid credit card number.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6467') == 'VALID!'
    assert verify_card_number('6011 1111 1111 1117') == 'VALID!'
    assert verify_card_number('3782-822463-10005') == 'VALID!'
    assert verify_card_number('5555 5555 5555 4444') == 'VALID!'
    assert verify_card_number('3056 9309 0259 04') == 'VALID!'
    assert verify_card_number('6011000990139424173') == 'VALID!'
    `)
})
```

`verify_card_number` should return `INVALID!` when called with an invalid credit card number.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6466') == 'INVALID!'
    assert verify_card_number('6011 1111 1111 1116') == 'INVALID!'
    assert verify_card_number('3782 822463 10004') == 'INVALID!'
    assert verify_card_number('5555-5555-5555-4440') == 'INVALID!'
    assert verify_card_number('3056 9309 0259 05') == 'INVALID!'
    assert verify_card_number('6011000990139424175') == 'INVALID!'
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def verify_card_number(card_number):
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    sum_of_odd_digits = 0
    card_number_reversed = translated_card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            number = (number // 10) + (number % 10)
        sum_of_even_digits += number

    total = sum_of_odd_digits + sum_of_even_digits
    if total % 10 == 0:
        return 'VALID!'
    else:
        return 'INVALID!'

# Example usage
verify_card_number('4111-1111-4555-1142')
```
