---
id: afcc8d540bea9ea2669306b6
title: Napraviti funkciju za ponavljanje stringa
challengeType: 26
dashedName: build-a-string-repeating-function
---

# --description--

U ovom laboratorijumu, kreiraćete funkciju koja ponavlja zadani niz određeni broj puta. Za potrebe ovog laboratorijuma, ne koristite ugrađenu metodu `.repeat()`.

**Cilj:** Ispunite korisničke priče ispod i osigurajte da svi testovi prođu kako biste završili laboratorijum.

**Korisničke priče:**

1. Trebalo bi da kreirate funkciju pod nazivom `repeatStringNumTimes` koja prima dva parametra: niz (string) i broj.
2. Funkcija treba da vrati niz ponovljen određeni broj puta.
3. Ako je broj manji ili jednak nuli, funkcija treba da vrati prazan niz (string).

# --hints--

You should create a function named `repeatStringNumTimes`.

```js
assert.isFunction(repeatStringNumTimes);
```

`repeatStringNumTimes` should take two parameters.

```js
assert.lengthOf(repeatStringNumTimes, 2);
```

The function `repeatStringNumTimes` should always return a string.

```js
assert.isString(repeatStringNumTimes('hello', 3));
```

`repeatStringNumTimes("*", 3)` should return the string `***`.

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)` should return the string `abcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)` should return the string `abcabcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)` should return the string `abc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)` should return the string `********`.

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)` should return an empty string (`""`).

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

`repeatStringNumTimes("abc", 0)` should return `""`.

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

The built-in `repeat()` method should not be used.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
