---
id: a39963a4c10bc8b4d4f06d7e
title: Implement funkciju za uklanjanje vrednosti
challengeType: 26
dashedName: implement-a-value-remover-function
---

# --description--

U ovom laboratorijumu, kreiraćete funkciju koja prima početni niz kao prvi argument, praćen jedan ili više dodatnih argumenata koji predstavljaju vrednosti koje treba ukloniti.

**Cilj**: Ispunite korisničke priče ispod i osigurajte da svi testovi prođu kako biste završili laboratorijum.

**Korisničke priče:**

1. Trebalo bi da imate funkciju `destroyer` koja prima niz i jedan ili više dodatnih argumenata.
1. Funkcija `destroyer` treba da vrati novi niz koji isključuje sve elemente iz prvog argumenta koji odgovaraju bilo kom od naknadnih argumenata.
1. Funkcija mora prihvatiti neodređeni broj argumenata.

# --hints--

You should have a `destroyer` function.

```js
assert.isFunction(destroyer);
```

`destroyer([1, 2, 3, 1, 2, 3], 2, 3)` should return `[1, 1]`.

```js
assert.deepEqual(destroyer([1, 2, 3, 1, 2, 3], 2, 3), [1, 1]);
```

`destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3)` should return `[1, 5, 1]`.

```js
assert.deepEqual(destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3), [1, 5, 1]);
```

`destroyer([3, 5, 1, 2, 2], 2, 3, 5)` should return `[1]`.

```js
assert.deepEqual(destroyer([3, 5, 1, 2, 2], 2, 3, 5), [1]);
```

`destroyer([2, 3, 2, 3], 2, 3)` should return `[]`.

```js
assert.deepEqual(destroyer([2, 3, 2, 3], 2, 3), []);
```

`destroyer(["tree", "hamburger", 53], "tree", 53)` should return `["hamburger"]`.

```js
assert.deepEqual(destroyer(["tree", "hamburger", 53], "tree", 53), ["hamburger"]);
```

`destroyer( ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"], "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan" )` should return `[12, 92, 65]`.

```js
assert.deepEqual(
  destroyer(
    ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"],
    "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan"
  ),
  [12, 92, 65]
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function destroyer(arr, ...valsToRemove) {
  return arr.filter(item => !valsToRemove.includes(item));
}
```

