---
id: 6852f06fdde1b4634a263f03
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Sada testiraj klasu `Email` kreiranjem objekta e-pošte i proveravanjem da li pravilno skladišti informacije. Isprižaće nekoliko atributa kako biste potvrdili da sve radi.

Kreirajte objekat e-pošte pod imenom `email_obj` sa `alice@example.com` kao pošiljaocem, `bob@example.com` kao primaocem, `Hello` kao predmetom i `Hi Bob!` kao telom. Zatim ispišite atribute `sender` i `subject` kao zasebne naredbe za štampanje kako biste potvrdili da su pravilno sačuvani.

# --hints--

You should create an email object named `email_obj` with the specified parameters.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_variable("email_obj").is_equivalent("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`))
    }

})
```

You should print `email_obj.sender`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.sender)")`))
    }

})
```

You should print `email_obj.subject`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.subject)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--

--fcc-editable-region--
```
