---
id: 6852f11d412c29665985a915
title: Korak 12
challengeType: 20
dashedName: step-12
---

# --description--

Sada kada korisnici mogu slati e-mailove, potreban im je način za skladištenje primljenih e-mailova. Svaki korisnik bi trebalo da ima inbox za prikupljanje svojih e-mailova.

Dodajte atribut `inbox` u metodu `__init__` klase `User` i inicijalizujte ga kao prazan list `[]`. Ovo će skladištiti sve e-mailove koje korisnik primi.

# --hints--

You should add `self.inbox = []` to the `User` class `__init__` method.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
```
