---
id: 6852f0c3fcedb964ab959921
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Sada proverite da je atribut `read` pravilno dodan u vašu klasu `Email`. Pošto već imate objekat e-pošte iz prethodnih koraka, štampajte atribut `read` kako biste videli da je sada podrazumevano `False`.

# --hints--

You should print `email_obj.read` to test the read attribute.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)

--fcc-editable-region--
```
