---
id: 6852ea0c3cadb54e897acdac
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovoj radionici, kreći ćete sa izgradnjom *Simulatora e-pošte* koji simulira slanje, prijem i upravljanje e-porukama između različitih korisnika. Naučićete o klasama, objektima i kako organizovati kod na objektnom nivou.

Počnite kreiranjem klase pod nazivom `Email` koristeći ključnu reč `class`.

# --hints--

You should create a class named `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`(_Node(_code).has_class("Email"))`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
