---
id: 6852ef95ccc27660275692cd
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Vaša `Email` klasa mora da skladišti informacije o svakom e-mailu kada se kreira.

Unutar vaše `Email` klase, uklonite postojeću ključnu reč `pass` i kreirajte metod `__init__` koji prima `self`, `sender`, `receiver`, `subject` i `body` kao parametre.

# --hints--

You should define an `__init__` method inside the `Email` class.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__init__")`))
    }

})
```

Your `__init__` method should have parameters: `self`, `sender`, `receiver`, `subject`, and `body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").has_args("self, sender, receiver, subject, body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    pass
--fcc-editable-region--
```
