---
id: 6859400f8f7f872e50e1238c
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Klasa ``User`` treba metod za slanje e-pošte drugim korisnicima. Kada šaljete e-mail, kreiraćete novi objekat ``Email`` i dodati ga u prijemnikovu inbox.

Kreirajte metod nazvan ``send_email`` u vašoj klasi ``User``. Ovaj metod treba da prima parametre za ``receiver``, ``subject`` i ``body``.

Unutar metoda, kreirajte novi objekat ``Email`` sa sledećim vrednostima parametara i dodelite ga promenljivoj nazvanoj ``email``:

- ``sender``: ``self`` (trenutni korisnik)
- ``receiver``: ``receiver`` (korisnik koji prima e-mail)
- ``subject``: ``subject`` (naslov e-maila)
- ``body``: ``body`` (sadržaj e-maila)

# --hints--

You should define a `send_email` method in the `User` class.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").has_function("send_email")`))
    }

})
```

The `send_email` method should have parameters: `self`, `receiver`, `subject`, and `body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_args("self, receiver, subject, body")`))
    }

})
```

The method should create an `Email` object named `email` with the correct parameters.

```js
({
  test: () => {
    assert(runPython(`
cond = _Node(_code).find_class("User").find_function("send_email").find_body()[0]
cond.is_equivalent("email = Email(sender=self, receiver=receiver, subject=subject, body=body)") or cond.is_equivalent("email = Email(self, receiver, subject, body)")
    `));
  }
});
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
```
