---
id: 685c1e32254e5993a5c8823e
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Sada kada imate početnu konfiguraciju klase `Email`, vreme je da kreirate korisnike koji mogu slati i primati e-mailove. Svaki korisnik imaće ime i biće u stanju da obavlja e-mail operacije.

Kreirajte klasu `User` sa metodom `__init__` koja prima `self` i `name` kao parametre. Unutar klase, dodelite parametar `name` na `self.name`.

# --hints--

You should create a `User` class.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_class("User")`))
    }

})
```

The `User` class should have an `__init__` method that takes `self` and `name` as parameters.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").has_args("self, name")`))
    }

})
```

You should assign the `name` parameter to `self.name`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.name").is_equivalent("self.name = name")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--

--fcc-editable-region--
```
