---
id: 685c1eba966c6396850736ea
title: Korak 13
challengeType: 20
dashedName: step-13
---

# --description--

Iako korisnici mogu slati e-mailove i imati pošte, potreban je posvećen klasa za efikasno upravljanje operacijama u pošti.

Klasa `Inbox` će čuvati listu e-mailova i pružiti metode za dodavanje novih e-mailova, listanje svih e-mailova i upravljanje njima.

Kreirajte novu klasu nazvanu `Inbox` sa metodom `__init__` koja inicijalizuje praznu listu nazvanu `emails`.

# --hints--

You should create an `Inbox` class.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_class("Inbox")`));
    }
})
```

The `Inbox` class should have an `__init__` method that takes `self` as a parameter.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").has_args("self")`));
    }
})
```

You should initialize `self.emails = []` in the `Inbox` class `__init__` method.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").find_variable("self.emails").is_equivalent("self.emails = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--

--fcc-editable-region--
```
