---
id: 688754d6d59578df1313cb3c
title: Korak 14
challengeType: 20
dashedName: step-14
---

# --description--

Sada je vreme da povežemo klase User i Inbox kako bi se e-poruke zaista mogle isporučiti! 

Korisnici moraju imati odgovarajuće Inbox objekte umesto jednostavnih lista, a metoda `send_email` bi trebalo da isporuči e-poruke u inbox primaoca.

Ažurirajte klasu `User` `__init__` na način da svaki korisnik dobije stvarni Inbox objekat umesto prazne liste kreirane ranije.

# --hints--

You should update the User class to initialize `self.inbox = Inbox()`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = Inbox()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []
--fcc-editable-region--

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

class Inbox:
    def __init__(self):
        self.emails = []

```
