---
id: 68b9366fae8b13c1278761b1
title: Korak 15
challengeType: 20
dashedName: step-15
---

# --description--

Vaš inbox zahteva način za primanje novih e-poruka. Kada neko pošalje e-poruku korisniku, ona bi trebalo dodati u njegov inbox.

Dodajte metod nazvan `receive_email` u vašu klasu `Inbox` koja prima `self` i objekat e-poruke `email` kao parametre. Unutar tela metode, dodajte `email` u listu `emails` koristeći metodu `append`.

# --hints--

You should define a method named `receive_email` in the `Inbox` class that takes `self` and `email` as parameters.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_args("self, email")`))
    }
})
```

You should append the `email` to `self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_stmt("self.emails.append(email)")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--
class Inbox:
    def __init__(self):
        self.emails = []

    
--fcc-editable-region--
```
