---
id: 68b930ee0e7b8fb81895e59a
title: Korak 7
challengeType: 20
dashedName: step-7
---

# --description--

Sada ćete dodati metodu za označavanje e-pošte kao pročitane. Kada neko otvori e-mail, želite da promenite njegov status sa nepročitanog na pročitano. Ova metoda će biti jednostavna - samo treba da postavi atribut `read` na `True`.

Dodajte metodu nazvanu `mark_as_read` u vašu `Email` klasu. Ova metoda bi trebalo da primi samo `self` kao parametar i postavi atribut `read` na `True`.

# --hints--

You should define a method named `mark_as_read` in the `Email` class.
that should only take `self` as a parameter.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").has_args("self")`))
    }

})
```

The method should set `self.read` to `True`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").find_variable("self.read").is_equivalent("self.read = True")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
```
