---
id: 68e97fe79367ad7b5dd6c9cd
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

U ovom koraku dodajte element `h1` sa tekstom `XYZ Bookstore`.

# --hints--

You should have an `h1` element.

```js
assert.exists(document.querySelector('h1'));
```

Your `h1` element's text should be `XYZ Bookstore`. Double check for spelling or capitalization errors.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

Your `h1` element should be inside the `body` element.

```js
assert.match(code, /<body>[\s\S]*<h1>[\s\S]*<\/h1>[\s\S]*<\/body>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
