---
id: 68ec9332a9b5b2b32487bd00
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Atribut `id` dodaje jedinstveni identifikator HTML elementu. Svaki `id` treba da bude jedinstven unutar stranice i trebalo bi ga koristiti samo jednom.

Vrednosti `id` ne smeju da sadrže razmake i trebalo bi da sadrže samo slova, cifre, donje linije (underscores) i crtice (dashes).

Evo primera:

```html
<p id="para">example paragraph</p>
```

Dodajte atribut `id` vašem elementu koji ima klasu `card` i postavite mu vrednost na `sally-adventure-book`.

# --hints--

Your element with a class of `card` should have an `id` attribute.

```js
assert.isTrue(document.querySelector('.card')?.hasAttribute('id'));
```

Your element having a class of `card` should have an `id` having the value of `sally-adventure-book`.

```js
assert.equal(document.querySelector('.card')?.id, 'sally-adventure-book');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
  --fcc-editable-region--
    <div class="card">
  --fcc-editable-region--
      
    </div>
  </div>
</body>

</html>
```
