---
id: 68eca2a795b333ca5fee30a8
title: Korak 13
challengeType: 0
dashedName: step-13
---

# --description--

Sada kreiraj drugu karticu za knjigu. Dodaj još jedan ``div`` element sa atributom ``class`` postavljenim na ``card``. Primećujte kako možete ponovo koristiti isto ime klase za više elemenata kako biste primenili dosledno stilizovanje.

# --hints--

You should have two `div` elements with a `class` attribute having a value of `card`.

```js
assert.lengthOf(document.querySelectorAll('.card'), 2);
```

The second element with a class of `card` should be inside the element having a class of `card-container`.

```js
assert.lengthOf(document.querySelectorAll('.card-container .card'), 2);
```

The two elements with a class of `card` should be siblings.

```js
assert.exists(document.querySelector('.card + .card'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
  --fcc-editable-region--
    
  --fcc-editable-region--
  </div>
</body>

</html>
```
