---
id: 68ec99e478211dc578699944
title: Korak 11
challengeType: 0
dashedName: step-11
---

# --description--

Ispod elementa `h2` u prvom elementu koji ima klasu `card`, dodajte element `p` sa sledećim tekstom:

```md
This is an epic story of Sally and her dog Rex as they navigate through other worlds.
```

# --hints--

You should have a `p` element nested inside the element having a class of `card`.

```js
assert.exists(document.querySelector('.card p'));
```

Your `p` element's text should be `This is an epic story of Sally and her dog Rex as they navigate through other worlds.`

```js
assert.equal(document.querySelector('.card p')?.innerText.trim(), "This is an epic story of Sally and her dog Rex as they navigate through other worlds.");
```

Your `p` element should be below your `h2` element.

```js
assert.exists(document.querySelector('.card h2 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
