---
id: 68ec6d9a315221aa31e54816
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Atribut `class` se koristi za identifikovanje jednog ili više elemenata radi stilizovanja. Za razliku od atributa `id`, imena klasa ne moraju biti jedinstvena: više elemenata može deliti istu klasu.

Evo primera:

```html
<p class="example">example paragraph</p>
```

Dodajte atribut `class` vašem `div` elementu i postavite mu vrednost na `card-container`.

# --hints--

Your `div` element should have a `class` attribute.

```js
assert.isTrue(document.querySelector('div')?.hasAttribute('class'));
```

Your `div` element should have a `class` attribute with the value of `card-container`.

```js
assert.equal(document.querySelector('div')?.className, 'card-container');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  <div>
--fcc-editable-region--
    
  </div>
</body>

</html>
```
