---
id: 68ec98b38d83a3c28dd30efe
title: Korak 10
challengeType: 0
dashedName: step-10
---

# --description--

Unutar prvog elementa koji ima klasu `card`, dodaj element `h2` sa tekstom `Sally's SciFi Adventure`.

# --hints--

You should have an `h2` element nested inside the element having a class of `card`.

```js
assert.exists(document.querySelector('.card h2'));
```

Your `h2` element's text should be `Sally's SciFi Adventure`. 

```js
assert.equal(document.querySelector('.card h2')?.innerText.trim(), "Sally's SciFi Adventure");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
    --fcc-editable-region--
      
    --fcc-editable-region--
    </div>
  </div>
</body>

</html>
```
