---
id: 68eca6364c5616d393389a30
title: Korak 18
challengeType: 0
dashedName: step-18
---

# --description--

Zapamtite, HTML element izgleda ovako:

```html
<element attribute="value">
    inner text
</element>
```

Ispod elementa sa klasom `card-container`, dodajte novi `p` element sa ovim tekstom:

```md
Review your selections and continue to checkout.
```

Ispod `p` elementa, kreirajte `div` element sa atributom `class` postavljenim na `btn-container`. Ovaj kontejner će grupisati vaše navigacione dugme elemente.

# --hints--

The newly added `p` element should be below the element with a class of `card-container`.

```js
assert.exists(document.querySelector('.card-container + p'));
```

Your new `p` element's text should be `Review your selections and continue to checkout.`

```js
const cardContainer = document.querySelector('.card-container');
assert.equal(cardContainer?.nextElementSibling?.innerText.trim(), 'Review your selections and continue to checkout.');
```

You should have a `div` element with a `class` of `btn-container`.

```js
assert.exists(document.querySelector('div.btn-container'));
```

Your `div` element having a `class` of `btn-container` should come after the newly added `p` element.

```js
assert.exists(document.querySelector('p + .btn-container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
  <div class="card-container">
    <div class="card" id="sally-adventure-book">
      <h2>Sally's SciFi Adventure</h2>
      <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
      <button class="btn">Buy Now</button>
    </div>
    
    <div class="card" id="dave-cooking-book">
      <h2>Dave's Cooking Adventure</h2>
      <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
      <button class="btn">Buy Now</button>
    </div>
  </div>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
