---
id: 68ec6cd5b7e8f5a8f7319e32
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Element ``div`` se koristi kao kontejner za grupisovanje drugih HTML elemenata. Glavnije ćete koristiti element ``div`` kada želite da grupišete HTML elemente koji će deliti određeni skup CSS stilova.

Ispod elementa ``p``, dodajte element ``div``. Ovaj ``div`` će biti kontejner za vaše kartice sa knjigama.

**Napomena**: Ova radionica ne primenjuje CSS. Klase i grupisan elementi korisni su za CSS stilizovanje, ali u ovoj radionici se koriste samo za strukturiranje i grupisovanje sadržaja. Naučićete kako funkcioniše stilizovanje u kasnijem modulu.

# --hints--

You should have a `div` element.

```js
assert.exists(document.querySelector('div'));
```

Your `div` element should be below your `p` element.

```js
assert.exists(document.querySelector('p + div'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
