---
id: 69134b6cb43aa5f254950a10
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Na ovom radionici, kreiraćete stranicu knjižare pravljenjem kartica knjiga koje prikazuju informacije o različitim knjigama. Vežvaćete organizovanje sadržaja korišćenjem elemenata `div`, klasa i ID-jeva.

Počnite sa stranicom knjižare kreiranjem HTML šablona.

Dodajte deklaraciju `<!DOCTYPE html>` i elemente `html` i `head`.

Dodajte atribut `lang` na element `html` i postavite ga na `"en"`.

# --hints--

Your code should include a `<!DOCTYPE html>` declaration.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Your code should include both opening and closing `<html>` tags.

```js
assert.match(code,/<html\b[^>]*>[\s\S]*<\/html>/i);
```

Your `html` element should have a `lang` attribute set to `"en"`.

```js
assert.match(code, /<html[^>]*\blang\s*=\s*["']en["'][^>]*>/i);
```

Your code should include both opening and closing `<head>` tags.

```js
assert.match(code,/<head>[\s\S]*<\/head>/i);
```

Your `head` element should be nested inside `html` element.

```js
assert.match(code, /<\/html>\s*$/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
