---
id: 69134d14150d51f8066a3246
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Sada, poboljšajte strukturu vašeg HTML dokumenta kako biste osigurali da je stranica pravilno kodirana.

Unutar elementa `head`, dodajte element `<meta charset="UTF-8">`.

Napokon, dodajte element `body` ispod sekcije `head`. Ovde će se nalaziti sav vidljivi sadržaj vaše stranice.

# --hints--

Your code should include a `<meta charset="UTF-8">` element.

```js
assert.exists(document.querySelector('head > meta[charset="utf-8"]'));
```

Your code should include opening and closing `<body>` tags.

```js
assert.match(code,/<body>[\s\S]*<\/body>/i);
```

Your `body` element should be below `head` element.

```js
assert.match(code, /<\/head>[\s\S]*<body>[\s\S]*<\/body>[\s\S]*<\/html>/i);
```

# --seed--
## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
  <title>XYZ Bookstore Page</title>
</head>

--fcc-editable-region--
</html>
```
