---
id: 699b6dd1a2961bd38491ef2e
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

Objekat `profile` je imao stroge ključeve i tipove. Ako ste pokušali dodati svojstvo koje nije definisano ili dodeliti nekompatibilan tip, TypeScript bi izbacio grešku.

Bilo bi lepo da objekat `userRoles` nije tako striktan i imao fleksibilnije ključeve. 

Kao što se sećate iz prethodnih lekcija, ovaj rezultat možete postići korišćenjem tipa `Record`:

```ts
const studentGrades: Record<string, number> = {
  Math: 95,
  English: 88,
  Science: 92
};
```

U ovom primeru sa `studentGrades`, TypeScript ne mora da zna sve predmete unapred. To vam omogućava da kasnije dodate više predmeta kao što su `Art` ili `Music`. TypeScript ga samo brine da vrednosti za svaki ključ budu brojevi. 

Ažurirajte svoj objekat `userRoles` dodavanjem inline tipa `Record`. Postavite ga na `<string, string>` tako da svaki ključ može biti bilo koji string, a svaka vrednost mora biti string.

# --hints--

Your `userRoles` object should have an inline `Record` type set to `Record<string, string>`.

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.annotation.matches("Record<string, string>"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles = {
--fcc-editable-region--
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

console.log(userRoles);
```
