---
id: 699b61ca8c9cb9bc1968364e
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Ako otvorite konzolu, trebalo bi da vidite sledeću poruku o grešci:

```md
Object literal may only specify known properties, and 'mood' does not exist in type '{ username: string; age: number; isLoggedIn: boolean; }'.
```

U čistom JavaScriptu (vanilla JavaScript), dozvoljeno vam je dodavanje bilo kojih svojstava i vrednosti koje ste želeli objektu `profile`. Ali u TypeScript-u, kompajler proverava da li objekat ima samo svojstva koja ste definisali (`username`, `age` i `isLoggedIn` u ovom slučaju). Pošto `mood` nije bio uključen u tip, TypeScript ga tretira kao nevažeće svojstvo i baca grešku.

Uklonite `mood: null` iz vašeg koda.

# --hints--

You should no longer have `mood: null` in your code.

```js
assert.isUndefined(profile.mood);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
--fcc-editable-region--
  mood: null
--fcc-editable-region--
};

console.log(profile);
```
