---
id: 699b71d98c820ac7ab9eea2a
title: Korak 13
challengeType: 1
dashedName: step-13
---

# --description--

Ako ponovo otvorite konzolu, videćete novu poruku o grešci:

```md
Type 'number' is not assignable to type 'string'.
```

Ovo se dešava jer pokušavate da dodate svojstvo ``guest`` sa numeričkom vrednošću. 

Da biste uklonili grešku TypeScript-a, promenite vrednost svojstva ``guest`` u ``"read-only"``.

Sa tom poslednjom promenom, vaš radionica za korisnički profil je gotova!

# --hints--

Your `guest` property for the `userRoles` object should have the value of `"read-only"`.

```js
assert.propertyVal(userRoles, "guest", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
--fcc-editable-region--
  guest: 3
--fcc-editable-region--
};

console.log(userRoles);
```

# --solutions--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
  guest: "read-only"
};

console.log(userRoles);
```
