---
id: 699b65caf17a8e8d340af27d
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Sada je vreme da dodamo poslednje svojstvo objektu `profile`.

U prethodnim lekcijama, naučili ste kako da radite sa tipovima nizova (array). Naučili ste o različitim tipovima nizova kao što su `string[]`, ili `number[]`:

```ts
const scores: number[] = [95, 87, 100];
```

U ovom primeru, niz `scores` prihvata samo brojeve. Ako pokušate da dodate drugi tip kao što je `string`, ili `null`, TypeScript će baciti grešku.

Ažurirajte svoj objekt `profile` da dodate svojstvo `programmingLanguages`, tipizovano kao `string[]`. Unutar objektnog literala, postavite ovo svojstvo na `["JavaScript", "Python", "C++"]`.

# --hints--

You should have a `programmingLanguages` property inside of your `profile` object.

```js
assert.property(profile, "programmingLanguages");
```

Your `programmingLanguages` property should have the value of `["JavaScript", "Python", "C++"]`.

```js
assert.deepEqual(profile.programmingLanguages, ["JavaScript", "Python", "C++"]);
```

Your `programmingLanguages` property should be typed as a `string[]`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "programmingLanguages", type: "string[]" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
--fcc-editable-region--
  
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
  
--fcc-editable-region--
};

console.log(profile);
```
