---
id: 699b68f94a0d154c2a7a5ce0
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Sada kada je objekat `profile` završen, vreme je da napravimo novi objekat koji predstavlja uloge korisnika.

Kreirajte novu varijablu nazvanu `userRoles` i dodelite joj objekat. Vaš objekat treba da ima sledeće svojstva i vrednosti:

```md
admin: "full-access",
editor: "limited-access",
viewer: "read-only"
```

# --hints--

You should have a variable called `userRoles`.

```js
assert.isDefined(userRoles);
```

Your `userRoles` variable should be an object.

```js
assert.isObject(userRoles);
```

Your `userRoles` object should have an `admin` property.

```js
assert.property(userRoles, "admin");
```

Your `admin` property should have a value of `"full-access"`.

```js
assert.propertyVal(userRoles, "admin", "full-access");
```

Your `userRoles` object should have an `editor` property.

```js
assert.property(userRoles, "editor");
```

Your `editor` property should have a value of `"limited-access"`.

```js
assert.propertyVal(userRoles, "editor", "limited-access");
```

Your `userRoles` object should have a `viewer` property.

```js
assert.property(userRoles, "viewer");
```

Your `viewer` property should have a value of `"read-only"`.

```js
assert.propertyVal(userRoles, "viewer", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--

--fcc-editable-region--
```
