---
id: 699b64096301fd75e3efeb1b
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Ako ponovo otvorite konzolu, videćete druga poruku o grešci:

```md
Property 'bio' is missing in type '{ username: string; age: number; isLoggedIn: false; }' but required in type '{ username: string; age: number; isLoggedIn: boolean; bio: string; }'.
```

Kao što ste se setili u prethodnim lekcijama, možete označiti svojstva kao opcionalna koristeći simbol `?` pored imena svojstva na ovaj način:

```ts
const musician: {
  name: string;
  instrument: string;
  isActive: boolean;
  numberOfAlbums?: number;
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

U ovom primeru, `numberOfAlbums` je označen kao opcionalno kako bi TypeScript u ovom slučaju ne izbacivao grešku. 

Ažurirajte svojstvo `bio` tako da bude označeno kao opcionalno.

# --hints--

Your `bio` property should be marked as optional.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string", isOptional: true };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
--fcc-editable-region--
  bio: string
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
