---
id: 699b5b7cb68b6d35afca61b7
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Iako je moguće nastaviti dodavati svojstva poput ovih objektu `profile`, bilo bi prijatno ograničiti tipove svojstava koja bi trebalo biti dozvoljena. Bolje je unapred definisati oblik objekta `profile`.

U prethodnim lekcijama, naučili ste kako da radite sa tipovima objekata poput ovih:

```ts
const composer: { 
  name: string; 
  instrument: string; 
  isActive: boolean; 
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

U ovom primeru, `composer` je objekat sa ugrađenim tipom objekta. Ovaj objekat ima tri obavezna svojstva: `name`, `instrument` i `isActive`. Prva dva svojstva moraju biti tipa `string`, a `isActive` mora biti tipa `boolean`.

Ažurirajte svoj postojeći objekat `profile` da uključite ugrađeni tip objekta. `username` bi trebalo da bude tipizovan na `string`, `age` bi trebalo da bude tipizovan na `number`, a `isLoggedIn` bi trebalo da bude tipizovan na `boolean`.

# --hints--

Your `profile` object should have an inline object type with a `username` property set to `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "username", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Your `profile` object should have an inline object type with an `age` property set to `number`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "age", type: "number" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

Your `profile` object should have an inline object type with an `isLoggedIn` property set to `boolean`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "isLoggedIn", type: "boolean" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile = {
--fcc-editable-region--
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  mood: null
}

console.log(profile);
```
