---
id: 699b4e284291adcfcb90df47
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, vežbaćete rad sa primitivnim tipovima, kao i sa objektima i nizovima, kreiranjem sigurnog korisničkog profila.

Počnite kreiranjem varijable nazvane `profile` i dodelite joj objekat. Vaš objekat treba da ima ova tri svojstva i vrednosti:

```md
username: "codeLearner",
age: 25,
isLoggedIn: false,
```

# --hints--

You should have a variable called `profile`.

```js
assert.exists(profile);
```

Your `profile` variable should be an object.

```js
assert.isObject(profile);
```

You should have a `username` property in your `profile` object.

```js
assert.property(profile, "username");
```

Your `username` property should have a value of `"codeLearner"`.

```js
assert.propertyVal(profile, "username", "codeLearner");
```

Your `profile` object should have an `age` property.

```js
assert.property(profile, "age");
```

Your `age` property should have a value of `25`.

```js
assert.propertyVal(profile, "age", 25);
```

Your `profile` object should have an `isLoggedIn` property.

```js
assert.property(profile, "isLoggedIn");
```

Your `isLoggedIn` property should have a value of `false`. 

```js
assert.propertyVal(profile, "isLoggedIn", false);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```
