---
id: 699b62c99ec8c63470313365
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Sada je vreme da dodate još jednu svojstvo vašem objektu `profile`.

Počnite ažuriranjem tipa ugrađenog objekta da uključite svojstvo `bio` sa tipom postavljenim na `string`.

# --hints--

You should update the inline object type to include a `bio` property set to `string`.

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
--fcc-editable-region--
  
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
