---
id: 68eb7f2397da251a40e04aa2
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Sledeći korak je dodati element ``div`` da sadrži sve elemente za kontrolu jačine zvuka.

Unutar elementa ``div``, ugnježđujte element ``span`` sa atributom ``id`` postavljenim na ``volume-label``.
Takođe, tekst za element ``span`` treba da bude ``Volume``.

Element ``id`` će kasnije biti referenciran putem ARIA atributa kako bi asistivne tehnologije mogle da identifikuju ovaj tekst kao oznaku za kontrolu jačine zvuka.

# --hints--

You should have a `div` element below your `button` element.

```js
assert.exists(document.querySelector('button + div'));
```

You should have a `span` element inside your `div` element.

```js
assert.exists(document.querySelector('button + div > span'));
```

Your `span` element should have an `id` attribute.

```js
const spanEl = document.querySelector('button + div > span');
assert.isTrue(spanEl?.hasAttribute('id'));
```

Your `span` element should have an `id` attribute with the value set to `volume-label`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.getAttribute('id'), 'volume-label');
```

Your `span` element should have the text `Volume`.

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.innerText.trim(), 'Volume');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

--fcc-editable-region--

--fcc-editable-region--

  </body>
</html>
```
