---
id: 68ec0a48c9d6db3594d768c3
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Sada je vreme da dodamo još jedan `span` element.

Ovaj `span` element treba da ima atribut `id` postavljen na `volume-description`.
Takođe, tekst za ovaj `span` element treba da bude `Adjust the sound level`.

Slično prethodnom `span` elementu, atribut `id` ovog `span` će kasnije biti referenciran putem ARIA atributa kako bi asistivne tehnologije mogle da identifikuju ovaj tekst kao opis za kontrolu jačine zvuka.

# --hints--

You should have two `span` elements inside the `div` element.

```js
const spans = document.querySelectorAll('div > span');
assert.lengthOf(spans, 2);
```

Your second `span` element should have an `id` attribute.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.isTrue(secondSpanEl?.hasAttribute('id'));
```

Your second `span` element should have an `id` attribute with the value of `volume-description`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.getAttribute('id'), 'volume-description');
```

Your second `span` element should have the text `Adjust the sound level`.

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.innerText.trim(), 'Adjust the sound level');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      --fcc-editable-region--
    
      --fcc-editable-region--
    </div>

  </body>
</html>
```
