---
id: a5deed1811a43193f9f1c841
title: Implementujte Element Skippera
challengeType: 26
dashedName: implement-an-element-skipper
---

# --description--

U ovom laboratorijumu ćete kreirati funkciju koja preskače elemente u nizu dok ne pronađe prihvatljiv element na osnovu specifične test funkcije.

Na primer, za niz poput `[1, 1, 1, 2, 1, 1, 1]` i test funkciju `function(n) {return n === 2}`, prvi element koji je prihvatljiv za ovo je onaj na indeksu 3, pa svi elementi pre toga moraju biti odbačeni, a izlaz bi trebalo da budu preostali elementi `[2, 1, 1, 1]`.

**Cilj**: Ispunite priče korisnika ispod i učinite da svi testovi prođu kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da imate funkciju `dropElements` koja prima niz (`arr`) i funkciju (`func`) kao argumente.
1. Funkcija `dropElements` bi trebalo da iterira kroz niz i ukloni elemente počevši od prvog sve dok `func` ne vrati `true` za element.
1. Funkcija `dropElements` bi trebalo da vrati preostale elemente u nizu ako je uslov zadovoljen.
1. Ako uslov nikada nije zadovoljen, treba da vrati prazan niz.

# --hints--

You should have a `dropElements` function.

```js
assert.isFunction(dropElements);
```

`dropElements([1, 2, 3, 4], function(n) {return n >= 3;})` should return `[3, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n >= 3;}), [3, 4]);
```

`dropElements([0, 1, 0, 1], function(n) {return n === 1;})` should return `[1, 0, 1]`.

```js
assert.deepEqual(dropElements([0, 1, 0, 1], function(n) {return n === 1;}), [1, 0, 1]);
```

`dropElements([1, 2, 3], function(n) {return n > 0;})` should return `[1, 2, 3]`.

```js
assert.deepEqual(dropElements([1, 2, 3], function(n) {return n > 0;}), [1, 2, 3]);
```

`dropElements([1, 2, 3, 4], function(n) {return n > 5;})` should return `[]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n > 5;}), []);
```

`dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;})` should return `[7, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;}), [7, 4]);
```

`dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;})` should return `[3, 9, 2]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;}), [3, 9, 2]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dropElements(arr, func) {
  while (arr.length > 0 && !func(arr[0])) {
    arr.shift();
  }
  return arr;
}
```
