---
id: a3566b1109230028080c9345
title: Dizajnirajte algoritam za sabiranje svih brojeva
challengeType: 26
dashedName: lab-sum-all-numbers-algorithm
---

# --description--

U ovom laboratorijumu, moraćete da dizajnirate algoritam za sabiranje svih brojeva.

**Cilj:** Ispunite priče korisnika ispod i osigurajte da svi testovi prođu kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da imate funkciju pod nazivom `sumAll` koja prima niz od dva broja.
1. `sumAll([n, m])` bi trebalo da vrati zbir `n` i `m` plus zbir svih brojeva između njih. Najmanji broj neće uvek biti prvi. Na primer, `sumAll([4,1])` bi trebalo da vrati `10` jer je zbir svih brojeva između `1` i `4` (oba uključujući) `10`.

# --hints--

You should have a function named `sumAll`.

```js
assert.isFunction(sumAll);
```

`sumAll([1, 4])` should return a number.

```js
assert.isNumber(sumAll([1, 4]));
```

`sumAll([1, 4])` should return `10`.

```js
assert.equal(sumAll([1, 4]), 10);
```

`sumAll([4, 1])` should return `10`.

```js
assert.equal(sumAll([4, 1]), 10);
```

`sumAll([5, 10])` should return `45`.

```js
assert.equal(sumAll([5, 10]), 45);
```

`sumAll([10, 5])` should return `45`.

```js
assert.equal(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function sumAll(arr) {
  let sum = 0;
  
  arr.sort((a, b) => a - b);

  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }

  return sum;
}
```
