id: 587d825b367417b2b2512c8b
title: Ukloni element iz maksimalnog kopca
challengeType: 1
forumTopicId: 301710
dashedName: remove-an-element-from-a-max-heap
---

# --description--

Sada kada možemo dodavati elemente u naš kopa, istražimo kako da uklanjamo elemente. Uklanjanje i umetanje svih elemenata zahteva sličnu logiku. U maksimalnom kopcu, obično želite da uklonite najveću vrednost, pa se ovo svodi samo na uzimanje iz korena našeg stabla. Ovo će narušiti svojstvo kopca našeg stabla, tako da ga moramo nekako popraviti. Obično, za maksimalni kopac, to se radi na sledeći način:

<ol>
  <li>Premestite poslednji element u kopa na poziciju korena.</li>
  <li>Ako bilo koji od sinova korena ima veću vrednost od samog korena, zamijenite koren sinom sa najvećom vrijednošću.</li>
  <li>Nastavite sa zamjenama dok roditelj ne bude veći od svih sinova ili dođete do najdublje grane u stablu.</li>
</ol>

# --instructions--

Instrukcije: Dodajte metodu na naš maksimalni kopac pod nazivom `remove`. Ova metoda treba da vrati najveću vrednost koja je uklonjena iz našeg maksimalnog kopca i da je ukloni iz kopa. Takođe bi trebalo da reorganizuje kopa kako bi svojstvo kopca ostalo očuvano. Nakon uklanjanja elementa, drugi najveći element koji preostaje u kopcu treba da bude koren.

# --hints--

Klasa `MaxHeap` mora biti definisana.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MaxHeap` mora imati metodu pod nazivom `print`.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

`MaxHeap` mora imati metodu pod nazivom `insert`.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MaxHeap` mora imati metodu pod nazivom `remove`.

```js
assert(
  (function () {
    let test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

Metoda `remove` mora ukloniti najveći element iz maksimalnog kopca, zadržavajući svojstvo maksimalnog kopca.

```js
function isHeap(arr, i, n) {
    if( arr[i] < arr[2 * i + 1] || arr[i] < arr[2 * i + 2] ){
        return false;
    }
    if (i > (n - 1) / 2) {
        return true;
    }
    if (isHeap(arr, 2 * i + 1, n) && isHeap(arr, 2 * i + 2, n)) {
        return true;
    }
    return false;
}


assert(
  (function () {
    let test = false;

    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    
    let max = Infinity;
    const [result, vals] = [[], [9, 3, 5, 2, 15, 3, 7, 12, 7, 10, 90]];

    vals.forEach((val) => test.insert(val));

    for (let i = 0; i < vals.length; i++) {
      const curHeap = test.print();
      const arr = curHeap[0] === null ? curHeap.slice(1) : curHeap;
      
      if (!isHeap(arr, 0, arr.length - 1)) {
        return false;
      }

      const removed = test.remove();
      if (!vals.includes(removed)) return false;
      if (removed > max) return false
      max = removed;
      result.push(removed);
    }
    
    for (let i = 0; i < vals.length; i++) {
      if (!result.includes(vals[i])) {
        return false;
      }
    }
    
    return true;
  })()
);
```

# --seed--

## --seed-contents--

```js
const MaxHeap = function () {
  this.heap = [];
  this.parent = index => {
    return Math.floor((index - 1) / 2);
  }
  this.insert = element => {
    this.heap.push(element);
    this.heapifyUp(this.heap.length - 1);
  }
  this.heapifyUp = index => {
    let currentIndex = index,
    parentIndex = this.parent(currentIndex);
    while (currentIndex > 0 && this.heap[currentIndex] > this.heap[parentIndex]) {
      this.swap(currentIndex, parentIndex);
      currentIndex = parentIndex;
      parentIndex = this.parent(parentIndex);
    }
  }
  this.swap = (index1, index2) => {
    [this.heap[index1], this.heap[index2]] = [this.heap[index2], this.heap[index1]];
  }
  this.print = () => {
    return this.heap;
  }
  // Only change code below this line

  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```