id: 587d825a367417b2b2512c87
title: Kreiraj dvosmernu listu veznih stavki
challengeType: 1
forumTopicId: 301626
dashedName: create-a-doubly-linked-list
---

# --description--

Sve listove veznih stavki koje smo do sada kreirali su jednostruvene liste veznih stavki. Ovde ćemo kreirati <dfn>dvosmernu listu veznih stavki</dfn>. Kao što ime kaže, čvorovi u dvosmernoj listi veznih stavki imaju reference na sledeći i prethodni čvor u listi.

Ovo nam omogućava da putujemo kroz listu u oba smera, ali zahteva i više memorije jer svaki čvor mora imati dodatnu referencu na prethodni čvor u listi.

# --instructions--

Dostavljena je klasa `Node` i kreiran je naš `DoublyLinkedList`. Moramo implementirati dve metode za našu dvosmernu listu veznih stavki, nazvane `add` i `remove`. Metoda `add` treba da doda dato element u listu, dok metoda `remove` treba da ukloni sve pojavljivanja datog elementa iz liste.

Pažljivo obradite sve uzročne slučajeve kada pišete ove metode, kao što su brisanje prvog ili poslednjeg elementa. Takođe, pokušaj na uklanjanje bilo kog elementa iz prazne liste treba da vrati `null`.

# --hints--

Mora postojati struktura podataka za `DoublyLinkedList`.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return typeof test == 'object';
  })()
);
```

`DoublyLinkedList` mora imati metodu nazvanu `add`.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.add == undefined) {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

`DoublyLinkedList` mora imati metodu nazvanu `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    if (test.remove == undefined) {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

Uklanjanje elementa iz prazne liste treba da vrati `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    return test.remove(100) == null;
  })()
);
```

Metoda `add` mora dodavati stavke u listu.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(5);
    test.add(6);
    test.add(723);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '56723';
  })()
);
```

Svaki čvor mora da prati prethodni čvor.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(50);
    test.add(68);
    test.add(73);
    var result = [];
    var curr = test.tail;
    while (curr) {
      result.push(curr.data);
      curr = curr.prev;
    }
    return result.join('') == '736850';
  })()
);
```

Prvi element treba ukloniti iz liste.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(25);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '3560';
  })()
);
```

Poslednji element treba ukloniti iz liste.

```js
assert(
  (function () {
    var test = false;
    if (typeof DoublyLinkedList !== 'undefined') {
      test = new DoublyLinkedList();
    }
    test.add(25);
    test.add(35);
    test.add(60);
    test.remove(60);
    var result = [];
    var curr = test.head;
    while (curr) {
      result.push(curr.data);
      curr = curr.next;
    }
    return result.join('') == '2535';
  })()
);
```


# --seed--

## --seed-contents--

```js
var Node = function(data, prev) {
  this.data = data;
  this.prev = prev;
  this.next = null;
};
var DoublyLinkedList = function() {
  this.head = null;
  this.tail = null;
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```