id: 587d8251367417b2b2512c62
title: Kreiranje klase za spojen list
challengeType: 1
forumTopicId: 301628
dashedName: create-a-linked-list-class
---

# --description--

Početićemo kreiranjem klase `linked list`. Svaki spojeni lista bi trebalo da počinje sa nekoliko osnovnih svojstava: `head` (prvi element u vašem spisku) i `length` (broj elemenata u vašem spisku). Ponekad ćete videti implementacije spojenih lista koje uključuju `tail` za poslednji element liste, ali za sada ćemo se držati samo ova dva. Svaki put kada dodajemo element na spojeni list, naše svojstvo `length` bi trebalo da se poveća za jedan.

Potreban nam je način za dodavanje elemenata u našu spojenu listu, pa prvu metodu koju želimo da kreiramo je metoda `add`.

Ako je naš spisak prazan, dodavanje elementa na spojeni list je jednostavno: kreiramo taj element kao čvor (node) u klasi `Node`, i postavljamo taj čvor kao `head` naše spojenog liste.

Ali šta ako naš spisak već ima jedan ili više članova? Kako dodajemo element na listu? Zapamtite da svaki čvor u spojenoj listi ima svojstvo `next`. Da bismo dodali čvor na listu, moramo pronaći poslednji čvor u listi, i usmeriti svojstvo `next` tog poslednjeg čvora ka našem novom čvoru. (Savet: znate da ste stigli do kraja spojenog liste kada je svojstvo `next` čvora `null`.)

# --instructions--

Napišite metodu za dodavanje koja postavlja prvi dodati čvor na spojenu listu kao `head`; nakon toga, svaki put kada dodajete čvor, svaki čvor bi trebalo da ima referencu na svojstvo `next` prethodnog čvora.

Napomena

`length` vaše liste bi trebalo da se poveća za jedan pri svakom dodatom elementu u spojenoj listi.

# --hints--

Vaša klasa `LinkedList` bi trebalo da ima metodu `add`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.add === 'function';
  })()
);
```

Vaša klasa `LinkedList` bi trebalo da postavi `head` na prvi dodati čvor.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    return test.head().element === 'cat';
  })()
);
```

Prethodni `node` u vašoj klasi `LinkedList` bi trebalo da ima referencu na najnovije kreirani node.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('fish');
    return test.head().next.element === 'dog' && test.head().next.next.element === 'fish';
  })()
);
```

`size` vaše klase `LinkedList` bi trebalo da odgovara broju čvorova u spojenoj listi.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.size() === 2;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() { 
  var length = 0; 
  var head = null; 

  var Node = function(element){
    this.element = element; 
    this.next = null; 
  }; 

  this.head = function(){
    return head;
  };

  this.size = function(){
    return length;
  };

  this.add = function(element){
    // Only change code below this line
    if (head == null) {
      head = new Node(element);
    } 
    else {
      let currentNode = head;
      while (currentNode.next != null) {
        // currentNode.next will be last node of linked list after loop
        currentNode = currentNode.next;
      }
      currentNode.next = new Node(element);
    }
    length++;
    // Only change code above this line
  };
}
```