id: 587d825b367417b2b2512c8d
title: Kreiraj Map u ES6 JavaScriptu
challengeType: 1
forumTopicId: 301635
dashedName: create-an-es6-javascript-map
---

# --description--

Nova verzija JavaScripta pruža ugrađenu Map strukturu koja nudi veliki deo funkcionalnosti koju smo ručno pisali u prethodnom izazovu. Ova Map struktura, iako slična običnim JavaScript objekatima, nudi ključne funkcije koje obični objekti nemaju. Na primer, ES6 Map prati redosled dodavanja elemenata. Evo kompletnog pregleda metoda:

- `.has(key)` vraća `true` ili `false` u zavisnosti od prisustva ključa
- `.get(key)` vraća vrednost povezana sa ključem
- `.set(key, value)` postavlja novu paru ključ-vrednost
- `.delete(key)` uklanja par ključ-vrednost
- `.clear()` briše sve pare ključ-vrednost
- `.keys()` vraća iterator svih ključeva Map u redosledu dodavanja
- `.values()` vraća iterator svih vrednosti Map u redosledu dodavanja
- `.entries()` vraća iterator svih parova ključ-vrednost Map u redosledu dodavanja

# --instructions--

Definisisać JavaScript Map strukturu i dodeliti je promenljivoj nazvanoj `myMap`. Dodajte par ključ-vrednost `freeCodeCamp` i `Awesome!` unutar nje.

# --hints--

`myMap` struktura mora da postoji.

```js
assert(typeof myMap === 'object');
```

`myMap` mora imati par ključ-vrednost `freeCodeCamp` i `Awesome!`.

```js
assert(myMap.get('freeCodeCamp') === 'Awesome!');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const myMap = new Map();

myMap.set("freeCodeCamp", "Awesome!");
```