id: 587d825d367417b2b2512c96
title: Dubinska pretraga
challengeType: 1
forumTopicId: 301640
dashedName: depth-first-search
---

# --description--

Kao što je sa <dfn>širokog pretraživanja</dfn>, ovde ćemo naučiti o drugom algoritmu za hodanje po grafu koji se zove <dfn>dubinska pretraga</dfn>.

Dok široka pretraga traži po koracima dužine ivica od izvornog čvora, <dfn>dubinska pretraga</dfn> počinje spuštanjem niz ivice dok ne stigne do kraja puta.

Kada stigne na kraj puta, pretraga se vraća do poslednjeg čvora sa ivicom koja nije posetena i nastavlja pretraživanje.

Slika ispod pokazuje kako algoritmi funkcionišu. Algoritam počinje od gornjeg čvora i poseti čvorove po numeričkom redosledu.

<img alt="animation showing the depth first search algorithm" src='https://cdn.freecodecamp.org/curriculum/coding-interview-prep/depth-first-search.gif'>

Primećujte kako, za razliku od široke pretrage, kada se čvor poseti, on ne posećuje sve susede. Umesto toga, počinje posetom jednog suseda i nastavlja po tom putu dok ne ostane nijedan drugi čvor za posetu na tom putu.

Da biste implementirali ovaj algoritam, moraćete da koristite stogu (stack). Stog je struktura podataka gde se poslednji dodati element prvi izvlači. Ovo je takođe poznato kao Last-In-First-Out (LIFO) struktura podataka. Staci su ključni u algoritmu dubinske pretrage jer, kada dodajemo susede na stog, želimo da posetimo najnovije dodate susede prvo i uklonimo ih sa stoga.

Jednostavan rezultat ovog algoritma je lista čvorova dostupnih od određenog čvora. Stoga ćete takođe morati da pratite posete čvorovima.

# --instructions--

Napišite funkciju `dfs()` koja prima neusmerizovanu matricu susedstva `graph` i čvor `root` kao parametre. Čvor će biti samo numerička vrednost čvora između `0` i `n - 1`, gde je `n` ukupan broj čvorova u grafu.

Vaša funkcija treba da vrati niz svih čvorova dostupnih od `root`-a.

# --hints--

Grafika unosa `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` sa početnim čvorom `1` treba da vrati niz podataka koji sadrži `0`, `1`, `2`, i `3`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })(),
  [0, 1, 2, 3]
);
```

Grafika unosa `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` sa početnim čvorom `3` treba da vrati niz podataka koji sadrži `3`, `2`, `1`, i `0`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3, 2, 1, 0]
);
```

Grafika unosa `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` sa početnim čvorom `1` treba da vrati niz podataka koji sadrži četiri elementa.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 1);
  })().length === 4
);
```

Grafika unosa `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` sa početnim čvorom `3` treba da vrati niz podataka koji sadrži `3`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })(),
  [3]
);
```

Grafika unosa `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` sa početnim čvorom `3` treba da vrati niz podataka koji sadrži jedan element.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 1, 0],
      [0, 1, 0, 0],
      [0, 0, 0, 0]
    ];
    return dfs(graph, 3);
  })().length === 1
);
```

Grafika unosa `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` sa početnim čvorom `3` treba da vrati niz podataka koji sadrži `2` i `3`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })(),
  [2, 3]
);
```

Grafika unosa `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` sa početnim čvorom `3` treba da vrati niz podataka koji sadrži dva elementa.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 3);
  })().length === 2
);
```

Grafika unosa `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` sa početnim čvorom `0` treba da vrati niz podataka koji sadrži `0` i `1`.

```js
assert.sameMembers(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })(),
  [0, 1]
);
```

Grafika unosa `[[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]]` sa početnim čvorom `0` treba da vrati niz podataka koji sadrži dva elementa.

```js
assert(
  (function () {
    var graph = [
      [0, 1, 0, 0],
      [1, 0, 0, 0],
      [0, 0, 0, 1],
      [0, 0, 1, 0]
    ];
    return dfs(graph, 0);
  })().length === 2
);
```

# --seed--

## --seed-contents--

```js
function dfs(graph, root) {

}

var exDFSGraph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
console.log(dfs(exDFSGraph, 3));
```

# --solutions--

```js
function dfs(graph, root) {
    var stack = [];
    var tempV;
    var visited = [];
    var tempVNeighbors = [];
    stack.push(root);
    while (stack.length > 0) {
        tempV = stack.pop();
        if (visited.indexOf(tempV) == -1) {
            visited.push(tempV);
            tempVNeighbors = graph[tempV];
            for (var i = 0; i < tempVNeighbors.length; i++) {
                if (tempVNeighbors[i] == 1) {
                    stack.push(i);
                }
            }
        }
    }
    return visited;
}
```