id: 587d8255367417b2b2512c74
title: Kreiranje klase za prioritetnu kolejku
challengeType: 1
forumTopicId: 301630
dashedName: create-a-priority-queue-class
---

# --description--

U ovom izazovu ćete kreirati prioritetnu kolejku. Prioritetna kolejka je posebna vrsta kolejke gdje elementi mogu imati dodatne informacije koje određuju njihov prioritet. To se lako može predstaviti cijelim brojem. Prioritet objekta mijenja redoslijed stavljanja, određujući niz elemenata koji se izvlače iz kolejke. Ako element s višim prioritetom bude postavljen nakon elemenata s nižim prioritetom, element s višim prioritetom će biti uklonjen prvi prije svih ostalih.

Na primjer, zamislite da imamo prioritetnu kolejku sa tri elementa:

```js
[['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

Ovdje druga vrijednost (cijeli broj) predstavlja prioritet objekta. Ako dodamo `['human', 1]` s prioritetom `1` (uzimajući u obzir da je niži prioritet viši prioritet), to će biti prvi element uklonjen iz kolejke. Kolekcija će izgledati ovako:

```js
[['human', 1], ['kitten', 2], ['dog', 2], ['rabbit', 2]]
```

Počeli smo sa pisanjem `PriorityQueue` u editoru koda. Morat ćete dodati metodu `enqueue` za dodavanje elemenata s prioritetom, metodu `dequeue` za uklanjanje i vraćanje elemenata, metodu `size` za vraćanje broja elemenata u kolejci, metodu `front` za vraćanje prvog elementa iz kolejke, i na kraju metodu `isEmpty` koja će vratiti `true` ako je kolejka prazna ili `false` ako nije.

Metoda `enqueue` treba da prima elemente u strukturi prikazanoj gore (`['human', 1']`) gdje predstavlja prioritet. Metode `dequeue` i `front` bi trebale vratiti samo ime objekta, ne njegov prioritet.

# --hints--

Vaša klasa `PriorityQueue` mora imati metodu `enqueue`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.enqueue === 'function';
  })()
);
```

Vaša klasa `PriorityQueue` mora imati metodu `dequeue`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.dequeue === 'function';
  })()
);
```

Vaša klasa `PriorityQueue` mora imati metodu `size`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.size === 'function';
  })()
);
```

Vaša klasa `PriorityQueue` mora imati metodu `front`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.front === 'function';
  })()
);
```

Vaša klasa `PriorityQueue` mora imati metodu `isEmpty`.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    return typeof test.isEmpty === 'function';
  })()
);
```

Vaša klasa `PriorityQueue` mora ispravno pratiti trenutni broj elemenata pomoću metode `size` prilikom dodavanja i uklanjanja elemenata iz kolejke.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    test.enqueue(['Jon Snow', 1]);
    var size1 = test.size();
    test.dequeue();
    var size2 = test.size();
    test.enqueue(['A', 3]);
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    return size1 === 2 && size2 === 1 && test.size() === 4;
  })()
);
```

Metoda `front` mora vratiti ispravan element koji je na vrhu kolejke prilikom dodavanja i uklanjanja elemenata.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['David Brown', 2]);
    var front1 = test.front();
    test.enqueue(['Jon Snow', 1]);
    var front2 = test.front();
    test.dequeue();
    test.enqueue(['A', 3]);
    var front3 = test.front();
    test.enqueue(['B', 3]);
    test.enqueue(['C', 3]);
    test.dequeue();
    var front4 = test.front();
    return (
      front1 === 'David Brown' &&
      front2 === 'Jon Snow' &&
      front3 === 'David Brown' &&
      front4 === 'A'
    );
  })()
);
```

Metoda `isEmpty` mora vratiti `true` kada je kolejka prazna.

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 1]);
    test.enqueue(['B', 1]);
    test.dequeue();
    var first = test.isEmpty();
    test.dequeue();
    return !first && test.isEmpty();
  })()
);
```

Prioritetna kolejka mora vratiti elemente s višim prioritetom prije elemenata s nižim prioritetom, a ostale elemente vraćati u redoslijedu prvog-ulaska-prvog-izlaska (FIFO).

```js
assert(
  (function () {
    var test = new PriorityQueue();
    test.enqueue(['A', 5]);
    test.enqueue(['B', 5]);
    test.enqueue(['C', 5]);
    test.enqueue(['D', 3]);
    test.enqueue(['E', 1]);
    test.enqueue(['F', 7]);
    var result = [];
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    result.push(test.dequeue());
    return result.join('') === 'EDABCF';
  })()
);
```

# --seed--

## --seed-contents--

```js
function PriorityQueue () {
  this.collection = [];
  this.printCollection = function() {
    console.log(this.collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function PriorityQueue() {
  this.collection = [];
  this.printCollection = function () {
    console.log(this.collection);
  };
  // Only change code below this line
  this.enqueue = function (newitem) {
    if (this.isEmpty()) {
      return this.collection.push(newitem);
    }

    this.collection = this.collection.reverse();
    var found_index = this.collection.findIndex(function (item) {
      return newitem[1] >= item[1];
    });
    if (found_index === -1) {
      this.collection.push(newitem);
    } else {
      this.collection.splice(found_index, 0, newitem);
    }
    this.collection = this.collection.reverse();
  };
  this.dequeue = function () {
    if (!this.isEmpty()) {
      return this.collection.shift()[0];
    } else {
      return "The queue is empty.";
    }
  };
  this.size = function () {
    return this.collection.length;
  };
  this.front = function () {
    return this.collection[0][0];
  };
  this.isEmpty = function () {
    return this.size() > 0 ? false : true;
  };
  // Only change code above this line
}
```