id: 587d8258367417b2b2512c80
title: Brisanje listnog čvora u binarnom stablu pretrage
challengeType: 1
forumTopicId: 301637
dashedName: delete-a-leaf-node-in-a-binary-search-tree
---

# --description--

Ovo je treći izazov od tri gde ćemo implementirati najtežu operaciju u binarnim stablima pretrage: brisanje. Brisanje je teško jer uklanjanje čvora razbija veze u stablo. Ove veze moraju pažljivo ponovo povezati kako bi se održala struktura binarnog stabla. Za neke slučajeve brisanja, to znači da stablo mora biti reorganizovano. Generalno, naići ćete na jedan od tri slučaja kada želite obrisati čvor: Listni čvor: Ciljni čvor nema potomke. Jedan potomak: Ciljni čvor ima samo jednog potomka. Dva potomka: Ciljni čvor ima dva potomka. Brisanje listnog čvora je jednostavno, samo ga brišemo. Brisanje čvora sa jednim potomkom je takođe jednostavno, samo ga brišemo i povežemo njegovog roditelja sa čvorom tog potomka. Brisanje čvora sa dva potomka je najteže, jer uvodi dva potomaka koja moraju biti ponovo povezana sa roditeljskim stabljem. Videćemo kako da rešimo ovaj slučaj u trećem izazovu. Štaviše, morate biti pažljivi sa nekim posebnim slučajevima kada upravljate brisanjem. Da li je stablo prazno? Da li je čvor za brisanje koren? Ima li stablo samo dva elementa? Za sada, hajde da se fokusiramo na prvi slučaj gde brišemo listni čvor.

# --instructions--

Napravite metodu u našem binarnom stablu pod nazivom `remove`. Ovde ćemo izgraditi logiku naše operacije brisanja. Prvo, želite napraviti funkciju unutar `remove` koja traži čvor koji želimo obrisati u trenutnom stablu. Ako čvor ne postoji u stablu, `remove` treba da vrati `null`. Sada, ako je ciljni čvor listni čvor bez potomaka, onda reference njegovog roditelja na taj čvor bi trebalo postaviti na `null`. Ovo efikasno briše taj čvor iz stabla. Da biste to uradili, takođe ćete morati pratiti roditeljski čvor čvora koji želimo obrisati. Takođe će biti korisno napraviti metodu za praćenje broja potomaka ciljnog čvora, jer će to odrediti koji slučaj brisanja obrađujemo. Obrašićemo drugi i treći slučaj u narednim izazovima. Srećno!

# --hints--

Struktura podataka `BinarySearchTree` bi trebalo da postoji.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Binarno stablo pretrage bi trebalo da ima metodu pod nazivom `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

Pokušaj brisanja elementa iz praznog stabla bi trebalo da vrati `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    return test.remove(100) == null;
  })()
);
```

Pokušaj brisanja elementa koji ne postoji bi trebalo da vrati `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(15);
    test.add(30);
    return test.remove(100) == null;
  })()
);
```

Ako koren čvora nema potomaka, brisanje bi trebalo da postavi koren na `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(500);
    test.remove(500);
    return test.inorder() == null;
  })()
);
```

Metoda `remove` bi trebalo da briše listne čvorove iz stabla.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.remove !== 'function') {
      return false;
    }
    test.add(5);
    test.add(3);
    test.add(7);
    test.add(6);
    test.add(10);
    test.add(12);
    test.remove(3);
    test.remove(12);
    test.remove(10);
    return test.inorder().join('') == '567';
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    },
    inorder: function() {
      if (this.root == null) {
        return null;
      } else {
        var result = new Array();
        function traverseInOrder(node) {
          if (node.left != null) {
            traverseInOrder(node.left);
          }
          result.push(node.value);
          if (node.right != null) {
            traverseInOrder(node.right);
          }
        }
        traverseInOrder(this.root);
        return result;
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}

function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
}
```

# --solutions--

```js
// solution required
```