id: 587d8254367417b2b2512c70
title: Kreiranje i dodavanje u Set-ove u ES6
challengeType: 1
forumTopicId: 301636
dashedName: create-and-add-to-sets-in-es6
---

# --description--

Sada, nakon rada sa ES5, uradićete nešto slično u ES6. Biće mnogo lakše. ES6 ima ugrađeni `Set` strukturu podataka, tako da su mnoge operacije koje ste ručno kreirali sada spremne za vas. Pogledajmo:

Kreiranje novog praznog seta:

```js
var set = new Set();
```

Možete kreirati set sa vrednostima:

```js
var set = new Set(1);
```

Možete kreirati set sa nizom podataka:

```js
var set = new Set([1, 2, 3]);
```

Nakon što kreirate set, možete dodati željene vrednosti koristeći metodu `add`:

```js
var set = new Set([1, 2, 3]);
set.add([4, 5, 6]);
```

Zapamtite, set je struktura podataka bez duplikata:

```js
var set = new Set([1, 2, 3, 1, 2, 3]);
// set sadrži samo [1, 2, 3]
```

# --instructions--

Za ovu vežbu, vratite set sa sledećim vrednostima: `1, 2, 3, 'Taco', 'Cat', 'Awesome'`

# --hints--

Vaš `Set` treba da sadrži samo vrednosti `1, 2, 3, Taco, Cat, Awesome`.

```js
assert(
  (function () {
    var test = checkSet();
    return (
      test.size == 6 &&
      test.has(1) &&
      test.has(2) &&
      test.has(3) &&
      test.has('Taco') &&
      test.has('Cat') &&
      test.has('Awesome')
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function checkSet() {
  var set = new Set([1, 2, 3, 3, 2, 1, 2, 3, 1]);
  // Only change code below this line

  // Only change code above this line
  console.log(Array.from(set));
  return set;
}

checkSet();
```

# --solutions--

```js
function checkSet(){var set = new Set([1,2,3,'Taco','Cat','Awesome']);
return set;}
```