id: 587d8255367417b2b2512c75
title: Kreiranje kružnog reda
challengeType: 1
forumTopicId: 301625
dashedName: create-a-circular-queue
---

# --description--

U ovom izazovu ćete kreirati kružni red (Circular Queue). Kružni red je struktura podataka koja dodaje podatke na kraj kolekcije, a zatim počinje ponovo zapisivati podatke na početku te kolekcije. Ovaj tip strukture podataka je koristan u određenim situacijama. Na primer, kružni red može se koristiti za prenosivost medija (streaming). Kada se red popuni, novi podaci o medijima će prepisati stare podatke.

Dobar način da se objasni ovaj koncept je korišćenjem niza podataka dužine `5`:

```js
[null, null, null, null, null]
 ^Čitanje @ 0
 ^Pisanje @ 0
```

Ovde su i čitaći i pišući pokazivači na poziciji `0`. Sada red prima tri nova elementa: `a`, `b` i `c`. Naš red sada izgleda ovako:

```js
[a, b, c, null, null]
 ^Čitanje @ 0
          ^Pisanje @ 3
```

Kada čitaći pročita, može ukloniti ili zadržati vrednost:

```js
[null, null, null, null, null]
                   ^Čitanje @ 3
                   ^Pisanje @ 3
```

Sada zapisujemo vrednosti `d`, `e` i `f` u red. Kada pišući pokazivač dođe na kraj niza, on se vraća na početak:

```js
[f, null, null, d, e]
                ^Čitanje @ 3
    ^Pisanje @ 1
```

Ovaj način zahteva fiksnu količinu memorije, ali omogućava rukovanje većim fajlovima.

# --instructions--

U ovom izazovu ćemo implementirati kružni red. Kružni red treba da pruži metode `enqueue` i `dequeue` koje će vam omogućiti čitanje i pisanje u red. Klasa objekta takođe mora prihvatiti ceo broj kao argument koji možete koristiti za određivanje veličine reda prilikom kreiranja. Predali smo vam početnu verziju ove klase u editoru koda.

Kada dodajete elemente u red, pokazivač za pisanje treba da se pomera napred i vraća na početak kada dođe do kraja reda. Metoda `enqueue` treba da vrati dodati element ako je uspešno; inače će vratiti `null`.

Slično tome, čitaći pokazivač treba da se pomera napred kada se elementi uklanjaju iz reda. Kada uklonite element, on treba biti vraćen. Ako ne možete ukloniti element, trebate vratiti `null`.

Pokaživač za pisanje ne bi trebalo da se pomeri dalje od čitaćeg pokazivača (naša klasa neće dozvoljavati pisanje na podatke koje još niste pročitali), a čitaći pokazivač ne bi trebalo da se pomera dalje od podataka koje ste napisali.

# --hints--

Metoda `enqueue` treba da dodaje elemente u kružni red.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

Ne smete dodati elemente više od čitaćeg pokazivača.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    test.enqueue(13);
    test.enqueue(25);
    test.enqueue(59);
    var print = test.print();
    return print[0] === 17 && print[1] === 32 && print[2] === 591;
  })()
);
```

Metoda `dequeue` treba da uklanja elemente iz reda.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 && test.dequeue() === 32 && test.dequeue() === 591
    );
  })()
);
```

Nakon uklanjanja elementa, njegova pozicija u nizu treba da se ponovo postavi na `null`.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(672);
    test.dequeue();
    test.dequeue();
    var print = test.print();
    return print[0] === null && print[1] === null && print[2] === 672;
  })()
);
```

Pokušaj uklanjanja više od pisanog pokazivača treba da vrati `null` i ne sme pomeriti pisajući pokazivač.

```js
assert(
  (function () {
    var test = new CircularQueue(3);
    test.enqueue(17);
    test.enqueue(32);
    test.enqueue(591);
    return (
      test.dequeue() === 17 &&
      test.dequeue() === 32 &&
      test.dequeue() === 591 &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.dequeue() === null &&
      test.enqueue(100) === 100 &&
      test.dequeue() === 100
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
class CircularQueue {
  constructor(size) {

    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line

    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line

    // Only change code above this line
  }
}
```

# --solutions--

```js
class CircularQueue {
  constructor(size) {
    this.queue = [];
    this.read = 0;
    this.write = 0;
    this.max = size - 1;

    while (size > 0) {
      this.queue.push(null);
      size--;
    }
  }

  print() {
    return this.queue;
  }

  enqueue(item) {
    // Only change code below this line
    console.log(this.write, this.max);
    if (this.queue[this.write] === null) {
      this.queue[this.write++] = item;

      if (this.write > this.max) {
        this.write = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }

  dequeue() {
    // Only change code below this line
    if (this.queue[this.read] !== null) {
      let item = this.queue[this.read];
      this.queue[this.read++] = null;
      if (this.read > this.max) {
        this.read = 0;
      }
      return item;
    }
    return null;
    // Only change code above this line
  }
}
```