id: 587d825b367417b2b2512c8e
title: Kreiraj hash tabelu
challengeType: 1
forumTopicId: 301627
dashedName: create-a-hash-table
---

# --description--

U ovom izazovu ćemo naučiti o hash tabelama. Hash tabela se koristi za implementaciju relacionalnih nizova podataka, ili mapa ključ-vrednost, kao što su `Object` i `Map` koje smo nedavno proučavali. JavaScript `Object` može funkcionisati kao hash tabela, na primer (njegova stvarna implementacija zavisi od okruženja u kojem se izvršava). Način na koji hash tabela radi je da uzme ulazni ključ i pouzdano ga hashuje do određene numeričke vrednosti. Ova numerička vrednost koristi se kao stvarni ključ po kojem se čuva povezana vrednost. Kada pokušate ponovo pristupiti istom ključu, funkcija hashovanja će obradovati ključ, vratiti sličan numerički rezultat, koji će zatim koristiti za pronalaženje povezane vrednosti. Ovo pruža prosečno O(1) vreme pristupa sa velikom efikasnošću.

Hash tabele mogu se implementirati kao nizovi podataka sa funkcijama hashovanja koje generišu indekse niza unutar definisanog opsega. Na ovaj način, izbor veličine niza je kritičan, baš kao i funkcija hashovanja. Na primer, da li će funkcija hashovanja generisati istu vrednost za dva različita ključa? Ovo se naziva kolizija (sukoba). Jedan od načina rešavanja kolizija je skladištenje svih parova ključ-vrednost u tom indeksu. Zatim, prilikom pristupa bilo kom od njih, moraćete da iterirate kroz skup objekata kako biste pronašli željeni ključ. Dobra funkcija hashovanja smanjuje kolizije kako bi se održavilo dobro vreme pretrage.

Ovde, nećemo se baviti detaljima hashovanja ili implementacijom hash tabele, samo ćemo pokušati da uhvatimo opšti osećaj kako funkcionišu.

# --instructions--

Implementirajte osnovnu funkcionalnost hash tabele. Kreirali smo standardnu funkciju hashovanja za vas. Možete proslediti niz znakova toj funkciji `hash` i ona će vam vratiti hash vrednost koju možete koristiti kao ključ za skladištenje. Skladištite objekte koristeći ovu hash vrednost u objektu `this.collection`. Kreirajte ove tri metode: `add`, `remove` i `lookup`. Prva treba da primi par ključ-vrednost koji se dodaje u hash tabelu. Druga treba da ukloni par ključ-vrednost kada se prosledi ključ. Treća treba da primi ključ i vrati povezanu vrednost ili `null` ako ključ ne postoji.

Obavezno osigurajte da svoj kod uzimate u obzir kolizije!

**Napomena:** Testiranje metode `remove` neće proći dok metode `add` i `lookup` nisu ispravno implementirane.

# --hints--

Funkcija `hash` mora biti validna.

```js
let calls = 0;
const ourHash = string => {
  calls++;
  let hashCode = 0;
  for (let i = 0; i < string.length; i++) {
    hashCode += string.charCodeAt(i);
  }
  return hashCode;
};

assert.strictEqual(hash('yek'),ourHash('yek'));

assert.strictEqual(hash('key'),ourHash('key'));
assert.strictEqual(hash('key1'),ourHash('key1'));
assert.strictEqual(hash('key2'),ourHash('key2'));
assert.strictEqual(hash('key3'),ourHash('key3'));

assert.strictEqual(hash('1key'),ourHash('1key'));
assert.strictEqual(hash('ke1y'),ourHash('ke1y'));
assert.strictEqual(hash('altKey'),ourHash('altKey'));

assert.strictEqual(called,calls); 
```

Struktura podataka `HashTable` mora postojati.

```js
assert.isDefined(HashTable); 
let newHashTable = new HashTable(); 
assert.isObject(newHashTable); 
```

`HashTable` mora imati metodu `add`.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.add); 
```

`HashTable` mora imati metodu `lookup`.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.lookup); 
```

`HashTable` mora imati metodu `remove`.

```js
let newHashTable = new HashTable(); 
assert.isFunction(newHashTable.remove); 
```

Metoda `add` mora dodati parove ključ-vrednost, a metoda `lookup` mora vratiti vrednosti povezane sa određenim ključem.

```js
let newHashTable = new HashTable(); 
newHashTable.add('key','value');
assert.strictEqual(newHashTable.lookup('key'),'value'); 
```

Metoda `remove` mora primiti ključ kao ulaz i ukloniti povezani par ključ-vrednost.

```js
let hashValue = hash('key');
let newHashTable = new HashTable(); 
newHashTable.add('key','value');
newHashTable.remove('key'); 
assert.notProperty(newHashTable.collection,hashValue); 
```

Metoda `remove` mora ukloniti samo odgovarajući par ključ-vrednost.

```js
let newHashTable = new HashTable(); 
let hashValue = hash('key');
newHashTable.add('key','value');
newHashTable.add('yek','value');
newHashTable.add('altKey','value');
newHashTable.remove('yek');

assert.notExists(newHashTable.lookup('yek'));
assert.exists(newHashTable.lookup('altKey'));
assert.exists(newHashTable.lookup('key'));

newHashTable.remove('key');

assert.notProperty(newHashTable.collection,hashValue); 
assert.exists(newHashTable.lookup('altKey'));

```

Objekti moraju biti dodati koristeći funkciju hashovanja.

```js
let newHashTable = new HashTable(); 
called = 0;
newHashTable.add('key1', 'value1')
newHashTable.add('key2', 'value2');
newHashTable.add('key3', 'value3');
assert.strictEqual(called,3); 
```

Hash tabela mora da rukuje kolizijama (sukobima).

```js
let newHashTable = new HashTable(); 
called = 0;
newHashTable.add('key1','value1');
newHashTable.add('1key', 'value2'); 
newHashTable.add('ke1y', 'value3'); 
assert.strictEqual(newHashTable.lookup('key1'),'value1');
assert.strictEqual(newHashTable.lookup('1key'),'value2');
assert.strictEqual(newHashTable.lookup('ke1y'),'value3');
```

# --seed--

## --seed-contents--

```js
var called = 0;
var hash = string => {
  called++;
  var hashed = 0;
  for (var i = 0; i < string.length; i++) {
    hashed += string.charCodeAt(i);
  }
  return hashed;
};
var HashTable = function() {
  this.collection = {};
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
var called = 0;
var hash = (string) => {
  called++;
  var hash = 0;
  for (var i = 0; i < string.length; i++) { hash += string.charCodeAt(i); }
  return hash;
};
var HashTable = function() {
  this.collection = {};
  // Only change code below this line

  this.add = function(key, val) {
    var theHash = hash(key);
    if (!this.collection.hasOwnProperty(theHash)) {
      this.collection[theHash] = {};
    }
    this.collection[theHash][key] = val;
  }

  this.remove = function(key) {
    var theHash = hash(key);
    var hashedObj = this.collection[theHash];
    if (hashedObj.hasOwnProperty(key)) {
      delete hashedObj[key];
    }
    if (!Object.keys(hashedObj).length) {
      delete this.collection[theHash];
    }
  }

  this.lookup = function(key) {
    var theHash = hash(key);
    if (this.collection.hasOwnProperty(theHash)) {
      return this.collection[theHash][key];
    }
    return null
  }
  // Only change code above this line
};
```