id: 587d8251367417b2b2512c64
title: Pronaﾄ訴 unutar spojenog popisa
challengeType: 1
forumTopicId: 301715
dashedName: search-within-a-linked-list
---

# --description--

Dodajmo joﾅ｡ nekoliko korisnih metoda u naﾅ｡u klasu spojenog popisa. Zar ne bi bilo korisno da moﾅｾemo reﾄ㌍ da li je naﾅ｡ popis prazan ili ne, kao ﾅ｡to je sluﾄ溝j sa naﾅ｡im klasama `Stack` i `Queue`?

Takoﾄ粗 bi trebalo da imamo moguﾄ㌻ost pronalaﾅｾenja specifiﾄ肱ih elemenata u naﾅ｡em spojenom popisu. Rad sa ovakvim strukturom podataka je neﾅ｡to za ﾅ｡ta ﾄ㌃te ﾅｾeleti mnogo veﾅｾbi! Napravimo metodu `indexOf` koja uzima `element` kao argument i vraﾄ㌢ `index` tog elementa u spojenom popisu. Ako element nije pronaﾄ粗n, vratiti `-1`.

Takoﾄ粗 napravimo suprotnu metodu: metodu `elementAt` koja uzima `index` kao argument i vraﾄ㌢ `element` na datoj `index`-u. Ako element ne postoji, vratiti `undefined`.

# --instructions--

Napiﾅ｡ite metodu `isEmpty` koja proverava da li je spojen popis prazan, metodu `indexOf` koja vraﾄ㌢ `index` odreﾄ粗nog elementa i metodu `elementAt` koja vraﾄ㌢ `element` na datoj `index`-u.

# --hints--

Vaﾅ｡a klasa `LinkedList` bi trebalo da ima metodu `isEmpty`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.isEmpty === 'function';
  })()
);
```

Vaﾅ｡a metoda `isEmpty` bi trebalo da vrati `false` kada postoji barem jedan element u spojenom popisu.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.isEmpty() === false;
  })()
);
```

Vaﾅ｡a metoda `isEmpty` bi trebalo da vrati `true` kada nema elemenata u spojenom popisu.

```js
assert(
  (function () {
    var test = new LinkedList();
    return test.isEmpty() === true;
  })()
);
```

Vaﾅ｡a klasa `LinkedList` bi trebalo da ima metodu `indexOf`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.indexOf === 'function';
  })()
);
```

Vaﾅ｡a metoda `indexOf` bi trebalo da vrati indeks odreﾄ粗nog elementa pronaﾄ粗nog u spojenom popisu.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.indexOf('cat') === 0;
  })()
);
```

Vaﾅ｡a metoda `indexOf` bi trebalo da vrati `-1` ako element nije pronaﾄ粗n u spojenom popisu.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.indexOf('pony') === -1;
  })()
);
```

Vaﾅ｡a klasa `LinkedList` bi trebalo da ima metodu `elementAt`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.elementAt === 'function';
  })()
);
```

Vaﾅ｡a metoda `elementAt` bi trebalo da vrati element na datoj poziciji u spojenom popisu.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.elementAt(1) === 'dog';
  })()
);
```

Vaﾅ｡a metoda `elementAt` bi trebalo da vrati `undefined` ako element na datoj poziciji u spojenom popisu ne postoji.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.elementAt(5) === undefined;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    var currentNode = head;
    var previousNode;
    if(currentNode.element === element){
      head = currentNode.next;
    } else {
      while(currentNode.element !== element) {
        previousNode = currentNode;
        currentNode = currentNode.next;
      }

      previousNode.next = currentNode.next;
    }

    length --;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
        var currentNode = head;

        while(currentNode.next){
            currentNode  = currentNode.next;
        }

        currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    var currentNode = head;
    var previousNode;
    if(currentNode.element === element){
        head = currentNode.next;
    } else {
        while(currentNode.element !== element) {
            previousNode = currentNode;
            currentNode = currentNode.next;
        }

        previousNode.next = currentNode.next;
    }

    length --;
  };

  this.indexOf = function(element) {
    if (head === null) return -1

    let current = head;
    let index = 0;

    while (current.element !== element && current.next !== null) {
      current = current.next;
      index++
    }

    if (current.element !== element && current.next === null) {
      return -1
    }

    return index;
  }

  this.elementAt = function(index) {
    if (head === null) return undefined;

    let current = head;
    let currentIndex = 0;

    while (currentIndex !== index && current.next !== null) {
      current = current.next;
      currentIndex++
    }

    if (currentIndex !== index && current.next === null) {
      return undefined;
    }

    return current.element;
  }

  this.isEmpty = function() {
    return length === 0;
  }
}
```