id: 587d825c367417b2b2512c90
title: Širinu-prva pretraga
challengeType: 1
forumTopicId: 301622
dashedName: breadth-first-search
---

# --description--

Dozej, učili smo različite načine predstavljanja grafova. Sada je pitanje: koja je udaljenost između bilo koje dve čvorove u grafu? Ovde dolaze algoritmi <dfn>traverširanja grafa</dfn>.

<dfn>Algoritmi traverširanja</dfn> su algoritmi za kretanje ili poseti čvorova u grafu. Jedan od tipova algoritama za kretanje je algoritam širinu-prva pretrage.

Ovaj algoritam počinje od jednog čvora i poseti sve njegove susede koji su udaljeni jednom ivicom. Zatim nastavlja da poseti susede svakog od njih i tako dalje, sve dok ne dosegnu sve čvorove.

Ključna struktura podataka koja će vam pomoći da implementirate algoritam širinu-prva pretrage je red (queue). To je niz podataka gde možete dodavati elemente na jednu stranu i uklanjati elemente sa druge strane. Takođe se naziva kao <dfn>FIFO</dfn> ili <dfn>First-In-First-Out</dfn> struktura podataka.

Pokazaće vam kako algoritam radi:
<img alt="animation showing the breadth first search algorithm" src='https://cdn.freecodecamp.org/curriculum/coding-interview-prep/breadth-first-search.gif'>

Siva boja pokazuje čvor koji se dodaje u red, a crna boja pokazuje čvor koji se uklanja iz reda. Posmatrajte kako svaki put kada se čvor ukloni iz reda (čvor postaje crn), svi njegovi susedi se dodaju u red (čvor postaje sivi).

Da biste implementirali ovaj algoritam, trebaće vam da unesete strukturu grafa i label čvora od kojeg želite početi.

Prvo, želite da znate udaljenost, ili broj ivica, od čvora za početak. Želite da postavite sve udaljenosti na veliku brojku, kao što je `Infinity`. Ovo sprečava probleme sa računanjem kada čvor nije dostupan od početnog čvora. Zatim, želite da krenete od početnog čvora ka njegovim susjedima. Ovi susedi su udaljeni jednom ivicom i trenutno trebate dodati jednu jedinicu udaljenosti na udaljenost koju čuvate.

# --instructions--

Napišite funkciju `bfs()` koja prima graf u matrici susednosti (dvodimenzionalni niz) i label čvora za koren kao parametre. Label čvora će biti samo numerička vrednost čvora između `0` i `n - 1`, gde je `n` ukupan broj čvorova u grafu.

Vaša funkcija treba da vrati JavaScript objekat sa parovima ključ-vrednost koji predstavljaju čvor i njegovu udaljenost od korena. Ako čvor nije dostupan, njegova udaljenost bi trebalo da bude `Infinity`.

# --hints--

Graf unosa `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` i početni čvor `1` bi trebalo da vrati `{0: 1, 1: 0, 2: 1, 3: 2}`

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 1);
assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: 2 });
```

Graf unosa `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]]` i početni čvor `1` bi trebalo da vrati `{0: 1, 1: 0, 2: 1, 3: Infinity}`

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 0],
  [0, 0, 0, 0]
];
var results = bfs(graph, 1);
 assert.deepEqual(results, { 0: 1, 1: 0, 2: 1, 3: Infinity });
```

Graf unosa `[[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]]` i početni čvor `0` bi trebalo da vrati `{0: 0, 1: 1, 2: 2, 3: 3}`

```js
var graph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
var results = bfs(graph, 0);
 assert.deepEqual(results, { 0: 0, 1: 1, 2: 2, 3: 3 });
```

Graf unosa `[[0, 1], [1, 0]]` i početni čvor `0` bi trebalo da vrati `{0: 0, 1: 1}`

```js
var graph = [
  [0, 1],
  [1, 0]
];
var results = bfs(graph, 0);
assert.deepEqual(results, { 0: 0, 1: 1 });
```

# --seed--

## --seed-contents--

```js
function bfs(graph, root) {
  var nodesLen = {};

  return nodesLen;
};

var exBFSGraph = [
  [0, 1, 0, 0],
  [1, 0, 1, 0],
  [0, 1, 0, 1],
  [0, 0, 1, 0]
];
console.log(bfs(exBFSGraph, 3));
```

# --solutions--

```js
function bfs(graph, root) {
  var nodesLen = {};
  // Set all distances to infinity
  for (var i = 0; i < graph.length; i++) {
    nodesLen[i] = Infinity;
  }
  nodesLen[root] = 0; // ...except root node
  var queue = [root]; // Keep track of nodes to visit
  var current; // Current node traversing
  // Keep on going until no more nodes to traverse
  while (queue.length !== 0) {
    current = queue.shift();
    // Get adjacent nodes from current node
    var curConnected = graph[current]; // Get layer of edges from current
    var neighborIdx = []; // List of nodes with edges
    var idx = curConnected.indexOf(1); // Get first edge connection
    while (idx !== -1) {
      neighborIdx.push(idx); // Add to list of neighbors
      idx = curConnected.indexOf(1, idx + 1); // Keep on searching
    }
    // Loop through neighbors and get lengths
    for (var j = 0; j < neighborIdx.length; j++) {
      // Increment distance for nodes traversed
      if (nodesLen[neighborIdx[j]] === Infinity) {
        nodesLen[neighborIdx[j]] = nodesLen[current] + 1;
        queue.push(neighborIdx[j]); // Add new neighbors to queue
      }
    }
  }
  return nodesLen;
}
```