id: 587d825b367417b2b2512c8c
title: Implementacija sortiranja pomoću minimalnog heapa
challengeType: 1
forumTopicId: 301643
dashedName: implement-heap-sort-with-a-min-heap
---

# --description--

Sada kada možemo dodavati i uklanjati elemente, istražimo neke aplikacije u kojima se heapi mogu koristiti. Heapi se često koriste za implementaciju prioritetskih redova jer uvek održavaju element sa najvećom ili najmanjom vrednošću na prvom mestu. Pored toga, koriste se za implementaciju algoritma sortiranja poznatog kao heap sort. Pokazaće nam kako to učiniti ovde. Heap sort koristi minimalni heap, koji je suprotan maksimalnom heapu. Minimalni heap uvek čuva element sa najmanjom vrednošću na mestu korena.

Heap sort radi uzimanjem neusumičnog niza podataka, dodavanjem svakog elementa u minimalni heap, a zatim uklanjanjem svakog elementa iz tog heapa i postavljanjem ga u novi niz. Struktura minimalnog heapa osigurava da će novi niz imati prethodne elemente sortirane od najmanjeg do najvećeg. Ovo je jedan od najefikasnijih algoritama sortiranja, sa prosečnom i najgorom kompleksnošću O(nlog(n)).

# --instructions--

Implementiraćemo heap sort koristeći minimalni heap. Molimo vas da ovde izmenite svoj kod za maksimalni heap. Kreirajte objekat `MinHeap` sa metodama `insert`, `remove` i `sort`. Metoda `sort` treba da vrati niz svih elemenata iz minimalnog heapa sortirane od najmanjeg do najvećeg.

# --hints--

Struktura podataka `MinHeap` mora postojati.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MinHeap` mora imati metodu pod nazivom `insert`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MinHeap` mora imati metodu pod nazivom `remove`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.remove == 'function';
  })()
);
```

`MinHeap` mora imati metodu pod nazivom `sort`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MinHeap !== 'undefined') {
      test = new MinHeap();
    } else {
      return false;
    }
    return typeof test.sort == 'function';
  })()
);
```

Metoda `sort` treba da vrati niz svih elemenata iz minimalnog heapa u sortiranom redosledu.

```js
assert(
  (() => {
    if (typeof MinHeap === 'undefined') {
      return false;
    }

    const heap = new MinHeap();
    const arr = createRandomArray(25);

    for (let i of arr) {
      heap.insert(i);
    }

    const result = heap.sort();
    arr.sort((a, b) => a - b);

    for (let i = 0; i < arr.length; i++) {
      if (arr[i] !== result[i]) {
        return false;
      }
    }
    return true;
  })()
);
```

# --seed--

## --seed-contents--

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
// Generate a randomly filled array
function createRandomArray(size = 5){
  let a = new Array(size);
  for(let i = 0; i < size; i++)
    a[i] = Math.floor(Math.random() * 100);
  
  return a;
}
const array = createRandomArray(25);

var MinHeap = function() {
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```