id: 587d825a367417b2b2512c8a
title: Ubacivanje elementa u maksimalni heap (Max Heap)
challengeType: 1
forumTopicId: 301703
dashedName: insert-an-element-into-a-max-heap
---

# --description--

Sada ćemo nastaviti sa drugom strukturom podataka tipa drvo, binarnim heap-om. Binarni heap je binarno stablo koje ima raspored elemenata koji poštuje svojstvo heap-a. Svojstvo heap-a opisuje odnos između roditeljskog čvora i njegovih potomaka. Možete imati maksimalni heap, gde svi roditeljski čvorovi su veći ili jednaki svojim potomcima, ili minimalni heap, gde je obrnuto tačno. Binarni heap-ovi su takođe potpuna binarna stabla. To znači da su sve nivoe stabla potpuno popunjeni i ako poslednji nivo nije potpuno popunjen, on se popunjava od leve ka desnoj strani.

Iako se binarni heap-ovi mogu implementirati kao strukture stabla sa čvorovima koji drže reference na levo i pravo, raspored elemenata prema svojstvu heap-a omogućava nam da predstavimo heap korišćenjem niza podataka. Odnos roditelj-potomek je ono što nas zanima i jednostavnim izračunavanjem možemo odrediti potomke bilo kog roditelja i roditelja bilo kog potonka.

Na primer, razmislite o ovom predstavljanju u obliku niza za minimalni binarni heap:

```js
[ 6, 22, 30, 37, 63, 48, 42, 76 ]
```

Korenski čvor je prvi element, `6`. Njegovi potomci su `22` i `30`. Ako pogledamo odnos između indeksa elemenata ovih vrednosti, za indeks `i`, potomci su `2 * i + 1` i `2 * i + 2`. Slično tome, element na indeksu `0` je roditelj ova dva potonaka koji se nalaze na indeksima `1` i `2`. Generalno, možemo pronaći roditelja bilo kog čvora za bilo koji indeks korišćenjem sledeće formule: `Math.floor((i - 1) / 2)`. Ove sistemi će biti tačni bez obzira na veličinu binarnog stabla. Konačno, možemo napraviti mala podešavanja kako bi ovo izračunavanje postalo jednostavnije tako što ćemo preskočiti prvi element u nizu podataka. To stvara sledeći odnos za bilo koji element na indeksu `i`:

Primer predstavljanja nizom:

```js
[ null, 6, 22, 30, 37, 63, 48, 42, 76 ]
```

Leva potomek elementa: `i * 2`

Desni potomek elementa: `i * 2 + 1`

Roditelj elementa: `Math.floor(i / 2)`

Čim shvatite ovo izračunavanje, korišćenje predstavljanja nizom je izuzetno važno jer se lokacije čvorova mogu brzo odrediti pomoću ovog izračunavanja i smanjuje se upotreba memorije jer nije potrebno skladištiti reference na potomke čvorove.

# --instructions--

Instrukcije: Ovde ćemo kreirati maksimalni heap (Max Heap). Počnite kreiranjem metode `insert` koja dodaje elemente u naš heap. Tokom ubacivanja, ključno je uvek održavati svojstvo heap-a. Za maksimalni heap, to znači da korenski čvor mora imati najveću vrednost u stablu i svi roditeljski čvorovi moraju biti veći od svojih potomaka. Implementacija heap-a korišćenjem niza podataka obično se dešava kroz tri koraka:

<ol>
  <li>Dodajte novi element na kraj niza podataka.</li>
  <li>Ako je element veći od svog roditelja, zamijenite njihove pozicije.</li>
  <li>Nastavite sa razmenama dok novi element ne bude manji od svog roditelja ili ne stignete do korena stabla.</li>
</ol>

Na kraju, dodajte metodu `print` koja vraća niz svih elemenata u heap-u.

# --hints--

Struktura podataka `MaxHeap` mora postojati.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    }
    return typeof test == 'object';
  })()
);
```

`MaxHeap` mora imati metodu pod nazivom `insert`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.insert == 'function';
  })()
);
```

`MaxHeap` mora imati metodu pod nazivom `print`.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

Metoda `insert` mora dodavati elemente prema svojstvu maksimalnog heap-a.

```js
assert(
  (function () {
    var test = false;
    if (typeof MaxHeap !== 'undefined') {
      test = new MaxHeap();
    } else {
      return false;
    }
    test.insert(50);
    test.insert(100);
    test.insert(700);
    test.insert(32);
    test.insert(51);
    test.insert(800);
    const result = test.print();
    const solution = JSON.stringify([null,800,51,700,32,50,100]);
    const solutionWithoutNull = JSON.stringify([800,51,700,32,50,100]);

    return (result.length == 6) ? (JSON.stringify(result) == solutionWithoutNull) : (JSON.stringify(result) == solution);
  })()
);
```

# --seed--

## --seed-contents--

```js
var MaxHeap = function() {
  // Only change code below this line
    
  // Only change code above this line
};
```

# --solutions--

```js
var MaxHeap = function() {
    // Only change code below this line
    this.heap = [];
    this.parent = index => {
      return Math.floor((index - 1) / 2);
    }
    this.insert = element => {
      this.heap.push(element);
      this.heapifyUp(this.heap.length - 1);
    }
    this.heapifyUp = index => {
      let currentIndex = index,
      parentIndex = this.parent(currentIndex);
      while (currentIndex > 0 && this.heap[currentIndex] > this.heap[parentIndex]) {
        this.swap(currentIndex, parentIndex);
        currentIndex = parentIndex;
        parentIndex = this.parent(parentIndex);
      }
    }
    this.swap = (index1, index2) => {
      [this.heap[index1], this.heap[index2]] = [this.heap[index2], this.heap[index1]];
    }
    this.print = () => {
      return this.heap;
    }
    // Only change code above this line
};
```