id: 587d8252367417b2b2512c67
title: Dodaj elemente na specifičnom indeksu u spojenoj listi
challengeType: 1
forumTopicId: 301619
dashedName: add-elements-at-a-specific-index-in-a-linked-list
---

# --description--

Napravimo metodu addAt(index,element) koja dodaje element na zadati indeks. Kao što koristimo za uklanjanje elemenata sa određenog indeksa, moramo pratiti trenutni indeks dok prolazimo kroz spojenu listu. Kada se trenutni indeks poklopi sa željenim indeksom, moramo postaviti novu vrednost svojstvu next prethodnog čvora kako bi on ukazivao na novi dodati čvor. A novi čvor mora imati referencu na čvor koji sledi na trenutnom indeksu. Vratimo se na primer sa lanca: nova osoba želi da se pridruži lancu, ali želi da se pridruži usred. Vi ste usred lanca, pa pružate ruke osobi koja je ispred vas. Nova osoba dolazi i stavlja svoje ruke u ruke osobe koja stoji u vašim rukama, a sada vi imate ruke na novoj osobi.

# --instructions--

Napravite metodu `addAt(index,element)` koja dodaje element na zadati indeks. Vratite `false` ako ne može se dodati element. **Napomena:** Proverite da li je zadani indeks negativan ili veći od dužine spojenoj liste.

# --hints--

Vaša metoda `addAt` treba da postavi vrednost `head` na novi čvor kada je zadani indeks 0.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'fish');
    return test.head().element === 'fish' && test.head().next.element === 'cat';
  })()
);
```

Vaša metoda `addAt` treba da poveća dužinu spojenoj liste za jedan svaki put kada se dodaje novi čvor.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.addAt(0, 'cat');
    return test.size() === 3;
  })()
);
```

Vaša metoda `addAt` treba da vrati `false` ako ne može da se doda čvor.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    return test.addAt(4, 'cat') === false;
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element) {
    this.element = element;
    this.next = null;
  };

  this.size = function() {
    return length;
  };

  this.head = function() {
    return head;
  };

  this.add = function(element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if (head === null){
        head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }
    length++;
  };
  this.addAt = function (index, element) {
    if (index > length || index < 0) {
      return false;
    }
    var newNode = new Node(element);
    var currentNode = head;
    if (index === 0) {
      head = newNode;
    } else {
      var previousNode = null;
      var i = 0;
      while (currentNode && i < index) {
        previousNode = currentNode;
        currentNode = currentNode.next;
        i++;
      }
      previousNode.next = newNode;
    }
    newNode.next = currentNode;
    length++;
  }
}
```