id: 587d8251367417b2b2512c61
title: Rad sa čvorovima u povezanoj listi
challengeType: 1
forumTopicId: 301721
dashedName: work-with-nodes-in-a-linked-list
---

# --description--

Drugi struktura podataka koja se često susreće u računarstvu je <dfn>povezana lista</dfn>. Povezana lista je sekvenca elemenata podataka, nazvana 'čvorovi', svaki usmeren ka sledećem. Svaki <dfn>čvor</dfn> u povezanoj listi ima dva ključna atributa: sam `element`, i referencu na sledeći `node`.

Zamislite da ste u redovima za pješačke. Vi imate ruke oko osobe koja sledi u redu, a osoba iza vas ima ruke oko vas. Možete direktno videti osobu ispred sebe, ali oni blokiraju pogled na druge ljude koji su dalje u redu. Čvor je kao osoba u redu: zna ko je i može da vidi samo sledeću osobu, ali ne zna ništa o ljudima koji su ispred ili iza njih.

# --instructions--

U našem editoru koda kreirali smo dva čvora, `Kitten` i `Puppy`, i ručno povezali čvor `Kitten` sa čvorom `Puppy`.

Kreirajte čvor `Cat` i čvor `Dog` i ručno ih povežite u listu.

# --hints--

Vaš `Puppy` čvor treba da pokazuje na `Cat` čvor.

```js
assert(Puppy.next.element === 'Cat');
```

Vaš `Cat` čvor treba da pokazuje na `Dog` čvor.

```js
assert(Cat.next.element === 'Dog');
```

# --seed--

## --seed-contents--

```js
var Node = function(element) {
  this.element = element;
  this.next = null;
};
var Kitten = new Node('Kitten');
var Puppy = new Node('Puppy');

Kitten.next = Puppy;
// Only change code below this line
```

# --solutions--

```js
// solution required
```