id: 587d8251367417b2b2512c65
title: Uklanjanje elemenata iz povezanog spiska po indeksu
challengeType: 1
forumTopicId: 301711
dashedName: remove-elements-from-a-linked-list-by-index
---

# --description--

Pre nego što nastavimo sa drugom strukturom podataka, hajde da uradimo nekoliko zadnjih vežbi sa povezanim spiskovima.

Implementirajte metodu `removeAt` koja uklanja element na određenom `index`-u. Metoda treba da se zove `removeAt(index)`. Da bismo uklonili element na određenu `index`, moraćemo da brojim čvorove dok prolazimo kroz povezani spisak.

Najčešći pristup za prolazak kroz elemente povezanog spiska je korišćenje <dfn>'iteratora'</dfn>, ili sentinel-a, koji 'pokazuje' čvorove koje vaš kod upoređuje. U našem slučaju, počevši od `head`-a našeg spiska, inicijalizujemo varijablu `currentIndex` koja počinje na `0`. `currentIndex` mora da se inkrementira za svaki čvor koji prođe.

Kao što je slučaj sa našom metodom `remove(element)`, koju smo <a href="/learn/coding-interview-prep/data-structures/remove-elements-from-a-linked-list" target="_blank" rel="noopener noreferrer nofollow">naučili u prethodnom kursu</a>, moramo biti pažljivi da ne ostavimo drugi deo našeg spiska bez čvorova kada uklanjamo čvor u našoj metodi `removeAt(index)`. Osiguravamo kontinuitet čvorova tako što garantujemo da čvor koji pokazuje na uklonjeni čvor, sada pokazuje na sledeći čvor.

# --instructions--

Napišite metodu `removeAt(index)` koja uklanja i vraća čvor na određenom `index`-u. Metoda treba da vrati `null` ako je dati `index` negativan ili veći/jednak dužini povezanog spiska.

**Napomena:** Ne zaboravite da brojiте `currentIndex`.

# --hints--

Vaša klasa `LinkedList` bi trebalo da ima metodu `removeAt`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.removeAt === 'function';
  })()
);
```

Vaša metoda `removeAt` bi trebalo da smanji dužinu povezanog spiska za jedan.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.removeAt(1);
    return test.size() === 2;
  })()
);
```

Vaša metoda `removeAt` bi trebalo da ukloni element na navedenom indeksu iz povezanog spiska.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.add('bird');
    test.removeAt(1);
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"kitten","next":{"element":"bird","next":null}}}'
    );
  })()
);
```

Kada postoji samo jedan element u povezanom spisku, vaša metoda `removeAt` bi trebalo da ukloni i vrati taj element, smanjujući dužinu povezanog spiska.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    var removedItem = test.removeAt(0);
    return test.head() === null && test.size() === 0 && removedItem === 'cat';
  })()
);
```

Vaša metoda `removeAt` bi trebalo da vrati element uklonjenog čvora.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    return test.removeAt(1) === 'dog';
  })()
);
```

Vaša metoda `removeAt` bi trebalo da vrati `null` i da spisak ostane nepromenjen ako je navedeni indeks manji od `0`.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(-1);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

Vaša metoda `removeAt` bi trebalo da vrati `null` i da spisak ostane nepromenjen ako je navedeni indeks veći ili jednak dužini.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    var removedItem = test.removeAt(3);
    return (
      removedItem === null &&
      JSON.stringify(test.head()) ===
        '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
      head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function (element) {
    this.element = element;
    this.next = null;
  };

  this.size = function () {
    return length;
  };

  this.head = function () {
    return head;
  };

  this.add = function (element) {
    var node = new Node(element);
    if (head === null) {
      head = node;
    } else {
      var currentNode = head;

      while (currentNode.next) {
        currentNode = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.removeAt = function (index) {
    var currentNode = head;
    var previous = head;
    var count = 0;
    if (index >= length || index < 0) {
      return null;
    }
    if (index === 0) {
      var removed = head.element;
      head = currentNode.next;
    } else {
      while (count < index) {
        previous = currentNode;
        currentNode = currentNode.next;
        count++;
      }
      var removed = previous.next.element;
      previous.next = currentNode.next;
    }
    length--;
    return removed;
  };
}
```