id: 587d8250367417b2b2512c60
title: Kreiranje klase reda
challengeType: 1
forumTopicId: 301631
dashedName: create-a-queue-class
---

# --description--

Kao i stike, red je kolekcija elemenata. Ali za razliku od stika, red prati princip FIFO (First In First Out - Prvo ulazi, prvo izlazi). Elementi koji se dodaju u red dodaju se na zadnji kraj, ili kraj reda, a samo element na prednjem kraju reda moéƒe biti uklonjen.

Moéƒemo koristiti niz podataka za predstavljanje reda, ali kao i stike, éƒelimo smanjiti nivo apstrakcije koji imamo nad naéÀim redovima.

Dve glavne metode klase reda su metoda `enqueue` i metoda `dequeue`. Metoda `enqueue` dodaje element na zadnji kraj reda, a metoda `dequeue` uklanja i vraáa element sa prednjeg kraja reda. Druge vaéƒne metode su `front`, `size` i `isEmpty`.

# --instructions--

NapiéÀite metodu `enqueue` koja dodaje element na zadnji kraj reda, metodu `dequeue` koja uklanja i vraáa prednji element, metodu `front` koja nam omoguáava da vidimo prednji element, metodu `size` koja prikazuje duéƒinu, i metodu `isEmpty` za proveru da li je red prazan.

# --hints--

VaéÀa klasa `Queue` mora imati metodu `enqueue`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.enqueue === 'function';
  })()
);
```

VaéÀa klasa `Queue` mora imati metodu `dequeue`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.dequeue === 'function';
  })()
);
```

VaéÀa klasa `Queue` mora imati metodu `front`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.front === 'function';
  })()
);
```

VaéÀa klasa `Queue` mora imati metodu `size`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.size === 'function';
  })()
);
```

VaéÀa klasa `Queue` mora imati metodu `isEmpty`.

```js
assert(
  (function () {
    var test = new Queue();
    return typeof test.isEmpty === 'function';
  })()
);
```

Metoda `dequeue` treba da ukloni i vrati prednji element reda

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.dequeue() === 'Smith';
  })()
);
```

Metoda `front` treba da vrati vrednost prednjeg elementa reda

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    test.enqueue('John');
    return test.front() === 'Smith';
  })()
);
```

Metoda `size` treba da vrati duéƒinu reda

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return test.size() === 1;
  })()
);
```

Metoda `isEmpty` treba da vrati `false` ako postoje elementi u redu

```js
assert(
  (function () {
    var test = new Queue();
    test.enqueue('Smith');
    return !test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Queue() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function Queue () { 
    var collection = [];
    this.print = function() {
        console.log(collection);
    };
    // Only change code below this line
    this.enqueue = function(item) {
        collection.push(item);
    }

    this.dequeue = function() {
        return collection.shift();
    }

    this.front = function() {
        return collection[0];
    }

    this.size = function(){
        return collection.length;
    }

    this.isEmpty = function() {
        return collection.length === 0 ? true : false;
    }
    // Only change code above this line
}
```