id: 587d8257367417b2b2512c7d
title: Pronađite minimalnu i maksimalnu visinu binarnog stabla pretrage
challengeType: 1
forumTopicId: 301641
dashedName: find-the-minimum-and-maximum-height-of-a-binary-search-tree
---

# --description--

U prethonom izazovu, napisali smo slučaj kada stablo može biti neravnoteženo. Da bismo razumeli koncept ravnoteže, istražimo drugu karakteristiku stabla: visinu. Visina u stablu predstavlja rastojanje od kore novog čvora do bilo kog listnog čvora. Različiti putevi u strukturi stabla sa mnogo grana mogu imati različite visine, ali za određeno stablo će postojati minimalna i maksimalna visina. Ako je stablo ravnomerno (balanced), ove vrednosti će biti različite za maksimalnu visinu. To znači da u ravnotežnom stablu su svi listni čvorovi na istoj visini, ili ako nisu na istoj visini, onda su barem na istoj udaljenosti.

Karakteristika ravnoteže je važna za stablo jer određuje efikasnost operacija stabla. Kao što smo objasnili u prethodnom izazovu, susrećemo se sa problemom najgoreg slučaja kod potpuno neravnoteženih stabala. Samoregulujuća stabla koriste se široko za rešavanje ovog problema u stablama sa promenljivim skupovima podataka. Uobiđi primeri su AVL stabla, crveno-crno stabla i B stabla. Sva ova stabla imaju dodatnu logiku koja popravlja ravnotežu stabla prilikom ubacivanja ili brisanja elemenata što dovodi do stanja neravnoteženosti.

**Napomena:** Karakteristika slična visini je dubina, koja se odnosi na rastojanje određenog čvora od kore novog čvora.

# --instructions--

Napišite dve funkcije za naše binarno stablo: `findMinHeight` i `findMaxHeight`. Ove funkcije treba da vrate celobrojnu vrednost minimalne i maksimalne visine unutar datog binarnog stabla, respektivno. Ako je čvor prazan (null), molimo vas da vratite visinu `-1` (to je bazni slučaj). Na kraju, dodajte treću funkciju `isBalanced` koja vraća `true` ili `false` u zavisnosti od toga da li je stablo ravnomerno. Možete koristiti dve funkcije koje ste prethodno napisali za ovo.

# --hints--

Struktura podataka `BinarySearchTree` bi trebalo da postoji.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Binarno stablo pretrage bi trebalo da ima funkciju pod nazivom `findMinHeight`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.findMinHeight == 'function';
  })()
);
```

Binarno stablo pretrage bi trebalo da ima funkciju pod nazivom `findMaxHeight`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.findMaxHeight == 'function';
  })()
);
```

Binarno stablo pretrage bi trebalo da ima funkciju pod nazivom `isBalanced`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.isBalanced == 'function';
  })()
);
```

Funkcija `findMinHeight` bi trebalo da vrati minimalnu visinu stabla.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMinHeight !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.findMinHeight() == 1;
  })()
);
```

Funkcija `findMaxHeight` bi trebalo da vrati maksimalnu visinu stabla.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMaxHeight !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.findMaxHeight() == 5;
  })()
);
```

Prazno stablo bi trebalo da vrati visinu `-1`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.findMaxHeight !== 'function') {
      return false;
    }
    return test.findMaxHeight() == -1;
  })()
);
```

Funkcija `isBalanced` bi trebalo da vrati `false` ako stablo nije ravnotežno binarno stablo pretrage.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isBalanced !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    return test.isBalanced() === false;
  })()
);
```

Funkcija `isBalanced` bi trebalo da vrati `true` ako je stablo ravnotežno binarno stablo pretrage.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isBalanced !== 'function') {
      return false;
    }
    test.add(10);
    test.add(3);
    test.add(22);
    test.add(1);
    test.add(4);
    test.add(17);
    test.add(32);
    return test.isBalanced() === true;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }

      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
  this.findMinHeight = function(root = this.root) {
    // empty tree.
    if (root === null) {
      return -1;
    }
    // leaf node.
    if (root.left === null && root.right === null) {
      return 0;
    }
    if (root.left === null) {
      return this.findMinHeight(root.right) + 1;
    }
    if (root.right === null) {
      return this.findMinHeight(root.left) + 1;
    }
    const lHeight = this.findMinHeight(root.left);
    const rHeight = this.findMinHeight(root.right);
    return Math.min(lHeight, rHeight) + 1;
  };
  this.findMaxHeight = function(root = this.root) {
    // empty tree.
    if (root === null) {
      return -1;
    }
    // leaf node.
    if (root.left === null && root.right === null) {
      return 0;
    }
    if (root.left === null) {
      return this.findMaxHeight(root.right) + 1;
    }
    if (root.right === null) {
      return this.findMaxHeight(root.left) + 1;
    }
    const lHeight = this.findMaxHeight(root.left);
    const rHeight = this.findMaxHeight(root.right);
    return Math.max(lHeight, rHeight) + 1;
  };
  this.isBalanced = function(root = this.root) {
    if (root === null) {
      return true;
    }

    if (root.left === null && root.right === null) {
      return true;
    }

    if (root.left === null) {
      return this.findMaxHeight(root.right) <= 0;
    }

    if (root.right === null) {
      return this.findMaxHeight(root.left) <= 0;
    }

    const lHeight = this.findMaxHeight(root.left);
    const rHeight = this.findMaxHeight(root.right);
    if (Math.abs(lHeight - rHeight) > 1) {
      return false;
    }
    return this.isBalanced(root.left) && this.isBalanced(root.right);
  };
}
```