id: 587d8259367417b2b2512c84
title: Kreiranje stabla pretrage sa stabljem reči
challengeType: 1
forumTopicId: 301634
dashedName: create-a-trie-search-tree
---

# --description--

Ovde ćemo otići od binarnih stabla pretrage i pogledati drugi tip strukture stabla koji se naziva stablo reči (Trie). Stablo reči je organizovano stablo za pretraživanje koje se često koristi za skladištenje sekvenci karaktera, ili generalizovanije nizove podataka ili promenljive skupove podataka gde su ključevi sekvence karaktera. Izuzetno je korisno za skladištenje skupova podataka kada će mnogi ključevi imati slične prefikse, na primer, sve reči u rječniku. Za razliku od binarnog stabla, čvorovi se ne odnose na stvarni vrednost. Umjesto toga, put do čvora predstavlja specifični ključ. Na primer, ako želimo da sačuvamo sekvencu karaktera 'code' u stablu reči, imali bismo četiri čvora, po jedan za svaki karakter: c — o — d — e. Praćenje te putanje kroz sve ove čvorove stvoriće 'code' kao sekvencu karaktera — ta putanja je ključ koji smo sačuvali. Zatim, ako želimo da dodamo sekvencu karaktera 'coding', podijeliti bismo tri prvih čvora od 'code' pre nego što se razlikuju nakon 'd'. Na ovaj način, veliki skupovi podataka mogu biti sačuvani veoma efikasno. Štaviše, pretraga može biti vrlo brza jer je efikasnost ograničena dužinom sekvence karaktera koja se skladišti. Pored toga, za razliku od binarnog stabla, čvor može sadržavati bilo koji broj čvorova potomaka. Kao što možete da zamislite iz gornjeg primera, određene metadate obično se skladištaju na čvorima koji označavaju kraj ključa kako bi u budućim pretragama taj ključ i dalje bio dostupan. Na primer, ako dodamo 'codes' u naš gornji primer, morali bismo imati način da znamo da 'e' u 'code' predstavlja kraj unetog ključa. Inače, ove informacije bi efikasno nestale kada dodajemo 'codes'.

# --instructions--

Napravićemo stablo reči koje skladišti reči. Primaćeš reči putem metode `add` i postaviće ih u strukturu podataka stabla reči. Takođe ćeš biti u mogućnosti da odgovoriš na pitanje da li je uneta sekvenca karaktera reč putem metode `isWord`, i da dobiješ sve reči unesene u stablo reči putem metode `print`. Metoda `isWord` treba da vrati boolean vrednost, a `print` treba da vrati niz svih reči kao sekvencu karaktera. Da bismo potvrdili da je ova struktura podataka implementirana pravilno, dostavili smo strukturu `Node` za svaki čvor u stablu. Svaki čvor će biti objekat sa svojstvom `keys`, koje je JavaScript `Map`. Ovo će držati pojedinačne karaktere koji su važeći ključevi svakog čvora. Takođe smo kreirali svojstvo `end` na čvorovima, koje se može postaviti na `true` ako čvor predstavlja kraj reči.

# --hints--

`Trie` treba da ima metodu `add`.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

`Trie` treba da ima metodu `print`.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    return typeof test.print == 'function';
  })()
);
```

`Trie` treba da ima metodu `isWord`.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    return typeof test.isWord == 'function';
  })()
);
```

Metoda `print` treba da vrati sve reči dodate u stablo reči kao niz karaktera.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    test.add('jump');
    test.add('jumps');
    test.add('jumped');
    test.add('house');
    test.add('mouse');
    var added = test.print();
    return (
      added.indexOf('jump') != -1 &&
      added.indexOf('jumps') != -1 &&
      added.indexOf('jumped') != -1 &&
      added.indexOf('house') != -1 &&
      added.indexOf('mouse') != -1 &&
      added.length == 5
    );
  })()
);
```

Metoda `isWord` treba da vrati `true` samo za reči unete u stablo reči i `false` za sve ostale.

```js
assert(
  (function testTrie() {
    var test = false;
    if (typeof Trie !== 'undefined') {
      test = new Trie();
    } else {
      return false;
    }
    test.add('hop');
    test.add('hops');
    test.add('hopped');
    test.add('hoppy');
    test.add('hope');
    return (
      test.isWord('hop') &&
      !test.isWord('ho') &&
      test.isWord('hopped') &&
      !test.isWord('hopp') &&
      test.isWord('hoppy') &&
      !test.isWord('hoping')
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
var Node = function() {
  this.keys = new Map();
  this.end = false;
  this.setEnd = function() {
    this.end = true;
  };
  this.isEnd = function() {
    return this.end;
  };
};
var Trie = function() {
  // Only change code below this line
  
  // Only change code above this line
};
```

# --solutions--

```js
// solution required
```