id: 587d8257367417b2b2512c7c
title: Proverite da li je element prisutan u binarnom stablu pretrage
challengeType: 1
forumTopicId: 301623
dashedName: check-if-an-element-is-present-in-a-binary-search-tree
---

# --description--

Sada kada imamo opšte razumevanje binarnog stabla pretrage, hajde da pričamo o tome detaljnije. Binarna stablo pretraga pružaju logaritamsko vreme za standardne operacije pretraživanja, ubacivanja i brisanja u proseku, i linearno vreme u najgorem slučaju. Zašto je tako? Svaka od ovih osnovnih operacija zahteva da izvršimo pretragu elementa u stablu (ili u slučaju ubacivanja, pretragu mesta gde bi trebalo postaviti) i zbog strukture stabla, za svakog roditeljskog čvora, delimo levi ili desni deo i efikasno smanjujemo veličinu preostalog stabla na pola. To čini pretraživanje sličnim logaritmu broja čvorova u stablu, što stvara logaritamsko vreme za ove operacije u proseku. Dobro, ali šta je sa najgorem slučaju? Pa, zamislite da gradimo stablo iz sledećih vrednosti, dodajući ih od leve ka desnoj: `10`, `12`, `17`, `25`. Prateći pravila našeg binarnog stabla pretrage, dodaćemo `12` na desnu stranu `10`, `17` na desnu stranu toga, i `25` na desnu stranu toga. Sada naše stablo izgleda kao spojen lista, i prolazak kroz njega za pretraživanje `25` zahteva prolazak kroz sve elemente u linearnom redosledu. Dakle, linearno vreme u najgorem slučaju. Problem je što stablo nije uravnoteženo. Detaljnije ćemo istražiti značenje ovoga u sledećim izazovima.

# --instructions--

U ovom izazovu, kreiraćemo alat za naše stablo. Napišite metodu `isPresent` koja prima vrednost celog broja kao ulaz i vraća tačno ili netačno da li je ta vrednost prisutna u binarnom stablu pretrage.

# --hints--

Klasa podataka `BinarySearchTree` bi trebalo da postoji.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Binarno stablo pretrage bi trebalo da ima metodu pod nazivom `isPresent`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.isPresent == 'function';
  })()
);
```

Metoda `isPresent` bi trebalo tačno da proverava prisustvo ili neprisutnost elemenata dodatih u stablo.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    test.add(4);
    test.add(7);
    test.add(411);
    test.add(452);
    return (
      test.isPresent(452) &&
      test.isPresent(411) &&
      test.isPresent(7) &&
      !test.isPresent(100)
    );
  })()
);
```

`isPresent` bi trebalo da rukuje slučajem kada je stablo prazno.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.isPresent !== 'function') {
      return false;
    }
    return test.isPresent(5) == false;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        function searchTree(node) {
          if (value < node.value) {
            if (node.left == null) {
              node.left = new Node(value);
              return;
            } else if (node.left != null) {
              return searchTree(node.left);
            }
          } else if (value > node.value) {
            if (node.right == null) {
              node.right = new Node(value);
              return;
            } else if (node.right != null) {
              return searchTree(node.right);
            }
          } else {
            return null;
          }
        }
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = (tree) => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.isPresent = function (value) {
    var current = this.root
    while (current) {
      if (value === current.value) {
        return true;
      }
      current = value < current.value ? current.left : current.right;
    }
    return false;
  }
}
```