id: 587d8257367417b2b2512c7b
title: Dodajte novi element u binarno traželjno stablo
challengeType: 1
forumTopicId: 301618
dashedName: add-a-new-element-to-a-binary-search-tree
---

# --description--

Ova serija izazova će predstaviti strukturu podataka stabla. Stablo je ključna i veoma upotrebljiva struktura u računarstvu. Naravno, ime potiče od činjenice da kada pogledate njegovu sliku, izgleda kao drveće koje poznajemo u stvarnom svetu. Struktura podataka stablo počinje sa jednim čvorom, obično se naziva koren, i odatle se širi ka drugim čvorovima, svaki sa više ćorsaka (podčvorova), i tako dalje. Ova struktura podataka se obično prikazuje sa korenom na vrhu; možete je zamisliti kao prirodno drvo koje je okrenuto glavom dole.

Prvo, hajde da objasnimo neke uobičajene termine koje ćemo naići kada je o stablima reč. Koreni čvor je na samom vrhu stabla. Tačke podataka u stablu nazivaju se čvorovi (nodes). Čvor sa ćorsakama koje vode ka drugim čvorovima naziva se roditelj tog ćvora ka kojem ta grana vodi (podčvor). Koriste se i drugi terminološki pojmovi, kao što očekujete. Podstablo se odnosi na sve potomke određenog čvora; grane mogu se zvati ivice (edges), a listni čvorovi su konačni čvorovi stabla koji nemaju potomke. Napokon, zapamtite da su stabla strukture podataka koje imaju rekurentnu (rekurzivnu) prirodu. To znači da su potomci čvora roditelji njegovog podstabla, i tako dalje. Razumevanje rekurzivne prirode stabala je ključno prilikom dizajniranja algoritama za standardne operacije na stablu.

Početi ćemo sa razmatranjem specifičnog tipa stabla, binarnog stabla. Zapravo, razmatraćemo binarno traželjno stablo (Binary Search Tree). Neka objasnimo šta to znači. Iako struktura podataka stabla može imati bilo koje grane za jedan čvor, binarno stablo može imati samo dve grane po čvoru. Štaviše, binarno traželjno stablo je organizovano prema podstablima potomaka, tako da vrednost svakog čvora u levom podstablu je manja ili jednaka vrednosti roditeljskog čvora, a vrednost svakog čvora u desnom podstablu je veća ili jednaka vrednosti roditeljskog čvora. Pomaže mnogo da se vide ove odnose kako bi se bolje razumelo:

<div style='width: 100%; display: flex; justify-content: center; align-items: center;'><img alt="an example of a binary search tree" style='width: 100%; max-width: 350px; background-color: var(--gray-05);' src='https://user-images.githubusercontent.com/18563015/32136009-1e665d98-bbd6-11e7-9133-63184f9f9182.png'></div>

Sada je ovaj organizovani odnos veoma jasan. Zapamtite da je svaka vrednost na levoj strani 8, korena, manja od 8, i svaka vrednost na desnoj strani je veća od 8. Takođe primetite da se ovaj odnos primenjuje za svaki podstabao. Na primer, levog potomka je podstablo. 3 je roditelj, i ima samo dva potomka — po pravilima binarnog traželjnog stabla, znamo čak i bez proveravanja da će levi potomek ovog čvora (i svi njegovi potomci) biti manji od 3, a desni potomek (i svi njegovi potomci) bit će veći od 3 (ali takođe manji od vrednosti korena strukture), i tako dalje.

Binarna traželjna stabla su veoma česte i korisne strukture podataka jer pružaju logaritamsko vreme u proseku za nekoliko standardnih operacija kao što su pretraživanje, ubacivanje i brisanje.

# --instructions--

Počećemo lako. Objasнили smo osnovnu strukturu binarnog traželjnog stabla ovde zajedno sa funkcijom za kreiranje čvora za naše stablo. Proverite da svaki čvor može imati levu i desnu vrednost. One će biti dodeljene podstablima potomaka, ako postoje. U našem binarnom traželjnom stablu, treba da napravite metodu za dodavanje nove vrednosti u stablo. Ova metoda bi trebalo da se zove `add` i da prima ceo broj koji treba dodati u stablo. Pazite na održavanje pravila binarnog traželjnog stabla: vrednost u svakom levom potomku mora biti manja ili jednaka vrednosti roditeljskog čvora, a vrednost u svakom desnom potomku mora biti veća ili jednaka vrednosti roditeljskog čvora. Ovde, učinite da naše stablo ne može da sadrži duplirane vrednosti. Ako pokušamo dodati vrednost koja već postoji, metoda bi trebalo da vrati `null`. Inače, ako je dodavanje uspešno, treba vratiti `undefined`.

**Hint:** Stabla su strukture podataka sa prirodnom rekurzivnom prirodom!

# --hints--

Struktura podataka `BinarySearchTree` mora postojati.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Binarno traželjno stablo mora imati metodu pod nazivom `add`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.add == 'function';
  })()
);
```

Metoda za dodavanje mora dodavati elemente u skladu sa pravilima binarnog traželjnog stabla.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    test.add(1);
    test.add(1);
    test.add(7);
    test.add(87);
    test.add(34);
    test.add(45);
    test.add(73);
    test.add(8);
    const expectedResult = [1, 4, 7, 8, 34, 45, 73, 87];
    const result = test.inOrder();
    return expectedResult.toString() === result.toString();
  })()
);
```

Dodavanje elementa koji već postoji mora vratiti `null`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.add !== 'function') {
      return false;
    }
    test.add(4);
    return test.add(4) == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    inOrder() {
      if (!this.root) {
        return null;
      }
      var result = new Array();
      function traverseInOrder(node) {
        node.left && traverseInOrder(node.left);
        result.push(node.value);
        node.right && traverseInOrder(node.right);
      }
      traverseInOrder(this.root);
      return result;
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.add = function(element) {
    let current = this.root;
    if (!current) {
      this.root = new Node(element);
      return;
    } else {
      const searchTree = function(current) {
        if (current.value > element) {
          if (current.left) {
            return searchTree(current.left);
          } else {
            current.left = new Node(element);
            return;
          }
        } else if (current.value < element) {
          if (current.right) {
            return searchTree(current.right);
          } else {
            current.right = new Node(element);
            return;
          }
        } else {
          return null;
        }
      };
      return searchTree(current);
    }
  };
}
```