id: 587d8255367417b2b2512c73
title: Koristite spread operator i notacije za integraciju ES5 Set()
challengeType: 1
forumTopicId: 301720
dashedName: use-spread-and-notes-for-es5-set-integration
---

# --description--

Da li se sećate ES6 spread operatora `...`?

`...` može uzeti iterabilne vrednosti u ES6 i pretvoriti ih u niz podataka.

Napravićemo Set, a zatim ćemo istražiti spread operaciju.

```js
var set = new Set([1,2,3]);
var setToArr = [...set]
console.log(setToArr) // returns [ 1, 2, 3 ]
```

# --instructions--

U ovom vežbanju, prosleđivaćemo Set objekat funkciji `checkSet`. Funkcija treba da vrati niz podataka sa vrednostima iz Set-a.

Sada ste uspešno naučili kako koristiti ES6 `Set()` objekat, odlično!

# --hints--

`checkSet(new Set([1,2,3,4,5,6,7])` bi trebalo da vrati `[1, 2, 3, 4, 5, 6, 7]`.

```js
var test = checkSet(new Set([1, 2, 3, 4, 5, 6, 7]));
assert.deepEqual(test, [1, 2, 3, 4, 5, 6, 7]);
```

# --seed--

## --seed-contents--

```js
function checkSet(set){
   // Only change code below this line

   // Only change code above this line
}
```

# --solutions--

```js
function checkSet(set){
return [...set];}
```