id: 587d8257367417b2b2512c7e
title: Koristite pretragu dubine u binarnom stablu za pretragu
challengeType: 1
forumTopicId: 301719
dashedName: use-depth-first-search-in-a-binary-search-tree
---

# --description--

Znamo kako da pronađemo određenu vrednost u binarnom stablu za pretragu. Ali šta ako želimo samo da istražimo celo stablo? Ili šta ako nemamo organizovano stablo i samo moramo da pronađemo vrednosti? Ovde ćemo predstaviti neke metode prolaženja kroz stablo koje se mogu koristiti za istraživanje struktura podataka u obliku stabla. Prva je pretraga dubine (Depth-First Search). U pretrazi dubine, manje stablo se istražuje što je moguće dublje pre nego što pretraga nastavi na drugo manje stablo. Postoje tri načina da to uradimo: Inorder: Počnite pretragu od najlevšeg čvora i završite na najdesnjem čvoru. Preorder: Istražite sve korene pre listova. Postorder: Istražite sve listove pre korena. Kao što možete da predvidite, možete izabrati različite metode pretraživaња u zavisnosti od vrste podataka koje vaše stablo skladišti i šta tražite. Za binarno stablo za pretragu, inorder prolaženje vraća čvorove u sortirano redosled.

# --instructions--

Ovde ćemo kreirati ova tri načina pretraživaња na našem binarnom stablu za pretragu. Pretraga dubine je rekurentna operacija koja nastavlja da istražuje sve manje stablo dok god postoje čvorovi-deti. Čim savladate ovaj osnovni koncept, možete samo promeniti redosled posete čvorovima i manjim stablima kako biste proizveli jednu od tri pretrage navedene iznad. Na primer, u postorder pretrazi želeli bismo da rekurentno idemo do listnog čvora pre nego što počnemo sa vraćanjem bilo kog samog čvora, dok u preorder pretrazi želimo prvo vratiti čvor, a zatim nastaviti sa rekurenškim prolaženjem ispod stabla. Implementirajte metode `inorder`, `preorder` i `postorder` na našem stablu. Svaka metoda treba da vrati niz vrednosti koji predstavljaju prolaženje kroz stablo. Uverite se da vraćate cele brojne vrednosti za svaki čvor u nizu, a ne sam čvor. Na kraju, vratite `null` ako je stablo prazno.

# --hints--

Struktura podataka `BinarySearchTree` mora postojati.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    }
    return typeof test == 'object';
  })()
);
```

Binarno stablo za pretragu mora imati metodu pod nazivom `inorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.inorder == 'function';
  })()
);
```

Binarno stablo za pretragu mora imati metodu pod nazivom `preorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.preorder == 'function';
  })()
);
```

Binarno stablo za pretragu mora imati metodu pod nazivom `postorder`.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    return typeof test.postorder == 'function';
  })()
);
```

Metoda `inorder` mora vratiti niz vrednosti čvorova dobijenih inorder prolaženjem.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.inorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.inorder().join('') == '012345678910';
  })()
);
```

Metoda `preorder` mora vratiti niz vrednosti čvorova dobijenih preorder prolaženjem.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.preorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.preorder().join('') == '710325469810';
  })()
);
```

Metoda `postorder` mora vratiti niz vrednosti čvorova dobijenih postorder prolaženjem.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.postorder !== 'function') {
      return false;
    }
    test.add(7);
    test.add(1);
    test.add(9);
    test.add(0);
    test.add(3);
    test.add(8);
    test.add(10);
    test.add(2);
    test.add(5);
    test.add(4);
    test.add(6);
    return test.postorder().join('') == '024653181097';
  })()
);
```

Metoda `inorder` mora vratiti `null` za prazno stablo.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.inorder !== 'function') {
      return false;
    }
    return test.inorder() == null;
  })()
);
```

Metoda `preorder` mora vratiti `null` za prazno stablo.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.preorder !== 'function') {
      return false;
    }
    return test.preorder() == null;
  })()
);
```

Metoda `postorder` mora vratiti `null` za prazno stablo.

```js
assert(
  (function () {
    var test = false;
    if (typeof BinarySearchTree !== 'undefined') {
      test = new BinarySearchTree();
    } else {
      return false;
    }
    if (typeof test.postorder !== 'function') {
      return false;
    }
    return test.postorder() == null;
  })()
);
```

# --before-each--

```js
BinarySearchTree.prototype = Object.assign(
  BinarySearchTree.prototype,
  {
    add: function(value) {
      function searchTree(node) {
        if (value < node.value) {
          if (node.left == null) {
            node.left = new Node(value);
            return;
          } else if (node.left != null) {
            return searchTree(node.left);
          }
        } else if (value > node.value) {
          if (node.right == null) {
            node.right = new Node(value);
            return;
          } else if (node.right != null) {
            return searchTree(node.right);
          }
        } else {
          return null;
        }
      }

      var node = this.root;
      if (node == null) {
        this.root = new Node(value);
        return;
      } else {
        return searchTree(node);
      }
    }
  }
);
```

# --seed--

## --seed-contents--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  // Only change code below this line
  
  // Only change code above this line
}
```

# --solutions--

```js
var displayTree = tree => console.log(JSON.stringify(tree, null, 2));
function Node(value) {
  this.value = value;
  this.left = null;
  this.right = null;
}
function BinarySearchTree() {
  this.root = null;
  this.result = [];

  this.inorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    if (node.left) this.inorder(node.left);
    this.result.push(node.value);
    if (node.right) this.inorder(node.right);
    return this.result;
  };
  this.preorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    this.result.push(node.value);
    if (node.left) this.preorder(node.left);
    if (node.right) this.preorder(node.right);
    return this.result;
  };
  this.postorder = function(node) {
    if (!node) node = this.root;
    if (!node) return null;

    if (node.left) this.postorder(node.left);
    if (node.right) this.postorder(node.right);
    this.result.push(node.value);

    return this.result;
  };
}
```