id: 587d8255367417b2b2512c72
title: Koristi .has i .size na ES6 Set
challengeType: 1
forumTopicId: 301717
dashedName: use--has-and--size-on-an-es6-set
---

# --description--

Istražićemo metode .has i .size dostupne na ES6 Set objekat.

Prvo, kreiraj ES6 Set

```js
var set = new Set([1,2,3]);
```

Metoda .has proverava da li vrednost postoji unutar seta.

```js
var hasTwo = set.has(2);
```

Metoda .size vraća celobrojnu vrednost koja predstavlja veličinu Seta

```js
var howBig = set.size;
```

# --instructions--

U ovom vežlaji, ćemo iterirati niz podataka i vrednosti za funkciju checkSet(). Tvoja funkcija treba da kreira ES6 set iz niza podataka kao argumenta. Proveri da li set sadrži vrednost argumenta. Pronađi veličinu seta. I vrati te dve vrednosti u nizu podataka.

# --hints--

`checkSet([4, 5, 6], 3)` bi trebalo da vrati [ false, 3 ]

```js
var test = checkSet([4, 5, 6], 3);
assert.deepEqual(test, [false, 3]);
```

# --seed--

## --seed-contents--

```js
function checkSet(arrToBeSet, checkValue){

   // Only change code below this line

   // Only change code above this line

}
```

# --solutions--

```js
function checkSet(arrToBeSet, checkValue){
var set = new Set(arrToBeSet);
var result = [
set.has(checkValue),
set.size
];
return result;
}
```