id: 587d8251367417b2b2512c63
title: Uklanjanje elemenata iz spojenog lista
challengeType: 1
forumTopicId: 301712
dashedName: remove-elements-from-a-linked-list
---

# --description--

Sljedeća ključna metoda koju bilo koja implementacija spojenog lista će zahtijevati je metoda `remove`. Ova metoda bi trebala uzeti element koji želimo ukloniti kao argument, a zatim potražiti listu kako bi pronašla i uklonila čvor koji sadrži taj element.

Svaki put kada uklanjamo čvor iz spojenog lista, važno je osigurati da ne ostavimo drugi dio liste bez veze. Zapamtite da svojstvo `next` svakog čvora pokazuje sljedeći čvor u listi. Ako uklanjamo središnji element, na primjer, želimo osigurati da postoji veza od svojstva `next` prethodnog čvora tog elementa do svojstva `next` središnjeg elementa (što je sljedeći čvor u listi!)

Ovo može izgledati vrlo zbunjujuće, pa se vratimo na primjer konvejkionog lanca kako bismo imali jasniju vizualizaciju koncepta. Stojite u konvejkciji, a osoba ispred vas napušta liniju. Osoba koja je napustila liniju više nema ruke preko nikoga na liniji—i vi više nemate ruke preko osobe koja je napustila. Pomjerate se naprijed i stavljate ruke preko osobe koja slijedi u redu.

Ako element koji želimo ukloniti je `head` element, mi ažuriramo `head` da bude drugi čvor spojenog lista.

# --instructions--

Napišite metodu `remove` koja uzima element i uklanja ga iz spojenog lista.

**Napomena:** Broj `length` liste bi trebao smanjiti za jedan svaki put kada se element ukloni iz spojenog lista.

# --hints--

Vaša klasa `LinkedList` bi trebala imati metodu `remove`.

```js
assert(
  (function () {
    var test = new LinkedList();
    return typeof test.remove === 'function';
  })()
);
```

Vaša metoda `remove` bi trebala ažurirati `head` na drugi čvor kada se uklanja prvi čvor.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.remove('cat');
    return test.head().element === 'dog';
  })()
);
```

Vaša metoda `remove` bi trebala smanjiti `length` spojenog lista za jedan svaki put kada se uklanja čvor.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('hamster');
    test.remove('cat');
    test.remove('fish');
    return test.size() === 2;
  })()
);
```

Vaša metoda `remove` bi trebala ažurirati referencu prethodnog čvora uklonjenog čvora na referencu `next` uklonjenog čvora.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('snake');
    test.add('kitten');
    test.remove('snake');
    return test.head().next.next.element === 'kitten';
  })()
);
```

Vaša metoda `remove` ne bi smjela mijenjati spojenu listu ako element nije u spojenoj listi.

```js
assert(
  (function () {
    var test = new LinkedList();
    test.add('cat');
    test.add('dog');
    test.add('kitten');
    test.remove('elephant');
    return (
      JSON.stringify(test.head()) ===
      '{"element":"cat","next":{"element":"dog","next":{"element":"kitten","next":null}}}'
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
      var currentNode = head;

      while(currentNode.next){
        currentNode  = currentNode.next;
      }

      currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    // Only change code below this line

    // Only change code above this line
  };
}
```

# --solutions--

```js
function LinkedList() {
  var length = 0;
  var head = null;

  var Node = function(element){
    this.element = element;
    this.next = null;
  };

  this.size = function(){
    return length;
  };

  this.head = function(){
    return head;
  };

  this.add = function(element){
    var node = new Node(element);
    if(head === null){
        head = node;
    } else {
        var currentNode = head;

        while(currentNode.next){
            currentNode  = currentNode.next;
        }

        currentNode.next = node;
    }

    length++;
  };

  this.remove = function(element){
    if (head === null) {
      return;
    }
    var previous;
    var currentNode = head;

    while (currentNode.next !== null && currentNode.element !== element) {
      previous = currentNode;
      currentNode = currentNode.next;
    }
    
    if (currentNode.next === null && currentNode.element !== element) {
      return;
    }
    else if (previous) {
      previous.next = currentNode.next;
    } else {
      head = currentNode.next;
    }

    length--;
  };
} 
```