id: 587d8250367417b2b2512c5f
title: Kreiraj klasu steka
challengeType: 1
forumTopicId: 301633
dashedName: create-a-stack-class
---

# --description--

U prethodnom delu, razgovarali smo o tome šta je stek i kako možemo koristiti niz podataka da ga predstavimo. U ovom delu, kreiraćemo našu klasu steka. Iako možete koristiti nizove podataka za kreiranje steka, ponekad je bolje smanjiti nivo kontrole koji imamo nad našim stekovima. Pored metoda `push` i `pop`, steci imaju druge važne metode. Dodajmo metode `peek`, `isEmpty` i `clear` u našu klasu steka.

# --instructions--

Napišite metodu `push` koja dodaje element na vrh steka, metodu `pop` koja uklanja i vraća element sa vrha steka, metodu `peek` koja proverava element na vrhu steka, metodu `isEmpty` koja proverava da li je stek prazan, i metodu `clear` koja uklanja sve elemente iz steka. Iako steci obično nemaju ovo, dodali smo pomoćnu metodu `print` koja ispisuje kolekciju u konzolu.

# --hints--

Vaša klasa `Stack` mora imati metodu `push`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.push === 'function';
  })()
);
```

Vaša klasa `Stack` mora imati metodu `pop`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.pop === 'function';
  })()
);
```

Vaša klasa `Stack` mora imati metodu `peek`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.peek === 'function';
  })()
);
```

Vaša klasa `Stack` mora imati metodu `isEmpty`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.isEmpty === 'function';
  })()
);
```

Vaša klasa `Stack` mora imati metodu `clear`.

```js
assert(
  (function () {
    var test = new Stack();
    return typeof test.clear === 'function';
  })()
);
```

Metoda `peek` treba da vrati element sa vrha steka

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.peek() === 'CS50' && test.peek() === 'CS50';
  })()
);
```

Metoda `pop` treba da ukloni i vrati element sa vrha steka

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    return test.pop() === 'CS50' && test.pop() === 'CS61';
  })()
);
```

Metoda `isEmpty` treba da vrati `true` ako stek nema nikakve elemente

```js
assert(
  (function () {
    var test = new Stack();
    return test.isEmpty();
  })()
);
```

Metoda `clear` treba da ukloni sve elemente iz steka

```js
assert(
  (function () {
    var test = new Stack();
    test.push('CS61');
    test.push('CS50');
    test.clear();
    return test.isEmpty();
  })()
);
```

# --seed--

## --seed-contents--

```js
function Stack() {
  var collection = [];
  this.print = function() {
    console.log(collection);
  };
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
class Stack {
  constructor() {
    this.collection = [];
  }
  print() {
    console.log(this.collection);
  }
  push(val) {
    this.collection.push(val);
  }
  pop() {
    return this.collection.pop();
  }
  peek() {
    return this.collection[this.collection.length - 1];
  }
  isEmpty() {
    return this.collection.length === 0;
  }
  clear() {
    return (this.collection.length = 0);
  }
}
```