---
id: 680a5e7f3d47715196b12ca0
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Sada kada imate ceo pomereni alfabet, uklonite poziv `print(shifted_alphabet)`.

# --hints--

You should remove `print(shifted_alphabet)` from your code.

```js
({ test: () => assert(runPython(`not _Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
print(shifted_alphabet)
--fcc-editable-region--
```
