---
id: 680a431692a5f863f726c438
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

U ovom radionici, kreći ćete da izgradite *šifru Cezara*. Ovo je jedna od najjednostavnijih tehnika za šifrovanje teksta, koja se sastoji od zamene svakog slova običnog teksta slovom koje se nalazi na fiksnom broju pozicija dalje u alfabetu. Na primer, sa pomeranjem od 5, `a` bi bio zamenjen sa `f`, `b` sa `g` i tako dalje.

Da biste implementirali ovu šifru, moraćete da kreirate novu verziju vašeg alfabeta koja počinje na poziciji koju ukazuje pomeranje. Kao što ste naučili u prethodnoj lekciji, možete izvući deo niza korišćenjem *string slicing*-a:

```py
fcc = 'freeCodeCamp'
print(fcc[8:]) # Camp
```

Kreirajte varijablu pod nazivom `shifted_alphabet` i koristite *slicing* sintaksu da joj dodelite deo `alphabet` koji počinje na indeksu `shift`. Zatim, pozovite ugrađenu funkciju `print()` da odštampa `shifted_alphabet` na terminal i pogledajte rezultat.

# --hints--

You should declare a new variable named `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shifted_alphabet")`)) })

```

You should assign `alphabet[shift:]` to `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shifted_alphabet").is_equivalent("shifted_alphabet = alphabet[shift:]")`)) })
```

You should call the `print()` function passing in `shifted_alphabet` as an argument.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--

--fcc-editable-region--
```
