---
id: 680a836235461fce1ff53b3f
title: Korak 8
challengeType: 20
dashedName: step-8
---

# --description--

Metoda `translate()` prima kao argument tablicu prevoda generisanu od strane `str.maketrans()`. Poziva se na string i vraća kopiju originalnog stringa gde su karakteri zamenjeni na osnovu tablice prevoda:

```py
t = str.maketrans('lk', 'br')
sentence = 'The tent gave in to the leaks.'

print(sentence.translate(t))
# Output: The tent gave in to the bears.
```

Pozovite metodu `translate()` na `text`, prosleđujući `translation_table` kao argument i dodeljujte rezultat promenljivoj nazvanoj `encrypted_text`.

# --hints--

You should have a variable named `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

You should assign `text.translate(translation_table)` to `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
--fcc-editable-region--

--fcc-editable-region--
```
