---
id: 680a549f60f5a3fedfc0edd9
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Kao što možete da vidite iz izlaza, pomereni alfabet počinje slovom `f` jer `shift` ima vrednost `5`. Ali sada, prvih pet slova alfabeta – `a`, `b`, `c`, `d` i `e` – nedostaje u pomerenom alfabetu, pa ćete morati da ih dodate na kraj pomerenog alfabeta. Možete koristiti operator `+` da spojite nedostajući deo alfabeta.

Izmenite postojeće dodeljivanje varijable `shifted_alphabet`: koristite sintaksu sečenja (slicing) da izvučete nedostajući prvi deo od `alphabet` i spojite ga sa `alphabet[shift:]`.

Kao podsetnik, `sentence[start:stop]` vraća karaktere iz `sentence` od pozicije `start` (uključeno) do `stop` (isključeno).

# --hints--

You should not directly modify the `alphabet` variable.

```js
({ test: () => runPython(`
  assert alphabet == "abcdefghijklmnopqrstuvwxyz"
`) })
```

You should assign the concatenation of `alphabet[shift:]` and the missing first portion of `alphabet` to the `shifted_alphabet` variable on the same line it was declared. Use `shift` to specify where to stop the slicing.

```js
({ test: () => runPython(`
_shifted_alpha = _Node(_code).find_variable("shifted_alphabet")
_first = _shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[:shift]")
_second =_shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[0:shift]")
assert _first or _second
`) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--
shifted_alphabet = alphabet[shift:]
--fcc-editable-region--
print(shifted_alphabet)
```
