---
id: 680a3d4c065557260de2ff33
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U ovoj radionici, kreiraćete program nazvan Cezarov šifar koji je metoda enkripcije koja pomera slova u alfabetu kako bi kodirala poruke.

Počnite kreiranjem varijable nazvane `shift` i dodeljujte vrednost `5` vašoj novoj varijabli.

# --hints--

You should declare a variable called `shift`. Pay attention to place the variable name at the beginning of the line.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

You should assign the value `5` to your `shift` variable.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
