---
id: 680a7e200400089967b8b1f4
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Metoda `str.maketrans()` prima dva niza jednakih dužine i vraća tabelu prevoda koja mapira svaki karakter prvog niza sa odgovarajućim karakterom drugog niza. Svaki karakter u tabeli prevoda pohranjen je kao Unicode ordinal, broj koji jedinstveno identifikuje karakter.

```py
lower_chars = 'abc'
upper_chars = 'ABC'

table = str.maketrans(lower_chars, upper_chars)
print(table) # {97: 65, 98: 66, 99: 67}
```

Vi ćete saznati više o tipu strukture koju vraća `str.maketrans()` kasnije u kurikulumu. Trenutno, kreirajte tabelu prevoda koja mapira karaktere iz `alphabet` na karaktere iz `shifted_alphabet` i dodelite je promenljivoj pod nazivom `translation_table`.

# --hints--

You should declare a variable named `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("translation_table")`)) })
```

You should call `str.maketrans()` passing in `alphabet` and `shifted_alphabet` as the arguments, and assign the result to `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("translation_table").is_equivalent("translation_table = str.maketrans(alphabet, shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--

--fcc-editable-region--
```
