---
id: 6810891bd6de8a87a833df42
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Kao što vidite iz izlaza, poruka je enkriptovana. Sledeći korak biće činizanje vašeg koda ponovo upotrebljivim u slučaju da želite da enkriptujete različite poruke.

Za to, moraćete da kreirate funkciju. Kao podsetnik, evo kako da kreirate funkciju nazvanu `spam` koja štampa `Spam!` na terminal:

```py
def spam():
    print('Spam!')
```

Kreirajte funkciju nazvanu `caesar`. Stavite sav svoj postojeći kod unutar tela funkcije. Obratite pažnju da zadržite isti nivo uvlačenja za sve linije unutar tela funkcije.

# --hints--

You should have a function named `caesar`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("caesar")`)) })
```

You should move all the code you wrote so far within the `caesar` function body. Make sure you keep the same indentation level for all of the lines within the function body...

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_body().is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'\\nshift = 5\\nshifted_alphabet = alphabet[shift:] + alphabet[:shift]\\ntranslation_table = str.maketrans(alphabet, shifted_alphabet)\\ntext = 'hello world'\\nencrypted_text = text.translate(translation_table)\\nprint(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
encrypted_text = text.translate(translation_table)
print(encrypted_text)
--fcc-editable-region--
```
