---
id: 6818b80519e6c6c47a8bb6e8
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Kao što ste naučili u prethodnoj lekciji, funkcije mogu imati parametre, koji su varijable koje se mogu referencirati unutar funkcije. Evo funkcije `sum` sa dva parametra, `num1` i `num2`.

```py
def sum(num1, num2):
    print(num1 + num2)
```

Poruka za šifrovanje i pomeraj su i dalje hardkodovane u vašoj funkciji. Dajte svojoj funkciji dva parametra pod nazivima `text` i `shift`, i obrišite deklaracije obe varijable, `text` i `shift`, iz tela funkcije `caesar`.

# --hints--

Your `caesar` function should have two parameters, `text` and `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift")`)) })
```

You should not have a `text` variable declared within the `caesar` function body.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("text")`)) })
```

You should not have a `shift` variable declared within the `caesar` function body.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("shift")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shift = 5
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    text = 'hello world'
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
```
