---
id: 6818fcb250f34e62cab32c39
title: Korak 15
challengeType: 20
dashedName: step-15
---

# --description--

Možda ste primetili da, iako je poruka enkriptovana, velika slova su ostala netaknuta. To se dešava zato što tabela prevoda ne uključuje velika slova.

Da biste to popravili, moraćete da modifikujete nizove prosleđene vašem pozivu `str.maketrans()` koristeći metod `upper()`:

```py
greet = "Hello"
print(greet.upper()) # HELLO
```

Ažurirajte vaš poziv `str.maketrans()` tako što ćete svakom argumentu konkatikovati veliku verziju samog tog argumenta.

# --hints--

You should modify the arguments of your `str.maketrans()` call by concatenating to each argument the uppercase version of the argument itself.

```js
({ test: () => assert(runPython(`caesar('This Is A TEST', 5) == 'Ymnx Nx F YJXY'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
