---
id: 6818ed1adc249bb91d373729
title: Korak 14
challengeType: 20
dashedName: step-14
---

# --description--

U prethodnim lekcijama, naučili ste da je naredba `return` korišćena za vraćanje vrednosti iz funkcije, tako da je možete koristiti negde drugde u svom kodu:

```py
def spam():
    return 'Spam!'
```

Uklonite poziv `print(encrypted_text)` iz vaše funkcije. Zatim, obrišite deklaraciju varijable `encrypted_text` i direktno vratite `text.translate(translation_table)` iz svoje funkcije.

# --hints--

You should not have `print(encrypted_text)` within the `caesar` function.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_call("print(encrypted_text)")`)) })
```

You should not have an `encrypted_text` variable within the `caesar` function.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_variable("encrypted_text")`)) })
```

Your `caesar` function should return `text.translate(translation_table)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_return("text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
