---
id: 6819ca464db219ab86f37dc4
title: Korak 20
challengeType: 20
dashedName: step-20
---

# --description--

Kao što ste već verifikovali, pomak prosleđen za enkripciju teksta treba da bude pozitivan ceo broj. Međutim, ne sme da premaši `25` (poslednji indeks `alphabet`).

Dodajte drugi uslov u izjavu `if` koji proverava da li je `shift` veći od `25`. Zapamtite da se logička OR operacija u Pythonu implementira pomoću operatora `or`.

Takođe, ažurirajte vraćenu poruku na `Shift must be an integer between 1 and 25.`.

# --hints--

You should use the `or` operator to add a second condition after `shift < 1` checking if `shift` is greater than `25`. You should not modify the existing condition.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift > 25", "25 < shift", "shift >= 26", "26 =< shift"]
assert any(_cond.is_equivalent(f"shift < 1 or {option}") for option in _options)
`) })
```

Your second `if` statement should return `Shift must be an integer between 1 and 25.`

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be an integer between 1 and 25.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
    
--fcc-editable-region--
    if shift < 1:
        return 'Shift must be a positive integer.'
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
