---
id: 6819bc0637b80256a33adccc
title: Korak 16
challengeType: 20
dashedName: step-16
---

# --description--

Sada kada ste implementirali osnovne funkcionalnosti cifara, vreme je da dodate validaciju. Za to, biće vam potreban `if` statement.

Na početku tela vaše funkcije, kreirajte `if` statement. Trenutno, koristite `True` kao uslov, a unutar tela `if` statement-a vratite string `Shift must be an integer value.`.

# --hints--

You should have an `if` statement inside your `caesar` function.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0]`)) })
```

Your `if` statement should use `True` as its condition.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("True")`)) })
```

You should return the string `Shift must be an integer value.` from your `if` statement.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_bodies()[0].has_return("'Shift must be an integer value.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
