---
id: 6819cf38880d8cc9b9f6af7a
title: Korak 21
challengeType: 20
dashedName: step-21
---

# --description--

Python vam omogućava da odredite podrazumevanu vrednost za parametre u funkciji, stvarajući funkciju koja se može pozvati sa manje ili bez argumenata. Evo kako kreirati funkciju sa parametrom `name` koji ima podrazumevanu vrednost:

```py
def greet(name='Polly'):
    return 'Hello ' + name
    
print(greet()) # Hello Polly
```

Dodajte treći parametar pod nazivom `encrypt` vašoj funkciji i dodelite mu podrazumevanu vrednost od `True`.

# --hints--

You should add a third parameter named `encrypt` with a default value of `True` to your function. Do not modify the existing parameters.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift, encrypt=True")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
--fcc-editable-region--

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
