---
id: 6819c5c9faf2548e556163d5
title: Korak 18
challengeType: 20
dashedName: step-18
---

# --description--

Kao što ste se setili iz prethodnih lekcija, operator `not` je unarni logički operator koji negira izraz:

```py
print(not True) # False
print(not False) # True
```

Koristite operator `not` da popravite uslov u izjavi `if` tako da vaša funkcija vrati `Shift must be an integer value.` kada `shift` nije ceo broj.

# --hints--

You should use the `not` operator to negate `isinstance(shift, int)` in your `if` statement.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if isinstance(shift, int):
--fcc-editable-region--
        return 'Shift must be an integer value.'
    
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
