---
id: 681a1d60e037127e7c8e0230
title: Korak 22
challengeType: 20
dashedName: step-22
---

# --description--

Koristićete parametar dodan u prethodnom koraku da utvrdite da li funkcija treba da enkriptuje tekst koji joj je prosleđen (podrazumevano ponašanje, `encrypt=True`), ili da dekriptuje šifrovsan poruku.

Kreirajte izjavu `if` koja proverava da li `encrypt` nije istinito. Unutar nove `if` izjave, postavite `shift` na `- shift`. Ovo je neophodno kako bi se pomeranje omogućilo u suprotnom smeru u odnosu na proces enkripcije.

# --hints--

You should set `shift` to `- shift` when `encrypt` is not truthy.

```js
({ test: () => assert(runPython(`caesar('Ymnx Nx F Yjxy', 5, False) == 'This Is A Test'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    
--fcc-editable-region--
    
--fcc-editable-region--
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
