---
id: 681a2a6cf95e82ca30866364
title: Korak 23
challengeType: 20
dashedName: step-23
---

# --description--

Deklarišite dve funkcije nazvane `encrypt` i `decrypt`, obe sa parametrima `text` i `shift`.

Koristićete `encrypt` za proces šifrovanja, a `decrypt` za dešifrovanje, označavanjem dve akcije opisnim imenom.

Vratite poziv `caesar` koji prosleđuje `text` i `shift` iz obe vaše nove funkcije, ali obavezno prosledite i `False` kao treći argument za funkciju `decrypt`.

# --hints--

You should have a function named `encrypt` with two parameters, `text` and `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_args("text, shift")`)) })
```

Your `encrypt` function should return `caesar(text, shift)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_return("caesar(text, shift)")`)) })
```

You should have a function named `decrypt` with two parameters, `text` and `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("decrypt").has_args("text, shift")`)) })
```

Your `decrypt` function should return `caesar(text, shift, False)`.

```js
({ test: () => runPython(`
_n = _Node(_code).find_function("decrypt")
assert _n.has_return("caesar(text, shift, encrypt=False)") or _n.has_return("caesar(text, shift, False)")
`) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

--fcc-editable-region--

--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
