---
id: 681a44f4c3235f7d8f428545
title: Korak 25
challengeType: 20
dashedName: step-25
---

# --description--

Sada ćete testirati funkciju `decrypt`. Zamenite vrednost dodeljenu `encrypted_text` sa sledećim nizom karaktera, koji predstavlja poruku za dešifrovanje: `Pbhentr vf sbhaq va hayvxryl cynprf.`

Zatim, deklarišite varijablu pod imenom `decrypted_text` i dodelite joj poziv na `decrypt` sa `encrypted_text` kao prvim argumentom i pomerajanjem (shift) od `13` kao drugim argumentom.

Konačno, ispišite `decrypted_text` na terminal. Time je Cesarov šifar završen.

# --hints--

You should assign `Pbhentr vf sbhaq va hayvxryl cynprf.` to `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'")`)) })
```

You should have a variable named `decrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("decrypted_text")`)) })
```

You should assign `decrypt(encrypted_text, 13)` to `decrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("decrypted_text").is_equivalent("decrypted_text = decrypt(encrypted_text, 13)")`)) })
```

You should print `decrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(decrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)
    
def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

--fcc-editable-region--
encrypted_text = encrypt('freeCodeCamp', 3)
print(encrypted_text)
--fcc-editable-region--
```

# --solutions--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)
    
def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'
decrypted_text = decrypt(encrypted_text, 13)
print(decrypted_text)
```
