---
id: 684cc59eed2706fa50e8933c
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Na ovom radionici, naučićete kako da počnete sa osnovama Tailwind CSS-a gradeći komponentu poruke o grešci.

Za prve korake, fokusiraćete se na kreiranje početnog HTML markupa.

Unutar elementa `body`, dodajte element `div`. Zatim unutar tog elementa `div`, dodajte element `p`.

# --hints--

You should have a `div` element.

```js
assert.exists(document.querySelector("div"));
```

You should have a `p` element.

```js
assert.exists(document.querySelector("p"));
```

Your `p` element should be inside of the `div` element.

```js
assert.exists(document.querySelector("div p"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
