---
id: 0095d0567b6817114989a1dd
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovom radionici, implementiraćete algoritam najkraćeg puta u JavaScript-u. Napišaćete funkciju koja računa najkraći put između čvorova u težinom grafu i vraća kako udaljenosti tako i pređene staze.

Na primer, dati graf gde su gradovi povezani putevima različitih udaljenosti, algoritam će pronaći najkraću rutu od jednog grada do drugog. Ako želite da putujete iz Grada A u Grad D, algoritam bi mogao da utvrdi da je putovanje A ⇨ B ⇨ C ⇨ D (ukupno: 15km) kraće nego direktno putovanje A ⇨ D (20km).

Da biste počeli, definišite varijablu pod nazivom `INF` i dodelite joj vrednost `Infinity`. Kasnije ćete je koristiti da naznačite da ne postoji direktna veza između dva čvora.

# --hints--

You should define a variable named `INF`.

```js
assert.isDefined(INF);
```

Your `INF` variable should have the value `Infinity`.

```js
assert.strictEqual(INF, Infinity);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
