---
id: 69c918841987bf16c1d1bac1
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Sada ćete dodati glavnu petlju koja pokreće algoritam. Ona radi jednom za svaki čvor u grafu, birajući najbliži neposetačeni čvor svaki put.

Unutar funkcije `shortestPath`, dodajte petlju `for` koja radi `n` puta koristeći `i` kao varijablu petlje.

Na početku tela petlje, koristeći `let`, deklarišite dve varijable `minDistance` i `current` postavljene na `INF`, i `-1` respektivno. One će pratiti najbliži neposetačeni čvor pronađen u svakoj iteraciji.

# --hints--

You should have a `for` loop inside `shortestPath` that runs `n` times. Don't forget to set `i` with `let` inside it.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+i\s*=\s*0\s*;\s*i\s*<\s*n\s*;\s*i\+\+\s*\)/
);
```

You should have a `minDistance` variable set to `INF` inside the loop.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+minDistance\s*=\s*INF/
);
```

You should have a `current` variable set to `-1` inside the loop.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+current\s*=\s*-1/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
--fcc-editable-region--

--fcc-editable-region--
}
```
