---
id: 69c918841987bf16c1d1bac7
title: Korak 16
challengeType: 1
dashedName: step-16
---

# --description--

Sada kada je trenutni čvor označen kao posetao, morate pogledati sve njegove susede kako biste videli da li ih možete doći efikasnije kroz trenutni čvor.

Nakon `visited[current] = true`, dodajte petlju suseda `for` koja iterira kroz svaki `nodeNo` od `0` do `n`. Napomena da je ovo još jedna unutrašnja petlja na istom nivou kao i prva.

Unutar petlje, deklarišite varijablu pod imenom `distance` i dodelite joj `matrix[current][nodeNo]`. Ovo vam daje težinu ivice između trenutnog čvora i njegovog suseda.

# --hints--

You should add a neighbor `for` loop after marking the current node as visited.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(/
);
```

The neighbor `for` loop should iterate from `0` to `n`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(\s*(var|let)\s+\w+\s*=\s*0\s*;\s*\w+\s*<\s*n\s*;\s*\w+\+\+\s*\)/
);
```

You should have a `distance` variable assigned to `matrix[current][nodeNo]` inside the neighbor loop.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distance\s*=\s*matrix\s*\[\s*current\s*\]\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
