---
id: 69c918841987bf16c1d1babd
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Morate pratiti najkraće poznato rastojanje od početnog čvora do svakog drugog čvora. Da biste počeli, pretpostavićete da je svaki čvor beskonačno daleko.

U JavaScriptu, možete kreirati niz prethodno popunjen vrednošću koristeći `new Array(n).fill(value)`. Na primer, `new Array(3).fill(0)` kreira `[0, 0, 0]`.

Unutar funkcije `shortestPath`, kreirajte varijablu pod nazivom `distances` i dodelite joj `new Array(n).fill(INF)`.

# --hints--

You should create a variable named `distances` inside `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distances\s*=/
);
```

Your `distances` variable should be set to `new Array(n).fill(INF)`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*INF\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;

}
--fcc-editable-region--
```
