---
id: 69c918841987bf16c1d1baca
title: Korak 19
challengeType: 1
dashedName: step-19
---

# --description--

Kada pronađete kraći put do suseda, morate ažurirati i snimanu stazu za dolazak do njega.

U JavaScriptu, možete kreirati novi niz kombinujući postojeći niz i novi element koristeći operater širenja (spread operator):

```js
const newPath = [...existingArray, newElement];
```

Unutar ugnježdene `if` blokove, ažurirajte `paths[nodeNo]` tako da bude put do trenutnog čvora sa dodanim `nodeNo` na kraju koristeći sintaksu širenja (spread syntax) kao što je: `[...paths[current], nodeNo]`.

# --hints--

You should update `paths[nodeNo]` inside the nested `if` block.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*\[\s*\w+\s*\]\s*=/
);
```

`paths[nodeNo]` should be set to `[...paths[current], nodeNo]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*\[\s*\w+\s*\]\s*=\s*\[\s*\.\.\.\s*paths\s*\[\s*current\s*\]\s*,\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          --fcc-editable-region--

          --fcc-editable-region--
        }
      }
    }
  }
}
```
