---
id: 5cc0bd7a49b71cb96132e54c
title: Izgradite generator raspona brojeva
challengeType: 26
dashedName: build-a-range-of-numbers-generator
---

# --description--

U ovom laboratorijumu, kreiraćete funkciju `rangeOfNumbers` koja generiše niz brojeva unutar definisanog opsega.

Primeri:

- `rangeOfNumbers(3, 9)` vraća `[3, 4, 5, 6, 7, 8, 9]`
- `rangeOfNumbers(5, 5)` vraća `[5]`

Zahtevi:

- Koristite rekurziju (funkcija mora da pozove samu sebe)
- Nema petlja niti metoda za niz (`for`, `while`, `Array.from`, `forEach`, `map`, `filter`, `reduce`)

**Cilj:** Ispunite priče korisnika ispod i učinite da svi testovi prođu kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da kreirate funkciju pod nazivom `rangeOfNumbers` koja prima dva parametra: `startNum` i `endNum`.
1. Funkcija bi trebalo da vrati niz celih brojeva koji počinje brojem predstavljenim parametrom `startNum` i završava se brojem predstavljenim parametrom `endNum` (uključujući oba).
1. `startNum` će uvek biti manji ili jednak `endNum`.
1. Vaša funkcija mora koristiti rekurziju pozivanjem sebe. Ne sme da koristi bilo koju sintaksu petlje (`for`, `while`, `Array.from()`, ili funkcije višeg reda kao što su `forEach`, `map`, `filter`, ili `reduce`).
1. Funkcija bi trebalo da obradi osnovni slučaj gde je `startNum` jednako `endNum` vraćanjem niza koji sadrži samo taj jedan broj.
1. Za rekurzivni slučaj, funkcija bi trebalo da pozove samu sebe sa modifikovanim parametrima kako bi izgradila niz, a zatim doda trenutni broj rezultatu.
1. Funkcija ne sme da se oslanja na globalne varijable za keširanje ili izgradnju niza.

# --hints--

You should have a function named `rangeOfNumbers`.

```js
assert.isFunction(rangeOfNumbers);
```

The `rangeOfNumbers` function should take two parameters.

```js
assert.match(rangeOfNumbers.toString(), /rangeOfNumbers\s*\(\s*\w+\s*,\s*\w+\s*\)/);
```

Your function should return an array.

```js
assert(Array.isArray(rangeOfNumbers(5, 10)));
```

Your code should not use any loop syntax (`for`, `while`, `Array.from()` or higher order functions such as `forEach`, `map`, `filter`, or `reduce`).

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce|Array\.from/g)
);
```

`rangeOfNumbers` should use recursion (call itself) to solve this challenge.

```js
assert.isAtLeast(
  __helpers.removeJSComments(rangeOfNumbers.toString()).match(/rangeOfNumbers\s*\(.+\)/g)?.length,
  2
);
```

`rangeOfNumbers(1, 5)` should return `[1, 2, 3, 4, 5]`.

```js
assert.deepStrictEqual(rangeOfNumbers(1, 5), [1, 2, 3, 4, 5]);
```

`rangeOfNumbers(6, 9)` should return `[6, 7, 8, 9]`.

```js
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

`rangeOfNumbers(4, 4)` should return `[4]`.

```js
assert.deepStrictEqual(rangeOfNumbers(4, 4), [4]);
```

`rangeOfNumbers(10, 15)` should return `[10, 11, 12, 13, 14, 15]`.

```js
assert.deepStrictEqual(rangeOfNumbers(10, 15), [10, 11, 12, 13, 14, 15]);
```

`rangeOfNumbers(2, 8)` should return `[2, 3, 4, 5, 6, 7, 8]`.

```js
assert.deepStrictEqual(rangeOfNumbers(2, 8), [2, 3, 4, 5, 6, 7, 8]);
```

Global variables should not be used to cache the array.

```js
rangeOfNumbers(1, 3);
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function rangeOfNumbers(startNum, endNum) {
  if (endNum - startNum === 0) {
    return [startNum];
  } else {
    const numbers = rangeOfNumbers(startNum, endNum - 1);
    numbers.push(endNum);
    return numbers;
  }
}
```
