---
id: 694bf6eb530e19a9c48c59f1
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Sada, vaša adresa izgleda nepotpuno. Takođe želite dodati broj stana gde živi zaposleni, pa biste trebali da modifikujete varijablu.

Kada želite dodati sadržaj na kraj postojeće string varijable, možete koristiti operater **proširenog dodeljivanja** (augmented assignment) ,`+=`. Ovo je kraće od pisanja `var = var + 'new text'`. Na primer:

```py
greeting = 'Hello'
greeting += ' World'
print(greeting) # Hello World
```

Zapamtite da su stringovi nemenjivi (immutable), stoga ova operacija ne menja originalni string. Umesto toga kreira novi string i ponovo mu dodeljuje vrednost varijabli.

Koristite operater `+=` da dodate string `, Apartment 4B` vašoj varijabli `address`.

# --hints--

You should use the `+=` operator to add the string `, Apartment 4B` to the `address` variable.

```js
({ test: () => assert(runPython(`_Node(_code).find_aug_variable("address").is_equivalent("address += ', Apartment 4B'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
address = '123 Main Street'  
--fcc-editable-region--  

--fcc-editable-region-- 
print(address)  
```
