---
id: 694bf6eb530e19a9c48c59f5
title: Korak 16
challengeType: 20
dashedName: step-16
---

# --description--

Kada radite sa nizovima znakova, često ćete morati da izvučete određeni deo niza. To se zove **seckanje** (slicing).

Sintaksa je `string[start:stop]`, gde:

* `start` je indeks gde se seckanje počinje (**uključivo**).
* `stop` je indeks gde se seckanje završava (**isključivo**).

Na primer, ako je `text = 'Python'`, onda `text[0:2]` daje `'Py'`.

Definišite `employee_code` kao `'DEV-2026-JD-001'`. Nakon toga, kreirajte varijablu `department` i dodelite joj seckanje od `employee_code` sa indeksa `0` do `3`. Zatim ispišite `department` u terminal.

# --hints--
You should have a variable named `employee_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_code")`)) })
```

You should assign the string `DEV-2026-JD-001` to your `employee_code` variable.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_code").is_equivalent("employee_code = 'DEV-2026-JD-001'")`)) })
```

You should have a variable named `department`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("department")`)) })
```

You should slice the first three characters of `employee_code` and assign them to your `department` variable.

```js
({ test: () => runPython(`
dep = _Node(_code).find_variable("department")
assert dep.is_equivalent("department = employee_code[0:3]") or dep.is_equivalent("department = employee_code[:3]")
`) })
```

You should print `department` to the terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(department)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)

--fcc-editable-region--

--fcc-editable-region--
```
