---
id: 694bf6eb530e19a9c48c59f6
title: Korak 17
challengeType: 20
dashedName: step-17
---

# --description--

Možete izvući deo iz bilo kog dela niza karaktera, ne samo od početka. I to je korisno u mnogim slučajevima.

Kreirajte varijablu `year_code` i dodelite joj izrezak od `employee_code` sa indeksa `4` do `8`. Ovo će izdvojiti `2026`.

Zatim kreirajte varijablu `initials` i dodelite joj izrezak od `employee_code` sa indeksa `9` do `11`. Ovo će izdvojiti `JD`.

Konačno, ispišite obe varijable u terminal.

# --hints--

You should have a variable named `year_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("year_code")`)) })
```

You should slice `employee_code` from index `4` to `8` and assign the result to your `year_code` variable.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("year_code").is_equivalent("year_code = employee_code[4:8]")`)) })
```

You should have a variable named `initials`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("initials")`)) })
```

You should slice `employee_code` from index `9` to `11` and assign the result to your `initials` variable.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("initials").is_equivalent("initials = employee_code[9:11]")`)) })
```

You should print `year_code` to the terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(year_code)")`)) })
```

You should print `initials` to the terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(initials)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department) 
--fcc-editable-region--  

--fcc-editable-region--  
```
