---
id: 694bf6ea530e19a9c48c59f0
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

Ako spojite dva niza kao što su `'John' + 'Doe'`, rezultat je `'JohnDoe'` bez razmaka. Da biste ovo popravili, morate spojiti niz koji sadrži razmak (`' '`) između njih.

Ažurirajte svoju varijablu `full_name` tako da spoji `first_name`, razmak i `last_name`.

# --hints--

You should add a space (`' '`) between `first_name` and `last_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + ' ' + last_name")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--
full_name = first_name + last_name
--fcc-editable-region--
print(full_name)
```
