---
id: 694bf6eb530e19a9c48c59f4
title: Korak 13
challengeType: 20
dashedName: step-13
---

# --description--

Konkateniranje mnogo stringova korišćenjem `+` i konvertovanje brojeva korišćenjem `str()` može postati neuredno i teško za čitanje.

Python 3.6 je uveo **f-strings** da reši ovo. Dodavanjem slova `f` pre otvarne navodnika, možete staviti varijable i izraze unutar polja za zamenu predstavljena vitičastim zagradama `{}`. Na primer:

```python
name = 'John'
print(f'Hello {name}') # Output: Hello John
```

Kreirajte varijablu `employee_card` i dodelite joj f-string koji prikazuje `Employee:` praćen razmakom i vrednost varijable `full_name`.

# --hints--

You should have a variable named `employee_card`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_card")`)) })
```

Your `employee_card` variable should have the value `f'Employee: {full_name}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name}'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

--fcc-editable-region--
```
