---
id: 694bf3af05d10fd89dab05cc
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

U Python, možete spojiti niz znakova koristeći operator `+`. Ovo se zove konkatenacija. Evo primera:

```python
greeting = 'Hello' + 'World'
print(greeting) # Output: HelloWorld
```

Kreiraj parametar `full_name` spajanjem `first_name` sa `last_name`. Zatim ispisi `full_name`.

# --hints--

Treba da imaš parametar koji se zove `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("full_name")`)) })
```

Trebalo bi da postavite vrednost niza karaktera kreiranog spajanjem `first_name` i `last_name` u vaš parametar `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + last_name")`)) })
```

Trebalo bi da ispišete parametar `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(full_name)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--

--fcc-editable-region--
```
