---
id: 694bf6eb530e19a9c48c59f7
title: Korak 18
challengeType: 20
dashedName: step-18
---

# --description--

Takođe možete koristiti negativne brojeve za sečenje (slicing) od kraja niza karaktera. Na primer, `-1` se odnosi na poslednji karakter, `-2` se odnosi na predposlednji karakter i tako dalje.

Da biste dobili poslednja tri karaktera niza, možete koristiti `string[-3:]`. Primećujte kako je parametar `stop` izostavljen nakon dvotačke. To znači da se sečenje treba nastaviti sve do granice niza.

Kreirajte varijablu pod nazivom `last_three`. Koristite negativno indeksiranje da izvučete poslednja tri karaktera iz `employee_code` (koji predstavljaju ID broj). Na kraju, ispišite `last_three` u terminal.

Time sa tim, radionica za kartice zaposlenih je završena.

# --hints--

You should have a variable named `last_three`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("last_three")`)) })
```

You should slice the last three characters of `employee_code` using negative indexing and assign them your `last_three` variable.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_three").is_equivalent("last_three = employee_code[-3:]")`)) })
```

You should print `last_three` to the terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_three)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
print(year_code)
initials = employee_code[9:11]
print(initials)
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
initials = employee_code[9:11]
print(year_code)
print(initials)
last_three = employee_code[-3:]
print(last_three)
```
