---
id: 694bf6eb530e19a9c48c59f3
title: Korak 12
challengeType: 20
dashedName: step-12
---

# --description--

Sada ćete koristiti funkciju `str()` još jednom. Kao i sa godinama, morate konvertovati bilo koju numeričku varijablu u string pre nego što je spojite sa drugim tekstom.

Kreirajte varijablu pod nazivom `experience_years` i dodelite joj ceo broj `5`.

Zatim, kreirajte varijablu `experience_info`. Dodelite joj string koji je formiran konkateniranjem `'Experience: '`, varijable `experience_years` (konvertovane u string) i `' years'`. Ispišite rezultat na terminal.

# --hints--

You should have a variable named `experience_years`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_years")`)) })
```

You should assign the integer `5` to your `experience_years` variable.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_years").is_equivalent("experience_years = 5")`)) })
```

You should have a variable named `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_info")`)) })
```

You should assign a string formed by concatenating `'Experience: '`, `str(experience_years)`, and `' years'` to your `experience_info` variable.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_info").is_equivalent("experience_info = 'Experience: ' + str(experience_years) + ' years'")`)) })
```

You should print `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(experience_info)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
--fcc-editable-region--

--fcc-editable-region--
```
