---
id: 695ea1e55c0e619e597fc5c9
title: Korak 10
challengeType: 20
dashedName: step-10
---

# --description--

Kao što vidite, Python je izbacio `TypeError: can only concatenate str (not "int") to str`. Ovo se dešava jer Python ne dozvoljava da direktno konkateniš tekst (stringove) i brojeve (integere).

Da biste ovo popravili, morate prvo konvertovati broj u string koristeći funkciju `str()`, koja vraća string verziju objekta:

```py
my_num = str(42)
print(type(my_num)) # <class 'str'>
```

Ažurirajte dodelu u vašem `employee_info` da konvertujete `employee_age` u string koristeći `str(employee_age)`.

# --hints--

You should modify the assignment of your `employee_info` variable by converting `employee_age` into a string with `str(employee_age)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + employee_age
--fcc-editable-region--
```
