---
id: 69727ea5c9606bedfe36d3d2
title: Korak 15
challengeType: 20
dashedName: step-15
---

# --description--

Sada je vreme da dodamo završne detalje na karticu.

Kreirajte varijablu pod imenom `position` sa vrednošću stringa `Data Analyst` i varijablu pod imenom `salary` sa vrednošću celog broja `75000`.

Zatim, ažurirajte svoj `employee_card` f-string da uključite poziciju i platu. Trebalo bi da prati ovaj tačan format: `Employee: [full_name] | Age: [employee_age] | Position: [position] | Salary: $[salary]`. Zamenite zamenske delove odgovarajućim varijablama.

Konačno, štampajte `employee_card` da biste videli rezultat.

# --hints--

You should have a variable named `position`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("position")`)) })
```

You should assign the string `Data Analyst` to your `position` variable.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("position").is_equivalent("position = 'Data Analyst'")`)) })
```

You should have a variable named `salary`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("salary")`)) })
```

You should assign the integer `75000` to your `salary` variable.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

You should update `employee_card` to use the f-string `f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: \${salary}'")`)) })
```

You should print `employee_card` to the terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_card)")`))})
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

employee_card = f'Employee: {full_name} | Age: {employee_age}'
--fcc-editable-region--
```
