---
id: 6980d3ecfe800d8a54130f30
title: Korak 14
challengeType: 20
dashedName: step-14
---

# --description--

Trenutno, `employee_card` prikazuje samo ime zaposlenog. Sada ćete mu dodati više informacija. 

Ažurirajte dodelu `employee_card` da uključite godine zaposlenog. Konačni niz bi trebalo da izgleda ovako: `Employee: [name] | Age: [age]` sa `[name]` zamenjenim imenom zaposlenog, a `[age]` zamenjenim godinama zaposlenog.

# --hints--

You should update `employee_card` to use the f-string `f'Employee: {full_name} | Age: {employee_age}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age}'")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--
employee_card = f'Employee: {full_name}'
--fcc-editable-region--
```
