---
id: a6e40f1041b06c996f7b2406
title: Izgradite Pronalazač Prvog Elementa
challengeType: 26
dashedName: build-a-first-element-finder
---

# --description--

U ovom laboratorijumu, kreiraćete funkciju koja prolazi kroz niz i vraća prvi element koji prođe test funkciju (poznatu i kao „test istine“).

Funkcija bi iterirala kroz niz i testirala svaki element koristeći dostavljenu test funkciju. Na kraju, vratila bi prvi element za koji test funkcija vraća `true`.

Primer:

```js
findElement([1, 3, 5, 8], num => num % 2 === 0) // returns 8
findElement([1, 3, 5], num => num % 2 === 0)    // returns undefined
```

**Cilj:** Ispunite priče korisnika ispod i osigurajte da svi testovi prođu kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da imate funkciju pod nazivom `findElement` koja prima niz i funkciju kao argumente.
2. Funkcija bi trebalo da vrati prvi stavku u nizu koja prođe test istine. To znači da, pozivanjem prosleđene funkcije `func`, dato element `x`, test istine je prošao ako je `func(x)` `true`.
3. Ako nijedan element ne prođe test, funkcija bi trebalo da vrati `undefined`.

# --hints--

You should have a `findElement` function.

```js
assert.isFunction(findElement);
```

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` should return `8`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` should return `undefined`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

`findElement([1, 2, 3, 4], function(num) { return num > 2; })` should return `3`.

```js
assert.strictEqual(
  findElement([1, 2, 3, 4], function (num) {
    return num > 2;
  }),
  3
);
```

`findElement(["hello", "world", "javascript"], function(str) { return str.length > 5; })` should return `"javascript"`.

```js
assert.strictEqual(
  findElement(["hello", "world", "javascript"], function (str) {
    return str.length > 5;
  }),
  "javascript"
);
```

`findElement(["cat", "dog", "bird"], function(str) { return str.length > 10; })` should return `undefined`.

```js
assert.strictEqual(
  findElement(["cat", "dog", "bird"], function (str) {
    return str.length > 10;
  }),
  undefined
);
```

`findElement([2, 4, 6, 8], function(num) { return num % 2 === 0; })` should return `2`.

```js
assert.strictEqual(
  findElement([2, 4, 6, 8], function (num) {
    return num % 2 === 0;
  }),
  2
);
```

`findElement([], function(num) { return num > 0; })` should return `undefined`.

```js
assert.strictEqual(
  findElement([], function (num) {
    return num > 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
