---
id: 68c343692285523aef0fc6bf
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Atribut ``scope`` se koristi za određivanje da li je zaglavna ćelija zaglavlje za red, kolonu ili grupu redova ili kolona. Evo primera:

```html
<th scope="col">Example Header</th>
```

Ovo pomaže čitačima ekrana da razumeju odnos između zaglavnih i podataka ćelija.

Za sve elemente ``th``, dodajte atribut ``scope`` sa vrednošću ``col``.

# --hints--

Each `th` element should have a `scope` attribute with the value `"col"`.

```js
const tableHeaders = document.querySelectorAll("th");
tableHeaders.forEach(header => {
  assert.equal(header?.getAttribute("scope"), "col");
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Tech Conference 2025 Schedule</title>
</head>
<body>
  <h1>Tech Conference 2025 Schedule</h1>

  <table>
    <caption>Schedule by Track and Time</caption>

    <thead>
      <tr>
      --fcc-editable-region--
        <th>Time</th>
        <th>Track A</th>
        <th>Track B</th>
        <th>Track C</th>
      --fcc-editable-region--
      </tr>
    </thead>
  </table>
</body>
</html>
```
