---
id: 6995d0325fd7da40854fac40
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U prethodnoj lekciji, naučili ste o binarnom pretraživanju i kako se razlikuje od linearnog pretraživanja.

U ovom radionici, primenićete to praksom kreiranjem algoritma binarnog pretraživanja.

Da biste počeli, kreirajte funkciju `binarySearch` sa parametrima `searchList` i `value`.

`searchList` je niz elemenata unutar kojih će funkcija pretraživati, a `value` je cilj koji će funkcija pokušati da pronađe unutar `searchList`.

# --hints--

You should create a `binarySearch` function.

```js
assert.isFunction(binarySearch);
```

Your `binarySearch` function should have a `searchList` and `value` parameters.

```js
const binSearchParams = __helpers.getFunctionParams(binarySearch.toString());

const binarySearchParamsNames = binSearchParams.map(param => param.name)
assert.include(binarySearchParamsNames, "searchList");
assert.include(binarySearchParamsNames, "value");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
