---
id: 6996c675b599baa7b103192d
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

U binarnom pretraživanju, postoji najniži mogući indeks i najviši mogući indeks.

Najniži mogući indeks predstavlja najleviju granicu u trenutnom delu vrednosti koje se pretražuju, u ovom slučaju, `searchList`. Najviši mogući indeks, s druge strane, predstavlja najdesniju granicu u vrednostima koje se pretražuju.

Da biste uzela u obzir ta dva, morate razmotriti ceo niz koji se pretražuje. Stoga, unutar funkcije `binarySearch`, definišite varijablu `low` sa vrednošću `0`, i varijablu `high` sa vrednošću koja uzima u obzir poslednji indeks u pretraživanom nizu.

# --hints--

You should have a variable named `low` set to `0` in your `binarySearch` function.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /(const|let|var)\s+low\s*=\s*0/);
```

You should have a variable named `high` set to `searchList.length - 1` in your `binarySearch` function.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+high\s*=\s*searchList\.length\s*-\s*1/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
