---
id: 6996c796385b6ea4e222d5e8
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Da biste počeli sa petljom, morate pronaći srednji indeks trenutnog opsega pretrage.

Srednji indeks možete dobiti izračunavanjem proseka promenljivih `low` i `high`. Moraćete da koristite celobrojno deljenje (floor division) kako biste dobili prosek nakon sabiranja dve vrednosti, tako da će odgovor uvek biti ceo broj zaokružen nadole, jer indeksi moraju biti cele brojeve.

Unutar vaše `while` petlje, kreirajte promenljivu `mid` postavljenu na prosek `low` i `high`.

# --hints--

Within the while loop you should create a `mid` variable set to `Math.floor((low + high) / 2)`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+mid\s*=\s*Math\.floor\(\s*\(\s*low\s*\+\s*high\s*\)\s*\/\s*2\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;

  while (low <= high) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
