---
id: 6996c95e5bede03bdf2937a2
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Sledeća stvar koju treba uraditi je kreirati uslov koji će proveriti da li je cilj `value` u sredini. 

Unutar petlje `while`, kreirajte izjavu `if` koja proverava da li je cilj `value` jednak `valueAtMiddle`. Ako jeste, vratite varijablu `pathToTarget`.

# --hints--

You should create an `if` statement with the condition `value === valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)/
);
```

You should return the `pathToTarget` variable inside the `if` statement.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?return\s+pathToTarget/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
