---
id: 6996ccadf9048c79b857bb15
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--


<0x0D>
<0x0D>
Možete videti da prvi poziv vraća `[3]`. To je zato što je `3` srednja vrednost u `[1, 2, 3, 4, 5]`.

<0x0D>
S druge strane, drugi poziv vraća prazan niz. To se dešava zato što petlja trenutno nema logiku za sužavanje opsega pretrage ažuriranjem promenljivih `low` ili `high`. Ona funkcioniše samo kao provera jedne središnje tačke i zatim završava bez daljeg pretraživanja.

<0x0D>
Da bi se omogućilo da binarna pretraga nastavi sužavanje svog opsega, dodajte izjavu `else if` koja proverava da li je `value` veće od `valueAtMiddle`.

# --hints--

You should add an `else if` statement with the condition `value > valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s+if\s*\(\s*value\s*>\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
--fcc-editable-region--
    } 
--fcc-editable-region--
    break;
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
