---
id: 6996c60281feca6718cc15c1
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Unutar vaše funkcije `binarySearch`, kreirajte varijablu `pathToTarget` i inicijalizujte je praznim nizom, za početak.

# --hints--

You should have a `pathToTarget` variable set to an empty array `[]` within your `binarySearch` function.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+pathToTarget\s*=\s*\[\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
