---
id: 6996c8375f4301004459ee85
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Sada kada ste izračunali indeks `mid`, morate zapravo dohvatiti vrednost sa `searchList` na tom indeksu. Ova vrednost je onaj parametar koji ćete uporediti sa vašim `value`, a to je cilj koji tražite.

Kreirajte varijablu `valueAtMiddle` i dodelite joj element iz `searchList`, koji se nalazi na indeksu `mid`.

# --hints--

You should create a `valueAtMiddle` variable set to `searchList[mid]`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+valueAtMiddle\s*=\s*searchList\s*\[\s*mid\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
