---
id: 6996c7363a4a437839f1f15d
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Sada kada ste definisali granice svog pretrage, vreme je da kreirate petlju koja će izvršiti algoritam binarnog pretraživanja.

Binarno pretraživanje radi ponavljajućim sužavanjem traženog prostora. Ovaj proces se nastavlja sve dok postoji validan opseg elemenata za proveru.

Ovo možete izraziti sa `while` petljom koja će se nastaviti sve dok je vaš pokazivač `low` manji ili jednak vašem pokazivaču `high`.

Ako `low` ikada postane veći od `high`, to znači da je traženi prostor postao prazan, i da `value` nije u nizu.

Unutar funkcije `binarySearch`, kreirajte `while` petlju sa uslovom koji proverava da li je `low` manji ili jednak high.

# --hints--

You should create a `while` loop that runs while `low` is less than or equal to `high`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /while\s*\(\s*low\s*<=\s*high\s*\)/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
