---
id: 6996cfb6c01e9a802add38be
title: Korak 13
challengeType: 1
dashedName: step-13
---

# --description--

Konačno, ako `value` nije pronađen u sredini i takođe nije veći od `valueAtMiddle`, onda mora biti da je `value` manji od `valueAtMiddle` i mora biti na levoj strani. Poslednji `else` blok će to obraditi.

Dakle, završi petlju uklanjanjem ključne reči `break` i dodavanjem bloka `else`. Unutar njega, ažuriraj varijablu `high` oduzimanjem `1` od `mid`.

# --hints--

You should replace the `break` keyword with an `else` block that updates the `high` variable to `mid - 1`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s*\{\s*high\s*=\s*mid\s*-\s*1\s*;?\s*\}/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
--fcc-editable-region--      
    } 
    break;
--fcc-editable-region--
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
