---
id: 6996ca21a74104721f2c7e5d
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

Moraš da isprobaš stvari kako bi razumeo tok algoritma u ovoj početnoj fazi.

Da bi to uradio, prvo moraš da izađeš iz petlje. To je zato što trenutna implementacija će dozvoliti samo jednu iteraciju, pa ako uslov nije zadovoljen, biće beskonačna petlja.

Neposredno nakon izjave `if`, koristi ključnu reč `break` da izađeš iz petlje `while`. Zatim, nakon petlje `while`, vrati prazan niz kako bi se signalizovalo da vrednost nije pronađena.

# --hints--

You should use the `break` keyword to break out of the `while` loop after the `if` statement.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

You should return an empty array after your `while` loop.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /return\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    
  }
  
--fcc-editable-region--
}
```
