---
id: 6996d0a2c2ac55041145a13f
title: Korak 14
challengeType: 1
dashedName: step-14
---

# --description--

Da li ste primetili da drugi poziv funkcije sada ima vrednosti `3, 5, 4` u listi? To ukazuje da je binarna pretraga funkcioniše kako što je predviđeno.

Ovo se dogodilo ovako: Algoritam je prvo proverio `3` kao sredinu početne liste. Pošto je `4` veće od `3`, pretraga se pomerila na desnu polovinu. Zatim je pregledao `5` kao novu sredinu. Pošto je `4` manje od `5`, pretraga se pomerila ulevo, na kraju identifikujući `4` kao sredinu konačnog raspona.

Da biste ponovo testirali funkciju, pozovite je sa `[1, 3, 5, 9, 14, 22], 10` i logujte je u konzolu. Ovo je situacija u kojoj vrednost neće biti pronađena.

# --hints--

You should call the function with `binarySearch([1, 3, 5, 9, 14, 22], 10)` and log it to the console.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*3\s*,\s*5\s*,\s*9\s*,\s*14\s*,\s*22\s*\]\s*,\s*10\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
--fcc-editable-region--

--fcc-editable-region--
```
