---
id: 698367b9d1c6914a22095aec
title: Implementiraj algoritam Bubble Sort
challengeType: 26
dashedName: implement-the-bubble-sort-algorithm
---

# --description--

Za ovaj laboratorijski zadatak, implementiraćete algoritam bubble sort (sortiranje mehurićem). Počinje od početka neuređenog niza i "mehuriće" neuređene vrednosti ka kraju, iterirajući kroz niz dok on nije potpuno sortirisan. To radi poređenjem susednih elemenata i zamjenjivanjem ih ako nisu u redu. Metoda nastavlja da prolazi kroz niz sve dok se ne izvrše nijedna zamena, tačkoj kada je niz sortiran.

Ova metoda zahteva više iteracija kroz niz i za prosečne i najgore slučajeve ima kvadratnu vremensku složenost. Iako je jednostavna, obično nije praktična u većini situacija.

**Cilj:** Ispunite priče korisnika ispod i učinite da svi testovi prođu kako biste završili laboratorijski zadatak.

**Priče korisnika:**

1. Napisite funkciju `bubbleSort` koja prima niz celih brojeva kao ulaz i vraća niz ovih celih brojeva u sortiranom redosledu od najmanjeg do najvećeg.

# --hints--

You should have a function named `bubbleSort`.

```js
assert.isFunction(bubbleSort);
```

`bubbleSort` should return a sorted array (least to greatest).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    bubbleSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`bubbleSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` should return an array that is unchanged except for order.

```js
assert.sameMembers(
  bubbleSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`bubbleSort` should not use the built-in `.sort()` method.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    bubbleSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function bubbleSort(array) {
  for (let i = 0; i < array.length; i++) {
    let swapped = false;
    for (let j = 1; j < array.length; j++) {
      if (array[j - 1] > array[j]) {
        let temp = array[j-1];
        array[j-1] =  array[j];
        array[j] = temp;
        swapped = true;
      }
    }
    if (swapped === false) {
      break;
    }
  }
  return array;
}
```
