id: 587d7fa8367417b2b2512bcb
title: Naučite o SVG u D3
challengeType: 6
forumTopicId: 301489
dashedName: learn-about-svg-in-d3
---

# --description--

<dfn>SVG</dfn> znači <dfn>Scalable Vector Graphics</dfn>.

Ovde "scalable" (skalabilan) znači da, bez obzira na povećanje ili smanjenje veličine objekta, on neće izgledati kao slika sa pikselima. On se skalira prema sistemu za prikaz, bilo da je na malom ekranu telefona ili na velikom televizoru.

SVG se koristi za kreiranje standardnih geometrijskih oblika. Pošto D3 predstavlja podatke vizuelno (grafički), on koristi SVG za kreiranje tih reprezentativnih oblika. SVG oblici za web stranicu moraju biti unutar HTML oznake `svg`.

CSS može biti skalabilan kada se stili koriste relativne jedinice (kao što su `vh`, `vw` ili procenti), ali korišćenje SVG je najjednostavniji način za kreiranje vizuelnih reprezentacija podataka.

# --instructions--

Dodajte `svg` čvor u `body` koristeći `append()`. Dodelite atribut `width` fiksnom vrednosti `w` i atribut `height` fiksnom vrednosti `h` koristeći metode `attr()` ili `style()` za svaki od njih. To ćete videti u rezultatu jer je na oznaci `style` korišćena pozadina boje ružičaste.

**Napomena:** Kada koristite atribute širine i visine sa `attr()`, oni nemaju jedinice. Ovo je osnova skaliranja – element će imati odnos širina/visina od 5:1, bez obzira na nivo povećanja ili smanjenja veličine.

# --hints--

Vaš dokument bi trebalo da ima jedan `svg` element.

```js
assert.lengthOf(document.querySelectorAll('svg'), 1);
```

`svg` element bi trebalo da ima atribut `width` postavljen na `500` ili da je stilizovan sa širinom od `500px`.

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const width = window.getComputedStyle(svg)['width'];
assert.isTrue(svg?.getAttribute('width') === '500' || width === '500px');
```

`svg` element bi trebalo da ima atribut `height` postavljen na `100` ili da je stilizovan sa visinom od `100px`.

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const height = window.getComputedStyle(svg)['height'];
assert.isTrue(svg?.getAttribute('height') === '100' || height === '100px');
```

# --seed--

## --seed-contents--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3.select('body');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
  </script>
</body>
```