id: 587d7fac367417b2b2512bdd
title: Koristite dinamičke skale
challengeType: 6
forumTopicId: 301495
dashedName: use-dynamic-scales
---

# --description--

Metode D3 `min()` i `max()` su ključne za postavljanje skale.

Kada imate složeni skup podataka, jedan od prioriteta je postavljanje skale kako bi prikaz odgovarao širini i visini SVG kontejnera. Želite da se svi podaci nacrtaju unutar SVG-a kako bi bili vidljivi na web stranici.

Primer ispod postavlja skalu za x-osu za podatke o točkama. Metoda `domain()` prenosi informacije skali o sirovim vrijednostima podataka za crtež. Metoda `range()` pruža informacije o stvarnom prostoru na web stranici za prikaz.

U ovom primjeru, domen počinje od 0 do najveće vrijednosti u skupu. Koristi `max()` metodu zajedno sa funkcijom povratne vrijednosti koja zavisi od x vrijednosti u podacima. Opseg koristi širinu SVG (`w`), ali uključuje i unutrašnji padding (razmak). Ovo ostavlja razmak između točaka crteža i ivica SVG-a.

```js
const dataset = [
  [34, 78],
  [109, 280],
  [310, 120],
  [79, 411],
  [420, 220],
  [233, 145],
  [333, 96],
  [222, 333],
  [78, 320],
  [21, 123]
];
const w = 500;
const h = 500;

const padding = 30;
const xScale = d3
  .scaleLinear()
  .domain([0, d3.max(dataset, d => d[0])])
  .range([padding, w - padding]);
```

Unutrašnji razmak može izgledati komplikovano na prvi pogled. Zamislite x-osu kao horizontalnu liniju od 0 do 500 (vrijednost širine SVG-a). Uključivanje unutrašnjeg razmaka u `range()` prisiljava crtež da počne sa 30 sa te linije (umjesto 0), i završi se na 470 (umjesto 500).

# --instructions--

Koristite varijablu `yScale` za kreiranje dinamičke skale y-ose za liniju. Domen treba da počne od nule i ide do najveće vrijednosti `y` u skupu. Opseg treba da koristi visinu SVG (`h`) i uključuje unutrašnji razmak.

**Napomena:** Zapamtite da postavite crtež ispravno. Kada postavljate opseg za y koordinate, najveća vrijednost (visina umanjena za unutrašnji razmak) je prvi argument, a najmanja vrijednost je drugi argument.

# --hints--

Tekst u `h2` treba da bude `30`.

```js
assert.strictEqual(output, 30);
assert.strictEqual(document.querySelector('h2')?.textContent, '30');
```

`domain()` za yScale treba da bude jednak `[0, 411]`.

```js
assert.deepEqual(yScale.domain(), [0, 411]);
```

`range()` za yScale treba da bude jednak `[470, 30]`.

```js
assert.deepEqual(yScale.range(), [470, 30]);
```

# --seed--

## --seed-contents--

```html
<body >
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    // Padding between the SVG boundary and the plot
    const padding = 30;

    // Create an x and y scale

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    // Add your code below this line

    const yScale = undefined;


    // Add your code above this line

    const output = yScale(411); // Returns 30
    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body >
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const padding = 30;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const output = yScale(411);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```