id: 587d7faa367417b2b2512bd6
title: Dodavanje alata za napomenu (tooltip) D3 elementu
challengeType: 6
forumTopicId: 301470
dashedName: add-a-tooltip-to-a-d3-element
---

# --description--

Alat za napomenu prikazuje više informacija o elementu na stranici kada korisnik pređe mišem preko tog elementa. Postoji nekoliko načina za dodavanje alata za napomenu u prikaz grafikona. Ovaj izazov koristi SVG `title` element.

`title` je vezan uz metodu `text()` da dinamički doda podatke na barove.

# --instructions--

Dodajte element `title` ispod svakog `rect` čvora. Zatim pozovite metodu `text()` sa funkcijom za pristup kako bi tekst prikazao vrednost podataka.

# --hints--

Vaš kod treba da ima 9 elemenata `title`.

```js
assert.lengthOf(document.querySelectorAll('title'), 9);
```

Prvi element `title` treba da ima tekst alata za napomenu `12`.

```js
assert.strictEqual(document.querySelectorAll('title')[0]?.textContent, '12');
```

Drugi element `title` treba da ima tekst alata za napomenu `31`.

```js
assert.strictEqual(document.querySelectorAll('title')[1]?.textContent, '31');
```

Treći element `title` treba da ima tekst alata za napomenu `22`.

```js
assert.strictEqual(document.querySelectorAll('title')[2]?.textContent, '22');
```

Četvrti element `title` treba da ima tekst alata za napomenu `17`.

```js
assert.strictEqual(document.querySelectorAll('title')[3]?.textContent, '17');
```

Peti element `title` treba da ima tekst alata za napomenu `25`.

```js
assert.strictEqual(document.querySelectorAll('title')[4]?.textContent, '25');
```

Šesti element `title` treba da ima tekst alata za napomenu `18`.

```js
assert.strictEqual(document.querySelectorAll('title')[5]?.textContent, '18');
```

Sedmi element `title` treba da ima tekst alata za napomenu `29`.

```js
assert.strictEqual(document.querySelectorAll('title')[6]?.textContent, '29');
```

Osmi element `title` treba da ima tekst alata za napomenu `14`.

```js
assert.strictEqual(document.querySelectorAll('title')[7]?.textContent, '14');
```

Deveti element `title` treba da ima tekst alata za napomenu `9`.

```js
assert.strictEqual(document.querySelectorAll('title')[8]?.textContent, '9');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar')
      .append('title')
      .text(d => d);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```