---
id: 587d7fac367417b2b2512bdc
title: >-
  Koristite funkcije d3.max i d3.min za pronalaženje minimalnih i maksimalnih vrednosti u skupu podataka
challengeType: 6
forumTopicId: 301496
dashedName: >-
  use-the-d3-max-and-d3-min-functions-to-find-minimum-and-maximum-values-in-a-dataset
---

# --description--

D3 funkcije `domain()` i `range()` postavljaju informacije na vaš nivo u zavisnosti od podataka. Postoji nekoliko načina da to učinite lakše.

Često kada postavljate domen, želite koristiti minimalne i maksimalne vrednosti unutar skupa podataka. Pokušaj ručnog pronalaženja ovih vrednosti, posebno u velikom skupu podataka, može dovesti do grešaka.

D3 ima dve funkcije - `min()` i `max()` koje vraćaju ove informacije. Evo primera:

```js
const exampleData = [34, 234, 73, 90, 6, 52];
d3.min(exampleData);
d3.max(exampleData);
```

Skup podataka može imati ugnježdene nizove podataka, kao što su koordinatni parovi `[x, y]` koji su bili u formatu tačke grafa. U tom slučaju, morate reći D3 kako da izračuna maksimum i minimum. Srećom, obe funkcije `min()` i `max()` uzimaju funkciju povratnika (callback function). U ovom primeru, argument funkcije povratnika `d` je za trenutni ugnježdeni niz. Funkcija povratnika mora vratiti element iz ugnježdenog niza (vrednost `x` ili `y`) koji želite da izračunate maksimum ili minimum. Evo primera kako pronaći minimalne i maksimalne vrednosti za nizove podataka:

```js
const locationData = [
  [1, 7],
  [6, 3],
  [8, 3]
];
const minX = d3.min(locationData, d => d[0]);
```

`minX` bi imao vrednost `1`.

# --instructions--

Niz `positionData` sadrži ugnježdene nizove koordinata x, y i z. Koristite D3 funkciju da pronađete najveću vrednost koordinate z (treći element) iz tih nizova i dodelite je promenljivoj `output`.

# --hints--

Tekst u `h2` treba da bude `8`.

```js
assert.strictEqual(output, 8);
assert.strictEqual(document.querySelector('h2')?.textContent, '8');
```

Vaš kod bi trebalo da koristi funkciju `max()`.

```js
assert.match(
  code,
  /\.max/g,
  'Your code should use the <code>max()</code> method.'
);
```

# --seed--

## --seed-contents--

```html
<body >
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];
    // Add your code below this line

    const output = undefined; // Change this line

    // Add your code above this line

    d3.select('body').append('h2').text(output);
  </script>
</body >
```

# --solutions--

```html
<body >
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];

    const output = d3.max(positionData, d => d[2]);

    d3.select('body').append('h2').text(output);
  </script>
</body >
```