id: 587d7faa367417b2b2512bd4
title: Dodaj efekat prelaska (hover) elementu D3
challengeType: 6
forumTopicId: 301469
dashedName: add-a-hover-effect-to-a-d3-element
---

# --description--

Moguće je dodati efekte koji ističu bar kada korisnik pređe mišem preko njega. Trenutno, stil pravougaonika postavljen je koristeći ugrađene D3 i SVG metode, ali možete koristiti i CSS.

Postavljate CSS klasu na SVG elemente pomoću `attr()`. Zatim, pseudo-klasa `:hover` za vašu novu klasu ima stilove za bilo koji efekat prelaska.

# --instructions--

Koristite metodu `attr()` da dodate klasu `bar` svim `rect` elementima. Ovo menja boju `fill` bara u smeđu kada pređete mišem preko njega.

# --hints--

Vaši `rect` elementi bi trebalo da imaju klasu `bar`.

```js
assert.isTrue(document.querySelector('rect')?.classList.contains('bar'));
```

# --seed--

## --seed-contents--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3);
  </script>
</body >
```

# --solutions--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy')
    // Add your code below this line
      .attr('class', 'bar');
    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3);
  </script>
</body >
```