id: 587d7fa9367417b2b2512bce
title: Postavite dinamički koordinatu za svaku šipku
challengeType: 6
forumTopicId: 301487
dashedName: dynamically-set-the-coordinates-for-each-bar
---

# --description--

Poslednji izazov kreirao je i dodao pravougaonik na `svg` element za svaku tačku u `dataset` kako bi predstavio šipku. Nažalost, sve su bile postavljene jedna preko druge.

Pozicija pravougaonika povezana je sa atributima `x` i `y`. Oni D3-u govore gde početi crtati oblik unutar `svg` područja. Poslednji izazov ih je sve postavio na 0, pa su sve šipke bile postavljene u gornji levi ugao.

Za linjasti grafikon, sve linije bi trebalo da budu na istoj vertikalnoj razini, što znači da vrednost `y` ostaje ista (na 0) za sve linije. Međutim, vrednost `x` mora se menjati kako dodajete nove šipke. Zapamtite da veće vrednosti `x` guraju stvari više u desno. Kada prolazite kroz elemente niza podataka u `dataset`-u, vrednost `x` bi trebalo da raste.

Metoda `attr()` u D3 prihvata funkciju povratnika za dinamičko postavljanje tog atributa. Funkcija povratnik prima dva argumenta: jedan je sama tačka podataka (obično `d`) a drugi je indeks te tačke podataka u nizu podataka. Drugi argument indeksa je opcionalan. Struktura izgleda ovako:

```js
selection.attr('property', (d, i) => {});
```

Važno je razumeti da NE TREBA da pišete `for` petlju niti koristite `forEach()` za iteraciju kroz elemente u skupu podataka. Zapamtite da metoda `data()` analizira skup podataka, a bilo koja metoda koja sledi `data()` izvršava se jednom za svaki element u skupu podataka.

# --instructions--

Izmenite funkciju povratnika atributa `x` tako da vrati indeks pomnožen sa 30.

**Napomena:** Svaka šipka ima širinu od 25, pa dodavanje vrednosti `x` za 30 povećava razmak između šipki. Bilo koja vrednost veća od 25 će raditi u ovom primeru.

# --hints--

Prvi `rect` bi trebalo da ima vrednost `x` od `0`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('x'),
  '0'
);
```

Drugi `rect` bi trebalo da ima vrednost `x` od `30`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('x'),
  '30'
);
```

Treći `rect` bi trebalo da ima vrednost `x` od `60`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('x'),
  '60'
);
```

Četvrti `rect` bi trebalo da ima vrednost `x` od `90`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('x'),
  '90'
);
```

Peti `rect` bi trebalo da ima vrednost `x` od `120`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('x'),
  '120'
);
```

Šesti `rect` bi trebalo da ima vrednost `x` od `150`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('x'),
  '150'
);
```

Sedmi `rect` bi trebalo da ima vrednost `x` od `180`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('x'),
  '180'
);
```

Osmi `rect` bi trebalo da ima vrednost `x` od `210`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('x'),
  '210'
);
```

Deveti `rect` bi trebalo da ima vrednost `x` od `240`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('x'),
  '240'
);
```

# --seed--

## --seed-contents--

```html
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        // Add your code below this line


      
        // Add your code above this line
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        return i * 30;
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```