id: 587d7faa367417b2b2512bd2
title: Dodaj oznake na elementima u D3
challengeType: 6
forumTopicId: 301476
dashedName: add-labels-to-d3-elements
---

# --description--

D3 vam omogućava dodavanje oznaka na grafikonskim elementima, kao što su barovi, korišćenjem SVG `text` elementa.

Za element tipa `rect`, element `text` zahteva atribute `x` i `y` kako bi mogao biti postavljen u SVG. Takođe mu je potreban pristup podacima da bi prikazao te vrednosti.

D3 vam daje veliku kontrolu nad načinom na koji označavate svoje barove.

# --instructions--

Kod u editoru već vezuje podatke za svaki novi element `text`. Prvo, dodajte `text` node-ove u `svg`. Zatim, dodajte koordinate atribute `x` i `y`. Trebalo bi da se izračunaju isto kao za `rect`, osim što vrednost `y` za `text` treba da učini oznaku tako da je 3 jedinice više nego bar. Na kraju, koristite D3 metodu `text()` da postavite oznaku jednaku vrednosti podataka.

**Napomena:** Da bi oznaka bila iznad (iznad) barova, odredite da li vrednost `y` za `text` treba da bude 3 više ili 3 manje od vrednosti `y` za bar.

# --hints--

Prvi element `text` treba da ima oznaku `12` i vrednost `y` od `61`.

```js
assert.strictEqual(document.querySelectorAll('text')[0]?.textContent, '12');
assert.strictEqual(
  document.querySelectorAll('text')[0]?.getAttribute('y'),
  '61'
);
```

Drugi element `text` treba da ima oznaku `31` i vrednost `y` od `4`.

```js
assert.strictEqual(document.querySelectorAll('text')[1]?.textContent, '31');
assert.strictEqual(
  document.querySelectorAll('text')[1]?.getAttribute('y'),
  '4'
);
```

Treći element `text` treba da ima oznaku `22` i vrednost `y` od `31`.

```js
assert.strictEqual(document.querySelectorAll('text')[2]?.textContent, '22');
assert.strictEqual(
  document.querySelectorAll('text')[2]?.getAttribute('y'),
  '31'
);
```

Četvrti element `text` treba da ima oznaku `17` i vrednost `y` od `46`.

```js
assert.strictEqual(document.querySelectorAll('text')[3]?.textContent, '17');
assert.strictEqual(
  document.querySelectorAll('text')[3]?.getAttribute('y'),
  '46'
);
```

Peti element `text` treba da ima oznaku `25` i vrednost `y` od `22`.

```js
assert.strictEqual(document.querySelectorAll('text')[4]?.textContent, '25');
assert.strictEqual(
  document.querySelectorAll('text')[4]?.getAttribute('y'),
  '22'
);
```

Šesti element `text` treba da ima oznaku `18` i vrednost `y` od `43`.

```js
assert.strictEqual(document.querySelectorAll('text')[5]?.textContent, '18');
assert.strictEqual(
  document.querySelectorAll('text')[5]?.getAttribute('y'),
  '43'
);
```

Sedmi element `text` treba da ima oznaku `29` i vrednost `y` od `10`.

```js
assert.strictEqual(document.querySelectorAll('text')[6]?.textContent, '29');
assert.strictEqual(
  document.querySelectorAll('text')[6]?.getAttribute('y'),
  '10'
);
```

Osmi element `text` treba da ima oznaku `14` i vrednost `y` od `55`.

```js
assert.strictEqual(document.querySelectorAll('text')[7]?.textContent, '14');
assert.strictEqual(
  document.querySelectorAll('text')[7]?.getAttribute('y'),
  '55'
);
```

Deveti element `text` treba da ima oznaku `9` i vrednost `y` od `70`.

```js
assert.strictEqual(document.querySelectorAll('text')[8]?.textContent, '9');
assert.strictEqual(
  document.querySelectorAll('text')[8]?.getAttribute('y'),
  '70'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');

    svg.selectAll('text').data(dataset).enter();
    // Add your code below this line




    // Add your code above this line
  </script>
  <body></body>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3)
      .text(d => d);
  </script>
  <body></body>
</body>
```