---
id: 587d7fa6367417b2b2512bc3
title: Izaberite grupu elemenata pomoću D3
challengeType: 6
forumTopicId: 301490
dashedName: select-a-group-of-elements-with-d3
---

# --description--

D3 takođe ima metodu `selectAll()` za odabir grupe elemenata. Vraća niz HTML čvorova za sve elemente u dokumentu koji odgovaraju nizu karaktera unosa. Evo primera kako izabrati sve oznake anchor taga u dokumentu:

```js
const anchors = d3.selectAll('a');
```

Kao metoda `select()`, `selectAll()` podržava povezivanje sekvenci metoda, i možete ga koristiti zajedno sa drugim metodama.

# --instructions--

Izaberite sve `li` oznake u dokumentu i promenite njihov tekst na niz karaktera `list item` koristeći metodu `.text()`.

# --hints--

Na stranici postoji 3 `li` elementa, a tekst u svakom mora biti `list item`. Velika slova i razmaci moraju biti potpuno identični.

```js
const listItems = document.querySelectorAll('li');
assert.lengthOf(listItems, 3);
for (let i = 0; i < listItems.length; i++) {
  assert.match(listItems[i]?.textContent, /list item/g);
}
```

Msimbo vašeg koda mora imati pristup objektu `d3`.

```js
assert.match(code, /d3/g);
```

Msimbo vašeg koda mora koristiti metodu `selectAll`.

```js
assert.match(code, /\.selectAll/g);
```

# --seed--

## --seed-contents--

```html
<body >
  <ul >
    <li >Example</li >
    <li >Example</li >
    <li >Example</li >
  </ul >
  <script>
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body >
  <ul >
    <li >Example</li >
    <li >Example</li >
    <li >Example</li >
  </ul >
  <script>
    d3.selectAll('li').text('list item');
  </script>
</body>
```