id: 587d7fab367417b2b2512bd7
title: Kreiraj scatter plot sa krugovima u SVG-u
challengeType: 6
forumTopicId: 301484
dashedName: create-a-scatterplot-with-svg-circles
---

# --description--

Raspršeni graf je još jedan tip prikaza podataka. Obično koristi krugove za prikaz tačaka podataka, od kojih svaka ima dve vrednosti. Ove vrednosti su povezane sa `x` i `y` osama i koriste se za pozicioniranje kruga u vizualizaciji.

SVG ima element `circle` za kreiranje oblika kruga. Radi na sličan način kao što ste koristili elemente `rect` za linearni grafikon.

# --instructions--

Koristite metode `data()`, `enter()` i `append()` da vezujete `dataset` za nove elemente `circle` koji se dodaju u SVG.

**Napomena:** Krugovi neće biti vidljivi jer nismo još definisali njihove atribute elemenata. To ćemo uraditi u sledećem izazovu.

# --hints--

Vaš kod bi trebalo da sadrži 10 elemenata `circle`.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
  </script>
</body>
```