id: 587d7fa8367417b2b2512bc9
title: Ažurirajte visinu elementa dinamički
challengeType: 6
forumTopicId: 301493
dashedName: update-the-height-of-an-element-dynamically
---

# --description--

Prethovni izazovi uključivali su kako prikazati podatke iz niza i kako dodati CSS klase. Možete kombinovati ove lekcije da biste kreirali jednostavan linijski grafikon. Postoje dva koraka za to:

1. Kreirati `div` za svaki element podataka u nizu podataka

2. Postaviti visinu `div`-a dinamički, koristeći funkciju povratne poziva (callback function) u metodi `style()` koja postavlja visinu jednaku vrednosti podataka

Zapamtite uzorak za podešavanje stila pomoću funkcije povratne poziva:

```js
selection.style('cssProperty', d => d);
```

# --instructions--

Dodajte metodu `style()` u kod u editoru da biste postavili svojstvo `height` za svaki element. Koristite funkciju povratne poziva da vratite vrednost elementa podataka zajedno sa podržanim sufiksom stringa `px`.

# --hints--

Prvi `div` bi trebalo da ima visinu od 12 pikseli.

```js
assert.strictEqual(document.querySelectorAll('div')[0]?.style.height, '12px');
```

Drugi `div` bi trebalo da ima visinu od 31 piksela.

```js
assert.strictEqual(document.querySelectorAll('div')[1]?.style.height, '31px');
```

Treći `div` bi trebalo da ima visinu od 22 piksela.

```js
assert.strictEqual(document.querySelectorAll('div')[2]?.style.height, '22px');
```

Četvrti `div` bi trebalo da ima visinu od 17 piksela.

```js
assert.strictEqual(document.querySelectorAll('div')[3]?.style.height, '17px');
```

Peti `div` bi trebalo da ima visinu od 25 piksela.

```js
assert.strictEqual(document.querySelectorAll('div')[4]?.style.height, '25px');
```

Šesti `div` bi trebalo da ima visinu od 18 piksela.

```js
assert.strictEqual(document.querySelectorAll('div')[5]?.style.height, '18px');
```

Sedmi `div` bi trebalo da ima visinu od 29 piksela.

```js
assert.strictEqual(document.querySelectorAll('div')[6]?.style.height, '29px');
```

Osmi `div` bi trebalo da ima visinu od 14 piksela.

```js
assert.strictEqual(document.querySelectorAll('div')[7]?.style.height, '14px');
```

Deveti `div` bi trebalo da ima visinu od 9 piksela.

```js
assert.strictEqual(document.querySelectorAll('div')[8]?.style.height, '9px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => `${d}px`);
  </script>
</body>
```