id: 587d7fa8367417b2b2512bcc
title: Prikažite Oblike korišćenjem SVG-a
challengeType: 6
forumTopicId: 301485
dashedName: display-shapes-with-svg
---

# --description--

Prethovni izazov kreirao je element `svg` sa definisanim širinom i visinom, koji je bio vidljiv zahvaljujući `background-color` primenjenom unutar stila. Kod je napravio prostor za definisanu širinu i visinu.

Sledeći korak je kreiranje oblika za postavljanje u oblast `svg`. U SVG-u podržano je nekoliko oblika, kao što su pravougaonik i krug. Oni se koriste za prikaz podataka. Na primer, SVG pravougaoničnik (`<rect>`) može napraviti liniju na grafikonu linija.

Kada postavite oblik u oblast `svg`, možete definisati gde je postavljen korišćenjem koordinata `x` i `y`. Početna tačka (0, 0) nalazi se u gornjem levom uglu. Pozitivne vrednosti za `x` pomeraju oblik udesno, a pozitivne vrednosti za `y` pomeraju oblik nadole od početne tačke.

Da postavimo oblik centriran u `svg`-u širine 500 i visine 100 iz prethodnog izazova, koordinata `x` će biti 250, a koordinata `y` će biti 50.

SVG element `<rect>` ima četiri atributa. Postoje koordinate `x` i `y` gde je postavljen u oblasti `svg`. Takođe ima `height` i `width` za definisanje veličine.

# --instructions--

Dodajte pravougaonični oblik (`<rect>`) u `svg` koristeći `append()`, i dodelite mu atribut `width` sa vrednošću `25` i atribut `height` sa vrednošću `100`. Takođe, dodelite `rect`-u atribute `x` i `y`, svaki sa vrednošću `0`.

# --hints--

Vaš dokument bi trebalo da sadrži jedan element `<rect>`.

```js
assert.lengthOf(document.querySelectorAll('rect'), 1);
```

Element `<rect>` bi trebalo da ima atribut `width` sa vrednošću `25`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('width'), '25');
```

Element `<rect>` bi trebalo da ima atribut `height` sa vrednošću `100`.

```js
assert.strictEqual(
  document.querySelector('rect')?.getAttribute('height'),
  '100'
);
```

Element `<rect>` bi trebalo da ima atribut `x` sa vrednošću `0`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('x'), '0');
```

Element `<rect>` bi trebalo da ima atribut `y` sa vrednošću `0`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('y'), '0');
```

# --seed--

## --seed-contents--

```html
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
    // Add your code below this line



    // Add your code above this line
  </script>
</body >
```

# --solutions--

```html
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h)
      .append('rect')
      .attr('width', 25)
      .attr('height', 100)
      .attr('x', 0)
      .attr('y', 0);
  </script>
</body >
```