id: 587d7fa8367417b2b2512bca
title: Promena izgleda stolobar charta
challengeType: 6
forumTopicId: 301481
dashedName: change-the-presentation-of-a-bar-chart
---

# --description--

Poslednji izazov kreirao je stolobar grafikon, ali postoje male izmene dizajna koje ga mogu poboljšati:

1. Dodajte razmak između svake trake kako bi bile vizuelno odvojene, što se postiže dodavanjem margina na CSS klasu `bar`

2. Povećajte visinu traka da bolje prikažu razliku vrednosti, što se postiže množenjem vrednosti sa određenim brojem radi proširenja visine

# --instructions--

Prvo, dodajte `margin` od `2px` na klasu `bar` u stilski deo. Zatim, promenite funkciju povratne vrednosti u metodi `style()` tako da vrati vrednost `10` puta stvarnu podatkovnu vrednost (uz `px`).

**Napomena:** Množenje svake podatkovne tačke sa _istom količinom_ samo menja skalu. Kao razvlačenje slike, ne menjajući značenje postojećih podataka ispod.

# --hints--

Prvi `div` treba da ima visinu od `120px` i desnu margina od `2px`.

```js
const firstDiv = document.querySelector('div');
assert.exists(firstDiv); 
const height = window.getComputedStyle(firstDiv)['height'];
const margin = window.getComputedStyle(firstDiv)['margin-right'];
assert.strictEqual(height, '120px');
assert.strictEqual(margin, '2px');
```

Drugi `div` treba da ima visinu od `310px` i desnu margina od `2px`.

```js
const secondDiv = document.querySelectorAll('div')[1];
assert.exists(secondDiv); 
const height = window.getComputedStyle(secondDiv)['height'];
const margin = window.getComputedStyle(secondDiv)['margin-right'];
assert.strictEqual(height, '310px');
assert.strictEqual(margin, '2px');
```

Treći `div` treba da ima visinu od `220px` i desnu margina od `2px`.

```js
const thirdDiv = document.querySelectorAll('div')[2];
assert.exists(thirdDiv); 
const height = window.getComputedStyle(thirdDiv)['height'];
const margin = window.getComputedStyle(thirdDiv)['margin-right'];
assert.strictEqual(height, '220px');
assert.strictEqual(margin, '2px');
```

Četvrti `div` treba da ima visinu od `170px` i desnu margina od `2px`.

```js
const fourthDiv = document.querySelectorAll('div')[3];
assert.exists(fourthDiv); 
const height = window.getComputedStyle(fourthDiv)['height'];
const margin = window.getComputedStyle(fourthDiv)['margin-right'];
assert.strictEqual(height, '170px');
assert.strictEqual(margin, '2px');
```

Peti `div` treba da ima visinu od `250px` i desnu margina od `2px`.

```js
const fifthDiv = document.querySelectorAll('div')[4];
assert.exists(fifthDiv); 
const height = window.getComputedStyle(fifthDiv)['height'];
const margin = window.getComputedStyle(fifthDiv)['margin-right'];
assert.strictEqual(height, '250px');
assert.strictEqual(margin, '2px');
```

Šesti `div` treba da ima visinu od `180px` i desnu margina od `2px`.

```js
const sixthDiv = document.querySelectorAll('div')[5];
assert.exists(sixthDiv); 
const height = window.getComputedStyle(sixthDiv)['height'];
const margin = window.getComputedStyle(sixthDiv)['margin-right'];
assert.strictEqual(height, '180px');
assert.strictEqual(margin, '2px');
```

Sedmi `div` treba da ima visinu od `290px` i desnu margina od `2px`.

```js
const seventhDiv = document.querySelectorAll('div')[6];
assert.exists(seventhDiv); 
const height = window.getComputedStyle(seventhDiv)['height'];
const margin = window.getComputedStyle(seventhDiv)['margin-right'];
assert.strictEqual(height, '290px');
assert.strictEqual(margin, '2px');
```

Osmi `div` treba da ima visinu od `140px` i desnu margina od `2px`.

```js
const eighthDiv = document.querySelectorAll('div')[7];
assert.exists(eighthDiv); 
const height = window.getComputedStyle(eighthDiv)['height'];
const margin = window.getComputedStyle(eighthDiv)['margin-right'];
assert.strictEqual(height, '140px');
assert.strictEqual(margin, '2px');
```

Deveti `div` treba da ima visinu od `90px` i desnu margina od `2px`.

```js
const ninthDiv = document.querySelectorAll('div')[8];
assert.exists(ninthDiv); 
const height = window.getComputedStyle(ninthDiv)['height'];
const margin = window.getComputedStyle(ninthDiv)['margin-right'];
assert.strictEqual(height, '90px');
assert.strictEqual(margin, '2px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    /* Add your code below this line */

    
    /* Add your code above this line */
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d + 'px'); // Change this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    margin: 2px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d * 10 + 'px');
  </script>
</body>
```