id: 587d7fa7367417b2b2512bc7
title: Promenite stilove na osnovu podataka
challengeType: 6
forumTopicId: 301479
dashedName: change-styles-based-on-data
---

# --description--

D3 se tiče predstavljanja i vizualizacije podataka. Možda želite da promenite stil elemenata na osnovu podataka.
Na primer, možda želite da podelite plavom bojom tačku podataka ako je njena vrednost manja od 20, a crvenom u suprotnom. Možete koristiti funkciju povratnika (callback function) u `style()` metodi i uključiti uslovnu logiku. Funkcija povratnik koristi parametar `d` za predstavljanje tačke podataka:

```js
selection.style('color', d => {});
```

Metoda `style()` nije ograničena samo na postavljanje `color` - može se koristiti i za druge CSS atribute.

# --instructions--

Dodajte `style()` metodu u kod u editoru kako biste uslovno postavili `color` elemenata `h2`. Napišite funkciju povratnika koja će vratiti crvenu boju ako je vrednost podataka manja od 20, a zelenu u suprotnom.

**Napomena:** Možete koristiti if-else logiku ili ternarni operator.

# --hints--

Prvi `h2` treba da ima crvenu boju.

```js
const headingTwoFirst = document.querySelector('h2');
assert.exists(headingTwoFirst); 
const color = window.getComputedStyle(headingTwoFirst)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Drugi `h2` treba da ima zelenu boju.

```js
const headingTwoSecond = document.querySelectorAll('h2')[1];
assert.exists(headingTwoSecond); 
const color = window.getComputedStyle(headingTwoSecond)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Treći `h2` treba da ima zelenu boju.

```js
const headingTwoThird = document.querySelectorAll('h2')[2];
assert.exists(headingTwoThird); 
const color = window.getComputedStyle(headingTwoThird)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Četvrti `h2` treba da ima crvenu boju.

```js
const headingTwoFourth = document.querySelectorAll('h2')[3];
assert.exists(headingTwoFourth); 
const color = window.getComputedStyle(headingTwoFourth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Peti `h2` treba da ima zelenu boju.

```js
const headingTwoFifth = document.querySelectorAll('h2')[4];
assert.exists(headingTwoFifth); 
const color = window.getComputedStyle(headingTwoFifth)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Šesti `h2` treba da ima crvenu boju.

```js
const headingTwoSixth = document.querySelectorAll('h2')[5];
assert.exists(headingTwoSixth); 
const color = window.getComputedStyle(headingTwoSixth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Sedmi `h2` treba da ima zelenu boju.

```js
const headingTwoSeventh = document.querySelectorAll('h2')[6];
assert.exists(headingTwoSeventh); 
const color = window.getComputedStyle(headingTwoSeventh)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Osmi `h2` treba da ima crvenu boju.

```js
const headingTwoEighth = document.querySelectorAll('h2')[7];
assert.exists(headingTwoEighth); 
const color = window.getComputedStyle(headingTwoEighth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Deveti `h2` treba da ima crvenu boju.

```js
const headingTwoNinth = document.querySelectorAll('h2')[8];
assert.exists(headingTwoNinth); 
const color = window.getComputedStyle(headingTwoNinth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD');
    // Add your code below this line

    // Add your code above this line
  </script>
</body >
```

# --solutions--

```html
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD')
      .style('color', d => (d < 20 ? 'red' : 'green'));
  </script>
</body >
```