---
id: 587d7fa7367417b2b2512bc5
title: Radi sa promenljivim podacima u D3
challengeType: 6
forumTopicId: 301498
dashedName: work-with-dynamic-data-in-d3
---

# --description--
Dve poslednje izazova se tiču osnova za prikazivanje dinamičkih podataka korišćenjem D3 za metode `data()` i `enter()`. Ove metode uzimaju set podataka i, zajedno sa metodom `append()`, kreiraju novi element DOM za svaki element u setu podataka.

U prethodnom izazovu, kreirao si novi element `h2` za svaku stavku u redu podataka `dataset`, ali svi su imali isti sadržaj, `New Title`. To je zato što ne koristiš podatke povezane sa svakim elementom od `h2`.

Putanja D3 za `text()` može prihvatiti niz znakova ili funkciju povratnog poziva kao argument:

```js
selection.text(d => d);
```

U gornjem primeru, parametar ``d`` se odnosi na jedan element u skupu podataka kojem je izabranje vezano.

Koristeći ovaj primer kao kontekst, prvi element od `h2` je povezan sa 12, drugi element od `h2` je povezan sa 31, treći element od `h2` je povezan sa 22, i tako dalje.

# --instructions--
Promeni putanju `text()` tako da svaki element `h2` ima odgovarajuću vrednost iz kolone podataka `dataset`, zajedno sa jednim razmakom i nizom slova `USD`. Na primer, prva zaglavlja bi trebalo da budu `12 USD`.

# --hints--
Prvi naslov za `h2` treba da sadrži tekst `12 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[0]?.textContent, '12 USD');
```

Kichwa cha pili cha `h2` kinapaswa kuwa na maandishi `31 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[1]?.textContent, '31 USD');
```

Kichwa cha tatu cha `h2` kinapaswa kuwa na maandishi `22 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[2]?.textContent, '22 USD');
```

Kichwa cha nne cha `h2` kinapaswa kuwa na maandishi `17 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[3]?.textContent, '17 USD');
```

Kichwa cha tano cha `h2` kinapaswa kuwa na maandishi `25 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[4]?.textContent, '25 USD');
```

Kichwa cha sita cha `h2` kinapaswa kuwa na maandishi `18 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[5]?.textContent, '18 USD');
```

Kichwa cha saba cha `h2` kinapaswa kuwa na maandishi `29 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[6]?.textContent, '29 USD');
```

Kichwa cha nane cha `h2` kinapaswa kuwa na maandishi `14 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[7]?.textContent, '14 USD');
```

Kichwa cha tisa cha `h2` kinapaswa kuwa na maandishi `9 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[8]?.textContent, '9 USD');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      // Add your code below this line

      .text('New Title');

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => `${d} USD`);
  </script>
</body>
```
