id: 587d7fac367417b2b2512bdb
title: Postavljanje domena i raspona vrednosti na skali
challengeType: 6
forumTopicId: 301491
dashedName: set-a-domain-and-a-range-on-a-scale
---

# --description--

Po podrazumevanju, skale koriste identičnu relaciju. To znači da ulazna vrednost odgovara izlaznoj vrednosti. Međutim, skalama je moguće dati veći efekat i privlačniji izgled.

Recimo da set podataka ima vrednosti od 50 do 480. Ovo je opseg unosa za skalu, poznat kao <dfn>domena</dfn>.

Želite da rasporedite te tačke duž ose `x` u SVG-u, između jedinica 10 i jedinica 500. Ovo je raspon izlaza, poznat kao <dfn>raspon vrednosti</dfn>.

Metode `domain()` i `range()` postavljaju ove vrednosti za skalu. Обе metode uzimaju niz podataka sa najmanje dva elementa kao argumente. Evo primera:

```js
scale.domain([50, 480]);
scale.range([10, 500]);
scale(50);
scale(480);
scale(325);
scale(750);
d3.scaleLinear();
```

U ovom slučaju, sledeće vrednosti će biti prikazane u konzoli: `10`, `500`, `323.37` i `807.67`.

Napomena da skala koristi linearnu relaciju između vrednosti domene i raspona vrednosti kako bi odredila izlaz za određeni broj. Najniža vrednost u domeni (50) odgovara najnižoj vrednosti (10) u rasponu vrednosti.

# --instructions--

Kreirajte skalu i postavite njenu domenu na `[250, 500]` i raspon vrednosti na `[10, 150]`.

**Napomena:** Možete povezati metode `domain()` i `range()` na promenljivu `scale`.

# --hints--

Vaš kod bi trebalo da koristi metodu `domain()`.

```js
assert.match(code, /\.domain/g);
```

`scale` domena bi trebalo da bude postavljena na `[250, 500]`.

```js
assert.deepEqual(scale.domain(), [250, 500]);
```

Vaš kod bi trebalo da koristi metodu `range()`.

```js
assert.match(code, /\.range/g);
```

`scale` raspon bi trebalo da bude postavljen na `[10, 150]`.

```js
assert.deepEqual(scale.range(), [10, 150]);
```

Tekst u `h2` mora biti `-102`.

```js
assert.strictEqual(document.querySelector('h2')?.textContent, '-102');
```

# --seed--

## --seed-contents--

```html
<body >
  <script>
    // Add your code below this line
    const scale = d3.scaleLinear();



    // Add your code above this line
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body >
```

# --solutions--

```html
<body >
  <script>
    const scale = d3.scaleLinear();
    scale.domain([250, 500]);
    scale.range([10, 150]);
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body >
```