id: 587d7fa7367417b2b2512bc8
title: Dodaj klase sa D3
challengeType: 6
forumTopicId: 301473
dashedName: add-classes-with-d3
---

# --description--

Korišćenje više stilova unutar jedne linije na elementima HTML može biti teško za upravljanje, čak i za male aplikacije. Lakše je dodati klase elementima i kreirati stil te klase odjednom koristeći CSS pravila. D3 ima metodu `attr()` za dodavanje bilo kog atributa HTML elementu, uključujući ime klase.

Metoda `attr()` radi isto kao i `style()`. Prihvata vrednosti razdvojene zarezima i može koristiti funkciju povratne vrednosti (callback function). Evo primera dodavanja klase `container` selekciji:

```js
selection.attr('class', 'container');
```

Zapamtite da će argument `class` ostati isti svaki put kada želite dodati klasu, a jedino što će se menjati je argument `container`.

# --instructions--

Dodajte metodu `attr()` u kod u editoru i dodelite klasu `bar` elementima `div`.

# --hints--

Vaši elementi `div` bi trebalo da imaju klasu `bar`.

```js
assert.isTrue(document.querySelector('div')?.classList.contains('bar'));
```

Vaš kod bi trebalo da koristi metodu `attr()`.

```js
assert.match(code, /\.attr/g);
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
    // Add your code below this line



    // Add your code above this line
  </script>
</body >
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      // Add your code below this line
      .attr('class', 'bar');
    // Add your code above this line
  </script>
</body >
```