id: 587d7fa7367417b2b2512bc5
title: Radite sa dinamičkim podacima u D3
challengeType: 6
forumTopicId: 301498
dashedName: work-with-dynamic-data-in-d3
---

# --description--

Poslednja dva izazova tiču osnova prikazivanja dinamičkih podataka korišćenjem D3 sa metodama `data()` i `enter()`. Ove metode uzimaju set podataka i, zajedno sa metodom `append()`, kreiraju novi DOM element za svaki element u setu podataka.

U prethodnom izazovu, kreirali ste novi `h2` element za svaki element u nizu podataka `dataset`, ali su svi imali isti tekst: `New Title`. To je zato što niste koristili podatke povezane sa svakim `h2` elementom.

D3 metoda `text()` može uzeti niz znakova ili funkciju povratnika kao argument:

```js
selection.text(d => d);
```

U gornjem primeru, parametar `d` predstavlja jedan element u setu podataka sa kojim je selekcija povezana.

Koristeći ovaj primer kao kontekst, prvi `h2` element je povezan sa 12, drugi `h2` element je povezan sa 31, treći `h2` element je povezan sa 22 i tako dalje.

# --instructions--

Izmenite metodu `text()` tako da svaki `h2` element ima odgovarajuću vrednost iz niza podataka `dataset`, zajedno sa razmakom i nizom znakova `USD`. Na primer, prvi naslov bi trebalo da bude `12 USD`.

# --hints--

Prvi `h2` element bi trebalo da ima tekst `12 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[0]?.textContent, '12 USD');
```

Drugi `h2` element bi trebalo da ima tekst `31 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[1]?.textContent, '31 USD');
```

Treći `h2` element bi trebalo da ima tekst `22 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[2]?.textContent, '22 USD');
```

Četvrti `h2` element bi trebalo da ima tekst `17 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[3]?.textContent, '17 USD');
```

Peti `h2` element bi trebalo da ima tekst `25 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[4]?.textContent, '25 USD');
```

Šesti `h2` element bi trebalo da ima tekst `18 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[5]?.textContent, '18 USD');
```

Sedmi `h2` element bi trebalo da ima tekst `29 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[6]?.textContent, '29 USD');
```

Osmi `h2` element bi trebalo da ima tekst `14 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[7]?.textContent, '14 USD');
```

Deveti `h2` element bi trebalo da ima tekst `9 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[8]?.textContent, '9 USD');
```

# --seed--

## --seed-contents--

```html
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      // Add your code below this line

      .text('New Title');

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => `${d} USD`);
  </script>
</body>
```