id: 587d7fa7367417b2b2512bc6
title: Dodajte inline stil elementima
challengeType: 6
forumTopicId: 301475
dashedName: add-inline-styling-to-elements
---

# --description--

D3 vam omogućava dodavanje inline CSS stilova na dinamički generisanim elementima koristeći metodu `style()`.

Metoda `style()` prima par ključ-vrednost, razdvojene zarezima, kao argumente. Evo primera kako postaviti boju teksta selekcije na plavu:

```js
selection.style('color', 'blue');
```

# --instructions--

Dodajte metodu `style()` u kod u editoru kako biste učinili da sav prikazani tekst ima `font-family` od `verdana`.

# --hints--

Vaši elementi `h2` bi trebalo da imaju `font-family` od `verdana`.

```js
const headingTwo = document.querySelector('h2');
assert.exists(headingTwo); 
const fontFamily = window.getComputedStyle(headingTwo)['font-family'];
assert.strictEqual(fontFamily, 'verdana');
```

Vaš kod bi trebalo da koristi metodu `style()`.

```js
assert.match(code, /\.style/g);
```

# --seed--

## --seed-contents--

```html
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD');
    // Add your code below this line



    // Add your code above this line
  </script>
</body >
```

# --solutions--

```html
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD')
      .style('font-family', 'verdana');
  </script>
</body >
```