id: 587d7fa9367417b2b2512bcf
title: Dinamički promenite visinu svake trake
challengeType: 6
forumTopicId: 301486
dashedName: dynamically-change-the-height-of-each-bar
---

# --description--

Visinu svake trake možete postaviti na vrednost podatkovnog elementa u nizu podataka, kao što je dinamički postavljena vrednost `x`.

```js
selection.attr('property', (d, i) => {});
```

Ovde će `d` biti vrednost podatkovnog elementa, a `i` će biti indeks tog podatkovnog elementa u nizu podataka.

# --instructions--

Izmenite funkciju za odgovor svojstva `height` da vrati vrednost podataka pomnoženu sa 3.

**Napomena:** Zapamtite da množenje svih elemenata podataka sa jednim brojem povećava veličinu podataka (kao dodavanje slike). Ovo pomaže u vidljivom razlike između vrednosti traka u ovom primeru.

# --hints--

Prvi `rect` treba da ima visinu od `36`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('height'),
  '36'
);
```

Drugi `rect` treba da ima visinu od `93`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('height'),
  '93'
);
```

Treći `rect` treba da ima visinu od `66`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('height'),
  '66'
);
```

Četvrti `rect` treba da ima visinu od `51`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('height'),
  '51'
);
```

Peti `rect` treba da ima visinu od `75`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('height'),
  '75'
);
```

Šesti `rect` treba da ima visinu od `54`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('height'),
  '54'
);
```

Sedmi `rect` treba da ima visinu od `87`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('height'),
  '87'
);
```

Osmi `rect` treba da ima visinu od `42`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('height'),
  '42'
);
```

Deveti `rect` treba da ima visinu od `27`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('height'),
  '27'
);
```

# --seed--

## --seed-contents--

```html
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        // Add your code below this line


        
        
        // Add your code above this line
      });
  </script>
</body >
```

# --solutions--

```html
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        return d * 3;
      });
  </script>
</body >
```