id: 587d7fa9367417b2b2512bd0
title: Preokretirovanje elemenata SVG
challengeType: 6
forumTopicId: 301488
dashedName: invert-svg-elements
---

# --description--

Verovatno ste primetili da linijski grafikon izgleda kao da je okrenut naopako, ili je okrenut. To se dešava zbog načina na koji SVG koristi koordinate (x, y).

U SVG-u, početna tačka koordinata nalazi se u gornjem levom uglu. Koordinata `x` sa vrednošću 0 postavlja oblik na leve ivici SVG područja. Koordinata `y` sa vrednošću 0 postavlja oblik na gornjoj ivici SVG područja. Veće vrednosti `x` guraju pravougaonik udesno. Veće vrednosti `y` guraju pravougaonik nadole.

Da bi linije bile poravnate sa gornjom stranom, morate promeniti način na koji se meri koordinata `y`. Morate uzeti u obzir dužinu linije i ukupnu visinu SVG područja.

Visina SVG područja je 100. Ako imate podatke od 0 u setu, želeli biste da linija počne ispod SVG područja (a ne na vrhu). Da biste to uradili, koordinata `y` mora imati vrednost 100. Ako je vrednost podataka 1, počeli biste sa koordinatom `y` koja ima vrednost 100 kako bi se linija postavila dole. Zatim morate uzeti u obzir dužinu linije od 1, pa će konačna koordinata `y` biti 99.

Koordinata `y` koja je `y = heightOfSVG - heightOfBar` učiniće da linije izgledaju poravnate sa gornjom stranom.

# --instructions--

Izmenite funkciju povratne vrednosti za atribut `y` kako biste postavili linije poravnate sa gornjom stranom. Zapamtite da je visina linije 3 puta vrednost podataka `d`.

**Napomena:** Generalno, odnos je `y = h - m * d`, gde je `m` konstanta koja meri podatke.

# --hints--

Prvi `rect` treba da ima vrednost `y` od `64`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('y'),
  (h - dataset[0] * 3).toString()
);
```

Drugi `rect` treba da ima vrednost `y` od `7`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('y'),
  (h - dataset[1] * 3).toString()
);
```

Treći `rect` treba da ima vrednost `y` od `34`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('y'),
  (h - dataset[2] * 3).toString()
);
```

Četvrti `rect` treba da ima vrednost `y` od `49`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('y'),
  (h - dataset[3] * 3).toString()
);
```

Peti `rect` treba da ima vrednost `y` od `25`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('y'),
  (h - dataset[4] * 3).toString()
);
```

Šesti `rect` treba da ima vrednost `y` od `46`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('y'),
  (h - dataset[5] * 3).toString()
);
```

Sedmi `rect` treba da ima vrednost `y` od `13`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('y'),
  (h - dataset[6] * 3).toString()
);
```

Osmi `rect` treba da ima vrednost `y` od `58`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('y'),
  (h - dataset[7] * 3).toString()
);
```

Deveti `rect` treba da ima vrednost `y` od `73`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('y'),
  (h - dataset[8] * 3).toString()
);
```

# --seed--

## --seed-contents--

```html
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => {
        // Add your code below this line


        
        // Add your code above this line
      })
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```

# --solutions--

```html
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```