id: 587d7fab367417b2b2512bd8
title: Dodajte atribute elementima kruga
challengeType: 6
forumTopicId: 301471
dashedName: add-attributes-to-the-circle-elements
---

# --description--

Poslednji izazov kreirao je `circle` elemente za svaku tačku u `dataset`-u, i dodao ih na SVG. Međutim, D3 zahteva više informacija o poziciji i veličini svakog `circle`-a kako bi ga pravilno prikazao.

`circle` u SVG ima tri glavna atributa. Atributi `cx` i `cy` su koordinata. Govore D3-u gde da postavi _centar_ oblika na SVG. Poluprečnik (atribut `r`) određuje veličinu `circle`-a.

Slično kao što je za `rect` koordinata `y`, atribut `cy` za `circle` meren je od vrha SVG-a, a ne od dna.

Sva tri atributa mogu koristiti funkciju vezivanja da postave svoje vrednosti putem transformacija. Zapamtite da se sve metode koje slede nakon `data(dataset)` izvršavaju jednom za svaki element u `dataset`-u. Parametar `d` u funkciji vezivanja odnosi se na trenutni element u `dataset`-u, što je niz podataka za svaku tačku. Koristite uglate zagrade, kao što je `d[0]`, da biste dobili vrednost iz tog niza.

# --instructions--

Dodajte atribute `cx`, `cy` i `r` na `circle` elemente. Vrednost `cx` treba da bude prvi broj u nizu svakog elementa u `dataset`-u. Vrednost `cy` treba da zavisi od drugog broja u nizu, ali obavezno osigurajte da grafikon prikazuje u normalnom smeru, a ne da je okrenut. Vrednost `r` treba da bude `5` za sve krugove.

# --hints--

Vaš kod bi trebalo da ima 10 `circle` elemenata.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

Prvi `circle` element bi trebalo da ima vrednost `cx` od `34`, vrednost `cy` od `422` i vrednost `r` od `5`.

```js
const circle1 = document.querySelector('circle');
assert.strictEqual(circle1?.getAttribute('cx'), '34');
assert.strictEqual(circle1?.getAttribute('cy'), '422');
assert.strictEqual(circle1?.getAttribute('r'), '5');
```

Drugi `circle` element bi trebalo da ima vrednost `cx` od `109`, vrednost `cy` od `220` i vrednost `r` od `5`.

```js
const circle2 = document.querySelectorAll('circle')[1];
assert.strictEqual(circle2?.getAttribute('cx'), '109');
assert.strictEqual(circle2?.getAttribute('cy'), '220');
assert.strictEqual(circle2?.getAttribute('r'), '5');
```

Treći `circle` element bi trebalo da ima vrednost `cx` od `310`, vrednost `cy` od `380` i vrednost `r` od `5`.

```js
const circle3 = document.querySelectorAll('circle')[2];
assert.strictEqual(circle3?.getAttribute('cx'), '310');
assert.strictEqual(circle3?.getAttribute('cy'), '380');
assert.strictEqual(circle3?.getAttribute('r'), '5');
```

Četvrti `circle` element bi trebalo da ima vrednost `cx` od `79`, vrednost `cy` od `89` i vrednost `r` od `5`.

```js
const circle4 = document.querySelectorAll('circle')[3];
assert.strictEqual(circle4?.getAttribute('cx'), '79');
assert.strictEqual(circle4?.getAttribute('cy'), '89');
assert.strictEqual(circle4?.getAttribute('r'), '5');
```

Peti `circle` element bi trebalo da ima vrednost `cx` od `420`, vrednost `cy` od `280` i vrednost `r` od `5`.

```js
const circle5 = document.querySelectorAll('circle')[4];
assert.strictEqual(circle5?.getAttribute('cx'), '420');
assert.strictEqual(circle5?.getAttribute('cy'), '280');
assert.strictEqual(circle5?.getAttribute('r'), '5');
```

Šesti `circle` element bi trebalo da ima vrednost `cx` od `233`, vrednost `cy` od `355` i vrednost `r` od `5`.

```js
const circle6 = document.querySelectorAll('circle')[5];
assert.strictEqual(circle6?.getAttribute('cx'), '233');
assert.strictEqual(circle6?.getAttribute('cy'), '355');
assert.strictEqual(circle6?.getAttribute('r'), '5');
```

Sedmi `circle` element bi trebalo da ima vrednost `cx` od `333`, vrednost `cy` od `404` i vrednost `r` od `5`.

```js
const circle7 = document.querySelectorAll('circle')[6];
assert.strictEqual(circle7?.getAttribute('cx'), '333');
assert.strictEqual(circle7?.getAttribute('cy'), '404');
assert.strictEqual(circle7?.getAttribute('r'), '5');
```

Osmi `circle` element bi trebalo da ima vrednost `cx` od `222`, vrednost `cy` od `167` i vrednost `r` od `5`.

```js
const circle8 = document.querySelectorAll('circle')[7];
assert.strictEqual(circle8?.getAttribute('cx'), '222');
assert.strictEqual(circle8?.getAttribute('cy'), '167');
assert.strictEqual(circle8?.getAttribute('r'), '5');
```

Deveti `circle` element bi trebalo da ima vrednost `cx` od `78`, vrednost `cy` od `180` i vrednost `r` od `5`.

```js
const circle9 = document.querySelectorAll('circle')[8];
assert.strictEqual(circle9?.getAttribute('cx'), '78');
assert.strictEqual(circle9?.getAttribute('cy'), '180');
assert.strictEqual(circle9?.getAttribute('r'), '5');
```

Deseti `circle` element bi trebalo da ima vrednost `cx` od `21`, vrednost `cy` od `377` i vrednost `r` od `5`.

```js
const circle10 = document.querySelectorAll('circle')[9];
assert.strictEqual(circle10?.getAttribute('cx'), '21');
assert.strictEqual(circle10?.getAttribute('cy'), '377');
assert.strictEqual(circle10?.getAttribute('r'), '5');
```

# --seed--

## --seed-contents--

```html
<body >
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
    // Add your code below this line



    // Add your code above this line
  </script>
</body >
```

# --solutions--

```html
<body >
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => d[0])
      .attr('cy', d => h - d[1])
      .attr('r', 5);
  </script>
</body >
```