id: 587d7fad367417b2b2512bdf
title: Dodajte ose na vizuelizaciju
challengeType: 6
forumTopicId: 301472
dashedName: add-axes-to-a-visualization
---

# --description--

Drugi način za poboljšanje grafika raspršenih tačaka je dodavanje x ose i y ose.

D3 ima dve metode, `axisLeft()` i `axisBottom()`, za prikaz y ose i x ose, respektivno. Evo primera kako kreirati x osu koristeći `xScale` iz prethodnih zadataka:

```js
const xAxis = d3.axisBottom(xScale);
```

Sledeći korak je prikazivanje ose na SVG-u. Za to, možete koristiti ceo SVG deo, element `g`. `g` označava grupu. Za razliku od `rect`, `circle` i `text`, osa je samo crvena linija kada se prikaže. Pošto je jednostavna forma, korišćenje `g` funkcioniše. Poslednji korak je korišćenje atributa `transform` za pozicioniranje ose na pravilnoj lokaciji u SVG-u. Inače, linija bi bila prikazana pored ivice SVG-a i nije bila vidljiva. SVG podržava različite tipove `transforms`, ali pozicioniranje ose zahteva `translate`. Kada se primeni na element `g`, pomera celu grupu sa strane i nadole za zadatu količinu. Evo primera:

```js
const xAxis = d3.axisBottom(xScale);

svg
  .append('g')
  .attr('transform', 'translate(0, ' + (h - padding) + ')')
  .call(xAxis);
```

Kod pomenut iznad postavlja x osu ispod SVG-a. Zatim se prosleđuje kao argument pomoću `call()`. Y osa radi na isti način, osim što je argument `translate` u formi `(x, 0)`. Pošto je `translate` niz karaktera u atributu `attr()` iznad, možete koristiti konkatenaciju za kombinovanje vrednosti transformacije za njegove argumente.

# --instructions--

Grafik raspršenih tačaka sada ima x osu. Kreirajte y osu u promenljivoj pod nazivom `yAxis` koristeći metodu `axisLeft()`. Zatim prikažite ozu korišćenjem elementa `g`. Obavezno koristite atribut `transform` da transformišete osovu za količinu jedinica unutrašnjeg razmaka sa desne strane i 0 jedinica nadole. Ne zaboravite da pozovete (`call()`) osovu.

# --hints--

Vaš kod bi trebalo da koristi metodu `axisLeft()` sa `yScale` prosleđenom kao argument.

```js
assert.match(code, /\.axisLeft\(yScale\)/g);
```

Element `g` za y osu treba da ima atribut `transform` koji transformiše ozu na `(60, 0)`.

```js
const yAxis = document.querySelectorAll('g')[10];
assert.match(yAxis?.getAttribute('transform'), /translate\(60\s*,\s*0\)/g);
```

Vaš kod bi trebalo da pozove `yAxis`.

```js
assert.match(code, /\.call\(\s*yAxis\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<body >
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);
    // Add your code below this line
    const yAxis = undefined;
    // Add your code above this line

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    // Add your code below this line



    // Add your code above this line
  </script>
</body >
```

# --solutions--

```html
<body >
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);

    const yAxis = d3.axisLeft(yScale);

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    svg
      .append('g')
      .attr('transform', 'translate(' + padding + ',0)')
      .call(yAxis);
  </script>
</body >
```