id: 587d7fa8367417b2b2512bcd
title: Kreiranje trake za svaku tačku podataka u skupu
challengeType: 6
forumTopicId: 301482
dashedName: create-a-bar-for-each-data-point-in-the-set
---

# --description--

Prethovno zadatje je dodalo samo jedan pravougaonik na `svg` element da ga predstavlja. Ovde, spojićete ono što ste naučili do sada o `data()`, `enter()` i SVG oblicima kako biste kreirali i dodali pravougaonik za svaku tačku podataka u `dataset`.

Prethodno zadatje je pokazalo strukturu kako da se kreira i doda `div` za svaki element u `dataset`:

```js
d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
```

Postoji nekoliko razlika kada koristite elemente `rect` umesto elemenata `div`. Elementi `rect` moraju se dodati na `svg` element, a ne direktno na `body`. Takođe, morate reći D3 gde da postavi svaki `rect` unutar područja `svg`. Pozicioniranje trake će biti obrađeno u sledećem zadatku.

# --instructions--

Koristite metode `data()`, `enter()` i `append()` za kreiranje i dodavanje `rect` za svaki element u `dataset`. Sve trake treba da se prikažu preko svake; ovo će biti popravljeno u sledećem zadatku.

# --hints--

Vaš dokument bi trebalo da ima 9 elemenata `rect`.

```js
assert.lengthOf(document.querySelectorAll('rect'), 9);
```

Vaš kod bi trebalo da koristi metodu `data()`.

```js
assert.match(code, /\.data/g);
```

Vaš kod bi trebalo da koristi metodu `enter()`.

```js
assert.match(code, /\.enter/g);
```

Vaš kod bi trebalo da koristi metodu `append()`.

```js
assert.match(code, /\.append/g);
```

# --seed--

## --seed-contents--

```html
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      // Add your code below this line

      // Add your code above this line
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body >
```

# --solutions--

```html
<body >
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body >
```