---
id: 67172b43f84bcd2dec238a3d
title: Napravite organizator knjiga
challengeType: 26
dashedName: build-a-book-organizer
---

# --description--

**Cilj:** Ispunite korisničke priče ispod i osigurajte da svi testovi prođu kako biste završili laboratoriju.

**Korisničke priče:**

1. Trebalo bi da imate niz objekata pod nazivom `books` gde svaki objekat u nizu treba da ima string `title`, drugi string `authorName` i broj `releaseYear`. 

2. Vaš `books` niz bi trebalo da ima minimum od tri objekta.

3. Trebalo bi da imate povratnu funkciju pod nazivom `sortByYear` koja prima dva romana kao parametre za sortiranje niza.

4. Funkcija `sortByYear` bi trebalo da vrati `-1` ako je `releaseYear` prvog romana manji od onoga drugog romana.

5. Funkcija `sortByYear` bi trebalo da vrati `1` ako je `releaseYear` prvog romana veći od onoga drugog romana.

6. Funkcija `sortByYear` bi trebalo da vrati `0` ako su obe vrednosti `releaseYear` jednake.

7. Trebalo bi da filtrirate knjige napisane nakon određenogodišta kao što je 1950 iz niza `books` i sačuvate filtrirani niz u novi niz pod nazivom `filteredBooks`.

8. Trebalo bi da sortirate knjige u nizu `filteredBooks` prema njihovoj `releaseYear` u uzlaznom redosledu. Naučili ste u prethodnoj lekciji da će metoda `sort()` sortirati niz na mestu (in place). To znači da će se mutirati niz `filteredBooks`.

# --hints--

You should have a function `sortByYear` in your code.

```js
assert.isFunction(sortByYear);
```

Your `sortByYear` function should take two parameters.

```js
assert.lengthOf(sortByYear, 2);
```

Your `sortByYear` function should return `-1` if the `releaseYear` of `book1` object is smaller than that of the `book2` object, `1` if the `releaseYear` of `book1` object is larger than that of the `book2` object, and `0` in all other scenarios.

```js
assert.equal(sortByYear({releaseYear: 1913}, {releaseYear: 1925}), -1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1913}), 1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1925}), 0);
```

You should have an array `books` in your code.

```js
assert.isArray(books)
```

Your `books` array should have at least three book objects.

```js
assert.isAtLeast(books.length, 3)
```

Your `books` array should contain objects each with a string `title`, another string `authorName`, and a number `releaseYear`.

```js

books.forEach((book) => {
  assert.isObject(book);
  assert.sameMembers(Object.keys(book), ["title", "authorName", "releaseYear"]);
  assert.isString(book.title);
  assert.isNotEmpty(book.title);
  assert.isString(book.authorName);
  assert.isNotEmpty(book.authorName);
  assert.isNumber(book.releaseYear);
})
```

You should have an array `filteredBooks` in your code.

```js
assert.isArray(filteredBooks)
```

The `filteredBooks` array should have some of the books filtered out from the `books` array and not be empty.

```js
assert.isBelow(filteredBooks.length, books.length)
assert.isNotEmpty(filteredBooks)
```

The `filteredBooks` array should only contain books from the original `books` array.

```js
filteredBooks.forEach(book => {
  assert.isTrue(
    books.some(original => 
      original.title === book.title &&
      original.authorName === book.authorName &&
      original.releaseYear === book.releaseYear
    )
  );
});
```

The `filteredBooks` array should include only books released on or before a specified year.

```js
const maxYear = Math.max(...filteredBooks.map(book => book.releaseYear));
const removedYears = books.filter(book => !filteredBooks.includes(book)).map(book => book.releaseYear);
assert.isTrue(removedYears.every(year => year > maxYear));
```

You should call the `sort` higher order function by passing the `sortByYear` callback function on the `filteredBooks` array.

```js
assert.match(__helpers.removeJSComments(code), /\s*filteredBooks\s*\.\s*sort\s*\(\s*sortByYear\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sortByYear(book1, book2) {
  if (book1.releaseYear < book2.releaseYear) {
    return -1;
  }

  if (book1.releaseYear > book2.releaseYear) {
    return 1;
  }

  return 0;
}

const books = [
  {
    title: "The Great Gatsby",
    authorName: "F. Scott Fitzgerald",
    releaseYear: 1925,
  },
  {
    title: "In Search of Lost Time",
    authorName: "Marcel Proust",
    releaseYear: 1913,
  },
  {
    title: "Ulysses",
    authorName: "James Joyce",
    releaseYear: 1922,
  },
  {
    title: "One Hundred Years of Solitude",
    authorName: "Gabriel García Márquez",
    releaseYear: 1967,
  },
  {
    title: "The Catcher in the Rye",
    authorName: "J. D. Salinger",
    releaseYear: 1951,
  },
];

let filteredBooks = books.filter((book) => book.releaseYear < 1950);

filteredBooks.sort(sortByYear);

console.log("Books Written Before 1950 (sorted according to release year)");
filteredBooks.forEach((book) => {
  console.log(`${book.title} by ${book.authorName} (${book.releaseYear})`);
});
```
