---
id: a6b0bb188d873cb2c8729495
title: Implementator za konverziju HTML entiteta
challengeType: 26
dashedName: implement-an-html-entity-converter
---

# --description--

Ovaj laboratorijum se bavi konverzijom specijalnih znakova u nizu sa njihovim odgovarajućim HTML entitetima.

**Cilj**: Ispunite priče korisnika ispod i učinite da svi testovi prođu kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da imate funkciju `convertHTML` koja prima niz kao argument.
1. Funkcija `convertHTML` bi trebalo da vrati novi niz konvertovanjem specijalnih znakova u argumentnom nizu u njihove odgovarajuće HTML entitete.

   - `&` bi trebalo da se konvertuje u `&amp;`.
   - `<` bi trebalo da se konvertuje u `&lt;`.
   - `>` bi trebalo da se konvertuje u `&gt;`.
   - `"` bi trebalo da se konvertuje u `&quot;`.
   - `'` bi trebalo da se konvertuje u `&apos;`.

# --hints--

You should have a `convertHTML` function.

```js
assert.isFunction(convertHTML);
```

`convertHTML("Dolce & Gabbana")` should return the string `Dolce &amp; Gabbana`.

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")` should return the string `Hamburgers &lt; Pizza &lt; Tacos`.

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sixty > twelve")` should return the string `Sixty &gt; twelve`.

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')` should return the string `Stuff in &quot;quotation marks&quot;`.

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("Schindler's List")` should return the string `Schindler&apos;s List`.

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")` should return the string `&lt;&gt;`.

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")` should return the string `abc`.

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const map = {
  '&': '&amp;',
  '<': '&lt;',
  '>': '&gt;',
  '"': '&quot;',
  "'": '&apos;'
};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(char) {
    return map[char];
  });
}
```
