---
id: 68ba9c4f1688914d72876098
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Sledeći atribut koji ćete dodati je `referrerpolicy`. To je pravilo koje određuje koliko detalja delite kada vaša stranica povezuje sa drugom stranicom.

Dodajte atribut `referrerpolicy` i podesite ga na `strict-origin-when-cross-origin`. Ovo deli punu adresu na istoj lokaciji (sajtu), samo naziv sajta na drugim sajtovima i ništa na nesigurnim sajtovima.

# --hints--

Your `iframe` element should have a `referrerpolicy` attribute set to `strict-origin-when-cross-origin`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('referrerpolicy'), 'strict-origin-when-cross-origin')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write encrypted-media gyroscope web-share"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
