---
id: 68ba9c4f1688914d72876097
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Dodajte `encrypted-media`, `gyroscope` i `web-share` postojećim vrednostima u atributu `allow`.

Ova tri će omogućiti korišćenje proširenja za šifrovane medije radi zaštite videa, pružanje pristupa senzorima pokreta i orijentacije uređaja, i omogućavanje deljenja sadržaja iframe-a putem nativnih dijaloga za deljenje uređaja.

# --hints--

The `allow` attribute of your `iframe` element should have `encrypted-media` as one of its values.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /encrypted-media/)
```

The `allow` attribute of your `iframe` element should have `gyroscope` as one of its values.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /gyroscope/)
```

The `allow` attribute of your `iframe` element should have `web-share` as one of its values.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /web-share/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      allow="accelerometer autoplay clipboard-write"
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
