---
id: 68ba9c4f1688914d72876095
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Element ``iframe`` takođe prima atribut ``src`` sa vrednošću koja ukazuje na URL ili putanju resursa koji treba prikazati.

Dodajte atribut ``src`` sa vrednošću ``https://www.youtube.com/embed/I0_951_MPE0`` vašem elementu ``iframe``.

U ovom trenutku, trebalo bi da vidite video koji se prikazuje na stranici, ali morate dodati još neke atribute.

# --hints--

Your `iframe` element should have an `src` attribute set to `https://www.youtube.com/embed/I0_951_MPE0`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('src'), 'https://www.youtube.com/embed/I0_951_MPE0')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
