---
id: 68ba9c4f1688914d72876096
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

Jedan od atributa je `allow`. To je kao lista dozvola koja pregledaču govori koje funkcije `iframe` sme da koristi.

Evo `iframe` elementa sa atributom `allow`:

```html
<iframe
  allow="accelerometer autoplay clipboard-write encrypted-media gyroscope picture-in-picture web-share"
></iframe>
```

Dodajte atribut `allow` sa vrednostima `accelerometer`, `autoplay` i `clipboard-write`.

`accelerometer` omogućava da `iframe` koristi senzore pokreta kako bi detektovao stvari kao što su naginjanje i rotacija uređaja. `autoplay` omogućava video snimku da počne automatsko reprodukovanje, a `clipboard-write` omogućava iframe-u da piše podatke u međupokloški sistem korisnika.

# --hints--

Your `iframe` element should have an `allow` attribute set to `accelerometer`.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /accelerometer/)
```

The `allow` attribute of your `iframe` element should have `autoplay` as one of its values.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /autoplay/)
```

The `allow` attribute of your `iframe` element should have `clipboard-write` as one of its values.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /clipboard-write/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
