---
id: 68ba9c4f1688914d72876094
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Na prvom času o elementu `iframe`, naučili ste da je to zamenljiv element baš kao i `img`. To znači da može prihvatiti i svojstva `width` i `height` kako bi odredilo koliko treba da bude visoko i široko.

Dajte svom elementu `iframe` svojstvo `width` sa vrednošću `560` i svojstvo `height` sa vrednošću `315`.

# --hints--

Your `iframe` element should have a `width` attribute set to `560`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('width'), '560')
```

Your `iframe` element should have a `height` attribute set to `315`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('height'), '315')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>

    <iframe
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
