---
id: 68ba9c4f1688914d72876093
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Sada, kreiraj element `iframe`. Nemoj ga još ništa popuniti.

# --hints--

You should have an `iframe` element.

```js
assert.exists(document.querySelector('iframe'))
```

Your `iframe` element should be empty.

```js
assert.equal(document.querySelector('iframe')?.innerHTML.trim(), '')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
