---
id: 6724e2dbf723fe1c8883cc69
title: Pregled front-end biblioteka
challengeType: 31
dashedName: review-front-end-libraries
---

# --description--
## JavaScript Biblioteke i Okviri

Biblioteke i framework-ovi za JavaScript nude brza rešenja za uobičajene probleme i ubrzavaju razvoj pružanjem predizgrađenog koda.
Biblioteke su generalno više fokusirane na pružanje rešenja za specifične zadatke, kao što je manipulisanje DOM-a, rukovanje događajima (events) ili upravljanje AJAX zahtevima. 
Neki primeri biblioteka za JavaScript su jQuery i React.
Framework-ovi, s druge strane, pružaju definisanu strukturu za izgradnju aplikacija. Često dolaze sa setom pravila i konvencija koje programeri moraju pratiti.
Primeri framework-ova uključuju Angular i Next.js, koji je meta framework za React.
**Jednostrani aplikacije** (Single-page applications - SPAs) su web aplikacije koje učitavaju jednu HTML stranicu i dinamički ažuriraju tu stranicu kako korisnik interaguje sa aplikacijom bez ponovnog učitavanja cele stranice.
SPAs koriste JavaScript za upravljanje stanjem (state) aplikacije i prikaz sadržaja (render content). Ovo se često radi korišćenjem framework-ova koji pružaju odlične alate za izgradnju složenih korisničkih interfejsa.
Neki problemi vezani za SPAs uključuju:

- Čitači ekrana imaju poteškoća sa dinamički ažuriranim sadržajem.
- URL se ne menja kada korisnik navigira unutar aplikacije, što može otežati označavanje, povratak ili deljenje određenih stranica.
- Vreme inicijalnog učitavanja može biti sporo ako je aplikacija velika, jer su svi resursi potrebni za učitavanje odjednom.

## React

React je popularna JavaScript biblioteka za izgradnju korisničkih interfejsa i veb aplikacija.
Osnovni koncept Reacta je kreiranje ponovno upotrebljivih UI komponenti koje mogu da se ažuriraju i renderuju nezavisno kada se podaci promene.
React omogućava programerima da opišu kako bi UI trebalo da izgleda na osnovu stanja aplikacije (application state). Zatim, React ažurira i renderuje prave komponente kada se podaci ili stanje (state) promene.

## React Komponente

Komponente su građevni blokovi React aplikacija koji omogućavaju razvojnim programerima da rastvore kompleksne korisničke interfejse u manje, upravljive delove.
UI se opisuje koristeći JSX, proširenje JavaScript sintakse, koje omogućava razvojnim programerima pisanje koda sličnog HTML unutar JavaScripta.
Komponente su u suštini JS funkcije ili klase koje vraćaju deo UI-a.

Evo primera jednostavnog React komponente koja renderuje poruku pozdrava:

```jsx
function Greeting() {
  const name = 'Anna';
  return <h1>Welcome, {name}!</h1>;
}
```

Da biste koristili komponentu, možete jednostavno pozvati:
    
```jsx
  <Greeting />
```

## Uvoz i izvoz komponenti React-a

- Eksport čini komponentu dostupnom drugim fajlovima. Import omogućava fajlu da koristi komponentu koja je eksportovana negde drugde.
- U React-u postoji dva tipa eksporta: podrazumevani eksporti (default exports) i imenovani eksporti (named exports).

**Podrazumevani izvozi**

- Datoteka može imati samo jedan podrazumevani izvoz (default export). Idealno je za datoteku koja primarno sadrži jednu komponentu.
- Komponentu možete izvesti kao podrazumevani izvoz na ovaj način:

```jsx
// City.js
function City() {
  return <p>New York</p>;
}

export default City;
```

- Takođe možete eksportovati na istoj liniji kao definicija komponente:

```jsx
export default function City() {
  return <p>New York</p>;
}
```

- Da biste uvezli podrazumevani izvoz, koristite ključnu reč `import` bez vitičastih zagrada:

```jsx
// App.js
import City from './City';

function App() {
  return (
    <div>
      <h1>My favorite city is:</h1>
      <City />
    </div>
  );
}
```

**Imenovani Izvozi**

* Named exports omogućavaju fajlu da deli više komponenti. Za razliku od default eksporta, moraju se uvoziti koristeći tačno ime pod kojim su eksportovani.
* Možete eksportovati više komponenti iz istog fajla na ovaj način:

```jsx
// Animals.jsx
export function Cat() {
  return <h2>Mr. Whiskers</h2>;
}

export function Dog() {
  return <h2>Fido</h2>;
}
```

- Za uvoz nazvanih eksportova, koristite vitičaste zagrade sa tačnim nazivima izvoženja:

```jsx
import { Cat, Dog } from './Animals';
```

- Možete preimenovati imenovani izvoz prilikom uvoza koristeći ključnu reč `as`:

```jsx
import { Cat as Kitty } from './Animals';
```

**Mešoviti podrazumevani i imenovani izvozi**

Fajl može imati jedan podrazumevani izvoz i više imenovanih izvoza istovremeno:

```jsx
// Animals.jsx
export default function Cat() {
  return <h2>Mr. Whiskers</h2>;
}

export function Dog() {
  return <h2>Fido</h2>;
}
```

- Prilikom uvoza mešovitog eksporta, podrazumevani eksport dolazi prvi (bez vitičastih zagrada), praćen imenovanima eksportima (unutar vitičastih zagrada):

```jsx
import Cat, { Dog } from './Animals';
```

## Postavljanje projekta sa Reactom korišćenjem Vite-a

Alati za podešavanje projekata i CLI-ji pružaju brz i jednostavan način za pokretanje novih projekata, omogućavajući programerima da se fokusiraju na pisanje koda umesto da se bave konfiguracijom.
- Vite je popularan alat za podešavanje projekata i može se koristiti sa React-om.
- Da biste kreirali novi projekat sa Vite-om, možete koristiti sledeću komandu u vašem terminalu:

```bash
npm create vite@latest my-react-app -- --template react
```

Ova komanda kreira novi React projekat pod nazivom `my-react-app` koristeći Vite-ov React šablon. U direktorijumu projekta, videćete fajl `package.json` sa zavisnostima i komandama projekta navedenim u njemu.

- Da biste pokrenuli projekat, navigirajte do direktorijuma projekta i izvršite sledeće komande:

```bash
cd my-react-app # path to the project directory
npm install # installs the dependencies listed in the package.json file
```

- Nakon instalacije zavisnosti, trebalo bi da primetite novu fasciklu u vašem projektu nazvanu `node_modules`.

- U folderu `node_modules` se čuvaju svi paketi i biblioteke koje je vaš projekat zahteva.

- Da biste pokrenuli svoj projekat, koristite sledeću komandu:

```bash
npm run dev
```

- Nakon toga, otvorite svoj pretraživač i navigujte do ``http://localhost:5173`` da biste videli kako vaša React aplikacija radi.

Da biste zaista videli kod za početni šablon, možete u svoj projekat, unutar foldera `src`, i trebalo bi da vidite fajl `App.jsx`.

## Prosleđivanje props-ova u React komponentama

- U React-u, `props` (skraćeno od *properties*, odnosno svojstava) predstavljaju način za prenošenje podataka sa roditeljske komponente na dete komponentu. Ovaj mehanizam je potreban za kreiranje ponovo upotrebljivih i dinamičkih UI elemenata.
- `Props` mogu biti bilo koja JavaScript vrednost. Da biste preveli `props` sa roditeljske na dete komponentu, dodajete ih kao atribute kada koristite detsku komponentu u JSX-u roditelja. Evo jednostavnog primera:

```jsx
// Parent component
function Parent() {
  const name = 'Anna';
  return <Child name={name} />;
}

// Child component
function Child(props) {
  return <h1>Hello, {props.name}!</h1>;
}
```

Možete proslediti više props-ova koristeći spread operator ``(...)``, nakon što ih konvertujete u objekat. Evo primera:
    
````jsx
// Parent component
function Parent() {
  const person = {
    name: 'Anna',
    age: 25,
    city: 'New York'
  };
  return <Child {...person} />;
}
````

U ovom kodu, operator za širenje (spread operator) `{...person}` konvertuje objekat `person` u pojedinačne props-ove koji se prosleđuju komponenti `Child`.

## Uslovno renderovanje u React-u

- Uslovno renderovanje u React-u omogućava vam da kreirate dinamičke korisničke interfejse. Koristi se za prikaz različitog sadržaja na osnovu određenih uslova ili stanja unutar vaše aplikacije.
- Postoji nekoliko načina za uslovno renderovanje sadržaja u React-u. Jedan uobičajen pristup je korišćenje ternarnog operatera. Evo primera:

```jsx
function Greeting({ isLoggedIn }) {
  return (
    <div>
      {isLoggedIn ? <h1>Welcome back!</h1> : <h1>Please log in</h1>}
    </div>
  );
}
```

- Drugi način za uslovno prikazivanje sadržaja je korišćenje logičkog AND (`&&`) operatera. Ovo je korisno kada želite da prikažete sadržaj samo ako je zadovoljen određeni uslov. Evo primera:

```jsx
function Greeting({ user }) {
  return (
    <div>
      {user && <h1>Welcome, {user.name}!</h1>}
    </div>
  );
}
```

U kodu iznad, element `h1` se renderuje samo ako je objekat korisnika istinit.

Takođe možete koristiti direktnu `if` izjavu na ovaj način:

```js
function Greeting({ isLoggedIn }) {
  if (isLoggedIn) {
    return <h1>Welcome back!</h1>;
  }
  return <h1>Please sign in</h1>;
}
```

## Renderovanje lista u Reactu

Renderovanje lista u Reactu je uobičajen zadatak prilikom izgradnje korisničkih interfejsa.
Liste se mogu renderovati korišćenjem metode JS array `map()` za iteraciju preko niza stavki i vraćanje novog niza JSX elemenata.
Na primer, ako imate niz imena koje želite da renderujete kao listu, možete uraditi sledeće:

```jsx
function NameList({ names }) {
  return (
    <ul>
      {names.map((name, index) => (
        <li key={`${name}-${index}`}>{name}</li>
      ))}
    </ul>
  );
}
```

Uvek zapamtite da pružite jedinstveni ključ za svaku stavku liste kako biste pomogli React-u u upravljanju procesima ažuriranja i renderovanja. Ovim tehnikama, možete kreirati fleksibilne, efikasne i dinamične liste u vašim React aplikacijama.

## Inline stili u React-u

- Inline stili u React-u omogućavaju vam da primenite CSS stilove direktno na JSX elemente koristeći JavaScript objekte.
- Da biste primenili inline stilove u React-u, možete koristiti `style` atribut na JSX elementima. `style` atribut prima objekat gde su ključevi CSS svojstva u formatu camelCase, a vrednosti su odgovarajuće vrednosti. Evo primera:

```js
function Greeting() {
  return (
    <h1
      style={{ color: 'blue', fontSize: '24px', backgroundColor: 'lightgray' }}
    >
      Hello, world!
    </h1>
  );
}

export default Greeting;
```

Takođe možete izvući stilove u zaseban objekat i referencirati ih u atributu `style` na ovaj način:

```jsx
function Greeting() {

  const styles = {
    color: 'blue',
    fontSize: '24px',
    backgroundColor: 'lightgray'
  };

  return <h1 style={styles}>Hello, world!</h1>;
}

export default Greeting;
```

Inline styles podržavaju dinamičko stilizovanje omogućavanjem uslovnog primenjivanja stilova na osnovu `props` ili `state`. Evo primera kako možete uslovno primeniti stilove na osnovu `prop`-a:

```jsx
function Greeting({ isImportant }) {

  const styles = {
    color: isImportant ? 'red' : 'black',
    fontSize: isImportant ? '24px' : '16px'
  };

  return <h1 style={styles}>Hello, world!</h1>;
}

export default Greeting;
```

U kodu iznad, stili `color` i `fontSize` se uslovno postavljaju na osnovu `isImportant` prop-a.

## Rad sa događajima u React-u

- **Sistem sintetičkih događaja**: Ovo je način na koji React obrađuje događaje. On služi kao omotač (wrapper) za naivne događaje poput `click`, `keydown` i `submit` događaja. Rukovaoci događajima u React-u koriste konvenciju imenovanja camel case. (Primer: `onClick`, `onSubmit`, itd)

Ovo je primer korišćenja atributa ``onClick`` za element ``button`` u React-u:

```jsx
function handleClick() {
  console.log("Button clicked!");
}

<button onClick={handleClick}>Click Me</button>;
```

U React-u, funkcije za rukovanje događajima obično počinju sa prefiksom `handle` kako bi ukazale da su odgovorne za rukovanje događajima, kao što su `handleClick` ili `handleSubmit`.

Kada akcija korisnika pokrene događaj, React prosleđuje Sintetički objekat događaja vašem handler-u. Ovaj objekat se ponaša veoma slično naivnom Event objektu u vanilla JavaScript-u, pružajući svojstva kao što su `type`, `target` i `currentTarget`.

Kako biste sprečili podrazumevano ponašanje, kao što je osvežavanje pregledača tokom događaja `onSubmit`, na primer, možete pozvati metodu `preventDefault()`:

```jsx
function handleSubmit(event) {
  event.preventDefault();
  console.log("Form submitted!");
}

<form onSubmit={handleSubmit}>
  <input type="text" />
  <button>Submit</button>
</form>;
```

Takođe možete obaviti funkciju hendlera u strelicastu funkciju na ovaj način:

```jsx
function handleDelete(id) {
  console.log("Deleting item:", id);
}

<button onClick={() => handleDelete(1)}>Delete Item</button>;
```

## Rad sa stanjem i hook-om `useState`

- **Definicija stanja (state)**: U Reactu, stanje je objekat koji sadrži podatke za komponentu. Kada se stanje ažurira, komponenta će ponovo biti renderovana. React tretira stanje kao nemutabilno (immutable), što znači da ga ne smete menjati direktno.
- **``useState()` Hook`**: ``useState`` hook je funkcija koja vam omogućava deklarisanje varijabli stanja u funkcionalnim komponentama. Evo osnovne sintakse:

```js
const [stateVariable, setStateFunction] = useState(initialValue);
```

U varijabli stanja imate sledeće:

- `stateVariable` drži trenutnu vrednost stanja
- `setStateFunction` (funkcija za postavljanje) ažurira varijablu stanja
- `initialValue` postavlja početno stanje

Evo kompletnog primera za `Counter` komponentu:

```jsx
import { useState } from "react";

function Counter() {
  const [count, setCount] = useState(0);

  return (
    <div>
      <h2>{count}</h2>

      <button onClick={() => setCount(count - 1)}>Decrement</button>
      <button onClick={() => setCount(count + 1)}>Increment</button>
    </div>
  );
}

export default Counter;
```

## Renderovanje i React Komponente

- **Definicija**: U React-u, renderovanje je proces kojim se komponente pojavljuju u korisničkom interfejsu (UI), obično u pregledaču. Proces renderovanja sastoji se od tri faze: trigger, render i commit.

Faza okidača nastaje kada React detektuje da se nešto promenilo i da je korisnički interfejs (UI) možda potreban za ažuriranje. Ova promena je često uzrokovana ažuriranjem stanja (state) ili propova (props).

Kada se okidač (trigger) dogodi, React ulazi u fazu renderovanja. Ovde, React ponovo procenjuje vaše komponente i utvrđuje šta treba prikazati. Da bi to uradio, React koristi laganu kopiju "pravog" DOM-a nazvanu Virtual DOM. Uz pomoć Virtual DOM-a, React može brzo da proveri šta treba da se promeni u komponenti.

Commit faza je mesto gde React uzima pripremljene promene iz virtual DOM-a i primenjuje ih na real DOM. Drugim rečima, ovo je faza gde vidite konačan rezultat na ekranu.

## Ažuriranje objekata i nizova u stanju

**Ažuriranje objekata u stanju (State)**: Ako trebate ažurirati objekat u stanju, prvo je trebalo napraviti novi objekat ili kopirati postojeći objekat, a zatim postaviti stanje za taj novi objekat. Svaki objekat koji se postavlja u stanje treba smatrati kao samočitan (read-only). Evo primera kako postaviti ime, godine i grad korisnika. Funkcija ``handleChange`` koristi se za rukovanje ažuriranjima informacija korisnika:

```jsx
import { useState } from "react";

function Profile() {
  const [user, setUser] = useState({ name: "John Doe", age: 31, city: "LA" });

  const handleChange = (e) => {
    const { name, value } = e.target;

    setUser((prevUser) => ({...prevUser, [name]: value}));
  };

  return (
    <div>
      <h1>User Profile</h1>
      <p>Name: {user.name}</p>
      <p>Age: {user.age}</p>
      <p>City: {user.city}</p>

      <h2>Update User Age </h2>
      <input type="number" name="age" value={user.age} onChange={handleChange} />

      <h2>Update User Name </h2>
      <input type="text" name="name" value={user.name} onChange={handleChange} />

      <h2>Update User City </h2>
      <input type="text" name="city" value={user.city} onChange={handleChange} />
    </div>
  );
}

export default Profile;
```

- **Ažuriranje nizova u stanju**: Kada ažurirate nizove u stanju, važno je da ne modifikujete niz direktno korišćenjem metoda kao što su `push()` ili `pop()`. Umesto toga, trebalo bi da kreirate novi niz prilikom ažuriranja stanja:

```jsx
const addItem = () => {
  const newItem = {
    id: items.length + 1,
    name: `Item ${items.length + 1}`,
  };

  // Creates a new array
  setItems((prevItems) => [...prevItems, newItem]);
};
```

Ako želite da uklonite elemente iz niza, trebalo bi da koristite metod ``filter()``, koji vraća novi niz nakon filtriranja svega što želite da uklonite:

```jsx
const removeItem = (id) => {
  setItems((prevItems) => prevItems.filter((item) => item.id !== id));
};
```

## Pristup Vrednostima Korišćenjem Refova

- **`ref` Atribut**: Možete pristupiti DOM čvoru u React-u korišćenjem atributa `ref`. Evo primera kako se može demonstrirati `ref` za fokusiranje elementa `input`. Svojstvo `current` se koristi za pristup trenutnoj vrednosti tog `ref`:

```jsx
import { useRef } from "react";

const Focus = () => {
  const inputRef = useRef(null);

  const handleFocus = () => {
    if (inputRef.current) {
      inputRef.current.focus();
    }
  };

  return (
    <div>
      <input ref={inputRef} type="text" placeholder="Enter text" />
      <button onClick={handleFocus}>Focus Input</button>
    </div>
  );
};

export default Focus;
```

## Rad sa `useEffect` Hook-om

- **`useEffect()` Hook**: U Reactu, efekat je sve što se dešava izvan procesa renderovanja komponente. To znači bilo šta što React ne obrađuje direktno kao deo renderovanja korisničkog interfejsa (UI). Česti primeri uključuju dohvaćanje podataka (fetching data), ažuriranje naslova kartice pretraživača, čitanje ili pisanje u lokalni skladište pretraživača (local storage), dobijanje lokacije korisnika i mnogo više. Ove operacije interaguju sa spoljnim svetom i poznate su kao nuspojave (side effects). React pruža `useEffect` hook da vam omogući upravljanje tim nuspojavama. `useEffect` vam omogućava pokretanje funkcije nakon što komponenta završi renderovanje ili ažuriranje.

```js
import { useEffect } from "react";

useEffect(() => {
 // Your side effect logic (usually a function) goes here
}, [dependencies]);
```

Funkcija efekta se izvršava nakon što komponenta renderuje, dok opcioni argument ``dependencies`` kontroliše kada će se efekat pokrenuti/izvršiti.

Imajte na umu da `dependencies` može biti niz 'reaktivnih vrednosti' (stanje, props, funkcije, varijable i slično), prazan niz ili je potpuno izostavljen. Evo kako sve te opcije kontrolišu način na koji `useEffect` radi:

- Ako je ``dependencies`` niz koji uključuje jedan ili više reaktivnih vrednosti, efekat će se izvršiti svaki put kada se oni promene.

Ako je `dependencies` prazan niz, `useEffect` se izvršava samo jednom kada komponenta prvi put renderuje.

- Ako izostavite ``dependencies``, efekat se izvršava svaki put kada komponenta renderuje ili ažurira.

## Kako kreirati prilagođene hookove

- **Prilagođeni Hookovi**: Prilagođeni hook vam omogućavaju da izvučete ponovno upotrebljivu logiku iz komponenti, kao što su dohvaćanje podataka (data fetching), upravljanje stanjem (state management), uključivanje/isključivanje (toggling) i bočni efekti poput praćenja statusa online-a. U Reactu, svi ugrađeni hookovi počinju rečju ``use``, pa vaš prilagođeni hook bi trebalo da prati istu konvenciju.

Evo primera kreiranja `useDebounce` hook-a:

```jsx
function useDebounce(value, delay) {
  const [debouncedValue, setDebouncedValue] = useState(value);

  useEffect(() => {
    const handler = setTimeout(() => {
      setDebouncedValue(value);
    }, delay);

    return () => {
      clearTimeout(handler);
    };
  }, [value, delay]);

  return debouncedValue;
}

export { useDebounce };
```

```markdown
## Rad sa formularima u React-u
```

**Kontrolisani unosi**: Ovo se dešava kada sačuvate vrednost polja za unos u stanje (state) i ažurirate je putem događaja ``onChange``. To vam daje potpunu kontrolu nad podacima forme i omogućava trenutnu validaciju i uslovno renderovanje.

```jsx
import { useState } from "react";

function App() {
  const [name, setName] = useState("");

  const handleChange = (e) => {
    setName(e.target.value);
  };

  const handleSubmit = (e) => {
    e.preventDefault();
    console.log(name);
  };

  return (
    <>
      <form onSubmit={handleSubmit}>
        <label htmlFor="name">Your name</label> <br />
        <input value={name} id="name" onChange={handleChange} type="text" />
        <button type="submit">Submit</button>
      </form>
    </>
  );
}

export default App;
```

- **Nekontrolisani unosi (Inputs)**: Umesto da se unosi obrađuju putem hooka ``useState``, nekontrolisani unosi u HTML zadržavaju svoj sopstveni interni status uz pomoć DOM-a. Pošto DOM kontroliše vrednosti polja za unos, potrebno je da povučete (ili dohvatite) vrednosti iz polja za unos pomoću ``ref``.

```jsx
import { useRef } from "react";

function App() {
 const nameRef = useRef();

 const handleSubmit = (e) => {
   e.preventDefault();
   console.log(nameRef.current.value);
 };

 return (
   <form onSubmit={handleSubmit}>
     <label htmlFor="name">Your</label>{" "}
     <input type="text" ref={nameRef} id="name" />
     <button type="submit">Submit</button>
   </form>
 );
}

export default App;
```

## Rad sa `useActionState` Hook-om

- **Akcije servera (Server Actions)**: Ovo su funkcije koje se izvršavaju na serveru kako bi omogućile obradu formulara direktno na serveru, bez potrebe za API endpoints-ima. Evo primera iz Next.js aplikacije:

```js
"use server";

async function submitForm(formData) {
 const name = formData.get("name");
 return { message: `Hello, ${name}!` };
}
```

Direktiva ``"use server"`` označava funkciju kao serversku akciju.

- **`useActionState` Hook**: Ova kuka ažurira stanje na osnovu ishoda slanja forme. Evo osnovne sintakse za `useActionState` hook:

```js
const [state, action, isPending] = useActionState(actionFunction, initialState, permalink);
```

- `state` je trenutno stanje koje akcija vraća.
- `action` je funkcija koja pokreće serversku akciju.
- `isPending` je boolean koji ukazuje da li je akcija trenutno u toku ili ne.
- Parametar `actionFunction` jeste sama serverska akcija.
- `initialState` je parametar koji predstavlja početnu tačku stanja pre nego što se akcija izvrši.
- `permalink` je opcijalni string koji sadrži jedinstvenu URL adresu stranice koju formular modifikuje.

## Dohvatanje Podataka u React-u

- **Opcije za dohvaćanje podataka**: Postoji mnogo različitih načina za dohvaćanje podataka u Reactu. Možete koristiti naivni Fetch API, ili alate treće strane kao što su Axios ili SWR.
- **Najčešće korišćene state varijable prilikom dohvaćanja podataka**: Bez obzira na način koji odaberete za dohvaćanje podataka u Reactu, potrebno je da pratite nekoliko stanja (state). Prvo je podaci sami po sebi. Drugo će pratiti da li su podaci još u procesu dohvaćanja. Treće je state varijabla koja će uhvatiti sve greške koje se mogu dogoditi tokom procesa dohvaćanja podataka.

```js
const [data, setData] = useState(null);
const [loading, setLoading] = useState(true);
const [error, setError] = useState(null);
```

Pošto je dohvaćanje podataka nuspojavnost, najbolje je koristiti ``Fetch API`` unutar ``useEffect`` hook-a.

```js
useEffect(() => {
  const fetchData = async () => {
    try {
      const res = await fetch("https://jsonplaceholder.typicode.com/posts");
      
      if (!res.ok) {
        throw new Error("Network response was not ok");
      }

      const data = await res.json();
      setData(data);
    } catch (err) {
      setError(err);
    } finally {
      setLoading(false);
    }
  };

  fetchData();
}, []);
```

Zatim možete prikazati poruku o učitavanju ako dohvaćanje podataka nije završeno, poruku o grešci ako je došlo do greške prilikom dohvaćanja podataka, ili rezultate.

```jsx
if (loading) {
  return <p>Loading...</p>;
}

if (error) {
  return <p>{error.message}</p>;
}

return (
  <ul>
    {data.map((post) => (
      <li key={post.id}>{post.title}</li>
    ))}
  </ul>
);
```

Ako želite da koristite Axios, morate ga instalirati i uvesti:

```bash
npm i axios
```

```js
import axios from "axios";
```

Zatim možete dohvatiti podatke korišćenjem `axios.get`:

```js
const [data, setData] = useState(null);
const [loading, setLoading] = useState(true);
const [error, setError] = useState(null);

useEffect(() => {
  const fetchData = async () => {
    try {
      const res = await axios.get(
        "https://jsonplaceholder.typicode.com/users"
      );
      setData(res.data);
    } catch (err) {
      setError(err);
    } finally {
      setLoading(false);
    }
  };

  fetchData();
}, []);
```

Da biste dohvatili podatke koristeći `useSWR` hook, morate ga prvo instalirati i uvesti.

```bash
npm i swr
```

```js
import useSWR from "swr";
```

Evo kako možete koristiti hook za dohvaćanje podataka:

```jsx
import useSWR from "swr";

const fetcher = (url) => fetch(url).then((res) => res.json());

const FetchTodos = () => {
 const { data, error } = useSWR(
   "https://jsonplaceholder.typicode.com/todos",
   fetcher
 );

 if (!data) {
   return <h2>Loading...</h2>;
 }
 if (error) {
   return <h2>Error: {error.message}</h2>;
 }

 return (
   <>
     <h2>Todos</h2>
     <div>
       {data.map((todo) => (
         <h3 key={todo.id}>{todo.title}</h3>
       ))}
     </div>
   </>
 );
};

export default FetchTodos;
```

## Rad sa `useOptimistic` Hook-om

**`useOptimistic` Hook**: Ovaj hook se koristi za održavanje UI-a (korisničkog interfejsa) reaktivnim dok čeka da neka asinhrona akcija završi u pozadini. Pomaže u upravljanju "optimističkim ažuriranjima" u UI-u, strategijom po kojoj se trenutno ažuriranja pružaju UI-u na osnovu očekivanog ishoda neke akcije, kao što je čekanje odgovora sa servera.

Ovo je osnovna sintaksa:

```js
const [optimisticState, addOptimistic] = useOptimistic(actualState, updateFunction);
```

- `optimisticState` je privremeno stanje koje se ažurira odmah radi boljeg korisničkog iskustva.
- `addOptimistic` je funkcija koja primenjuje optimističko ažuriranje pre nego što dođe do stvarnih promena stanja.
- `actualState` je stvarna vrijednost stanja koja dolazi iz rezultata akcije, kao što je preuzimanje podataka sa servera.
- `updateFunction` je funkcija koja određuje kako bi optimistično stanje trebalo da se ažurira kada bude pozvano.

Evo primera korišćenja `useOptimistic` hook-a u komponenti `TaskList`:

```js
"use client";

import { useOptimistic } from "react";

export default function TaskList({ tasks, addTask }) {
  const [optimisticTasks, addOptimisticTask] = useOptimistic(
    tasks,
    (state, newTask) => [...state, { text: newTask, pending: true }]
  );

  async function handleSubmit(e) {
    e.preventDefault();
    const formData = new FormData(e.target);

    addOptimisticTask(formData.get("task"));

    addTask(formData);
    e.target.reset();
  }

  return <>{/* UI */}</>;
}
```

- **`startTransition`**: Ovo se koristi za renderovanje dela UI-ja i označavanje ažuriranja stanja kao nehitnog prelaska (transition). To omogućava da je UI responsivan tokom skupih ažuriranja. Evo osnovne sintakse:

```js
startTransition(action)
```

Komponenta ``action`` izvršava ažuriranje stanja ili pokreće neku logiku vezanu za tranziciju. Ovo osigurava da hitna ažuriranja korisničkog interfejsa (UI), poput kucanja ili klika, nisu blokirana.

## Rad sa `useMemo` Hook-om

- **Memoizacija**: Ovo je tehnika optimizacije u kojoj se rezultati skupih funkcionih poziva keširaju (pamte) na osnovu specifičnih argumenata. Kada se ponovo pruže isti argumenti, vraća se keširani rezultat umesto da se funkcija ponovno izračunava.
- **`useMemo` Hook**: Ovaj hook se koristi za memoizovanje izračunatih vrednosti. Evo primera kako se memoizuje rezultat sortiranja velikog niza (array). `expensiveSortFunction` će se izvršiti samo kada se promeni `largeArray`:

```js
const memoizedSortedArray = useMemo(
  () => expensiveSortFunction(largeArray),
  [largeArray]
);
```

## Rad sa `useCallback` Hook-om

- **`useCallback` Hook**: Koristi se za memoizovanje referenci funkcija.

```js
const handleClick = useCallback(() => {
 // code goes here
}, [dependency]);
``` 

- **`React.memo`**: Ovo se koristi za memorizovanje komponente kako bi se sprečilo nepotrebno ponovno renderovanje kada se njene propove nisu promenile.

```js
const MemoizedComponent = React.memo(({ prop }) => {
 return (
   <>
     {/* Presentation */}
   </>
 )
});
``` 

## Alati za upravljanje zavisnostima

- **Definicija zavisnosti**: U softveru, zavisnost je situacija u kojoj jedna komponenta ili modul u aplikaciji zavisi od druge kako bi pravilno funkcionisala. Zavisnosti su česte u softverskim aplikacijama jer omogućavaju programerima da koriste predizgrađene funkcije ili alate koje su kreirali drugi. Dve osnovne zavisnosti potrebne za React projekat biće paketi `react` i `react-dom`:

```json
"dependencies": {
  "react": "^18.3.1",
  "react-dom": "^18.3.1"
}
```

- **Definicija menadžera paketa**: Da biste upravljali zavisnostima softvera u projektu, moraćete da koristite menadžer paketa. Menadžer paketa je alat koji se koristi za instalaciju, ažuriranje i uklanjanje zavisnosti. Mnogi popularni programski jezici kao što su JavaScript, Python, Ruby i Java, svi koriste menadžere paketa. Popularni menadžeri paketa za JavaScript uključuju npm, Yarn i pnpm.
- **`package.json` Fajl**: Ovo je ključni konfiguracioni fajl u projektima koji sadrži metapodatke o vašem projektu, uključujući njegovo ime, verziju i zavisnosti. Takođe definiše skripte, informacije o licenciranju i druge postavke koje pomažu u upravljanju projektom i njegovim zavisnostima.
- **`package-lock.json` Fajl**: Ovaj fajl će zaključati tačne verzije svih paketa koje vaš projekat koristi. Kada ažurirate paket, nove verzije će takođe biti ažurirane u fajlu za zaključavanje (lock file).
- **`node_modules` Folder**: Ovaj folder sadrži stvarni kod zavisnosti navedenih u vašem `package.json` fajlu, uključujući i direktne zavisnosti vašeg projekta i sve zavisnosti tih zavisnosti.
- **Dev Dependencies (Zavisnosti za razvoj)**: Ovo su paketi koji se koriste samo za razvoj, a ne u produkciji. Primer ovoga bi bila biblioteka za testiranje kao što je Jest. Instalirali biste Jest kao dev zavisnost jer je potreban za testiranje vaše aplikacije lokalno, ali nije potreban da bi aplikacija radila u produkciji.

```json
"devDependencies": {
  "@eslint/js": "^9.17.0",
  "@types/react": "^18.3.18",
  "@types/react-dom": "^18.3.5",
  "@vitejs/plugin-react": "^4.3.4",
  "eslint": "^9.17.0",
  "eslint-plugin-react": "^7.37.2",
  "eslint-plugin-react-hooks": "^5.0.0",
  "eslint-plugin-react-refresh": "^0.4.16",
  "globals": "^15.14.0",
  "vite": "^6.0.5"
}
```

## React Router

- **Uvod**: React Router je biblioteka treće strane koja vam omogućava dodavanje rutiranja u vaše React aplikacije. Da biste počeli, moraćete da instalirate React Router u postojeći React projekat na sledeći način:

```bash
npm i react-router
```

Nakon toga, unutar fajla ``main.jsx`` ili ``index.jsx``, moraćete da postavite strukturu rute na ovaj način:

```jsx
import { StrictMode } from "react";
import { createRoot } from "react-dom/client";
import { BrowserRouter, Routes, Route } from "react-router";
import App from "./App.jsx";

import "./index.css";

createRoot(document.getElementById("root")).render(
  <StrictMode>
    <BrowserRouter>
      <Routes>
        <Route path="/" element={<App />} />
      </Routes>
    </BrowserRouter>
  </StrictMode>
);
```

``path`` i ``element`` se koriste za povezivanje URL i UI komponenti. U ovom slučaju, postavljamo rutu za naslovnu stranicu koja ukazuje na ``App`` komponentu.

- **Višestruki prikazi i podešavanje ruta**: U većim aplikacijama je uobičajeno imati više prikaza i ruta podešenih na ovaj način:

```jsx
<Routes>
  <Route index element={<Home />} />
  <Route path="about" element={<About />} />

  <Route path="products">
    <Route index element={<ProductsHome />} />
    <Route path=":category" element={<Category />} />
    <Route path=":category/:productId" element={<ProductDetail />} />
    <Route path="trending" element={<Trending />} />
  </Route>
</Routes>
```

Atribut ``index`` u ovim primerima služi da predstavlja podrazumevanu rutu za dati segment puta. Stoga, komponenta ``Home`` će biti prikazana na putanji na korenu ``/``, dok će komponenta ``ProductsHome`` biti prikazana na putanji ``/products``.

- **Ugnježđivanje Ruta**: Možete ugnježđivati rute unutar drugih ruta, što rezultira time da se putanja detalne rute spaja sa putanjom roditeljske rute.

```jsx
<Route path="products">
  <Route path="trending" element={<Trending />} />
</Route>
```

U navedenom primeru, putanja za popularne proizvode će biti products/trending.

- **Dinamički segmenti**: Dinamički segment je deo putanje URL-a koji može biti dinamičan/promenljiv.

```jsx
<Route path=":category" element={<Category />} />
```

U ovom primeru imamo dinamički segment nazvan `category`. Kada korisnik navigira na URL poput products/brass-instruments, prikaz će se promeniti u komponentu `Category` i možete dinamički dohvatiti odgovarajuće podatke na osnovu tog segmenta.

- **`useParams` Hook**: Ovaj hook se koristi za pristup dinamičkim parametrima iz putanje URL-a.

```js
import { useParams } from "react-router";

export default function Category() {
  let params = useParams();
   {/* Accessing the category param: params.category */}
   {/* rest of code goes here */}
}
```

## React Okvirnici

- **Uvod**: React framework-ovi pružaju funkcije kao što su rutiranje (routing), optimizacija slika, dohvaćanje podataka (data fetching), autentifikacija i još mnogo toga. To znači da možda nećete morati da postavljate odvojene front-end i back-end aplikacije za određene slučajeve upotrebe (use cases). Primeri React Frameworks uključuju Next.js i Remix.
- **Next.js Routing**: Ovaj sistem rutiranja (routing) uključuje podršku za dinamičke rute (dynamic routes), paralelne rute (parallel routes), rukovače ruta (route handlers), preusmeravanja (redirects), internacionalizaciju i još mnogo toga.

Evo primera kreiranja prilagođenog rukovaoca za zahtev:

```js
export async function GET() {
  const res = await fetch("https://example-api.com");
  const data = await res.json();

  return Response.json({ data });
}
```

**Optimizacija slika sa Next.js**: Komponenta `Image` proširuje nativni HTML element `img` i omogućava brže učitavanje stranica i optimizaciju veličine. To znači da će slike biti učitane samo kada uđu u vidokrug (viewport), a komponenta `Image` će automatski servirati ispravno dimenzionirane slike za svaki uređaj.

```jsx
import Image from "next/image";

export default function Page() {
  return (
    <Image src="link-to-image-goes-here" alt="descriptive-title-goes-here" />
  );
}
``` 

## Prošifrovanje propova

- **Definicija**: Prop drilling je proces prosleđivanja `props`-ova iz roditeljske komponente do duboko ugnježanih detet komponenti, čak i kada neke od tih detet komponenti ne zahtevaju te `props`-ove.

## Upravljanje stanjem

- **Context API**: Kontekst se odnosi na situaciju kada roditeljska komponenta čini informacije dostupnim detilnim komponentama bez potrebe da ih eksplicitno prosleđuje kroz props. `createContext` se koristi za kreiranje objekta konteksta koji predstavlja kontekst koji će druge komponente čitati. `Provider` se koristi za dostavljanje vrednosti konteksta detilnim komponentama.

```js
import { useState, createContext } from "react";

const CounterContext = createContext();

const CounterProvider = ({ children }) => {
  const [count, setCount] = useState(0);

  return (
    <CounterContext.Provider value={{ count, setCount }}>
      {children}
    </CounterContext.Provider>
  );
};

export { CounterContext, CounterProvider };
```

- **Redux**: Redux upravlja upravljanjem stanjem (state management) pružajući centralni *store* i strogu kontrolu nad ažuriranjem stanja. Koristi predvidljiv obrazac sa akcijama (*actions*), reducerima (*reducers*) i *middleware*-om. Akcije su nosioci informacija koji šalju podatke iz vaše aplikacije u Redux *store*, često pokrenuti interakcijama korisnika. Reduceri su funkcije koje određuju kako bi se stanje trebalo promeniti kao odgovor na te akcije, osiguravajući da je stanje ažurirano na nepromenljiv način (*immutable*). Middleware, s druge strane, deluje kao most između dispečovanja akcija i reducer-a, omogućavajući vam da proširite funkcionalnost Redux-a (npr. logovanje, rukovanje asinhronim operacijama) bez modifikovanja osnovnog toka.

- **Zustand**: Ovo rešenje za upravljanje stanjem je idealno za aplikacije srednje i male veličine. Radi korišćenjem ``useStore`` hooka kako bi se stanju pristupilo direktno u komponentama i stranicama. To vam omogućava da modifikujete i pristupate podacima bez potrebe za akcijama, reducerima ili provider-om.

## Debagovanje React Komponenata Korišćenjem React DevTools

- **React Developer Tools**: Ovo je ekstenzija za pretraživače koju možete koristiti u Chrome-u, Firefox-u i Edge-u za inspekciju React komponenti i identifikovanje problema sa performansama. Za Safari, moraćete da instalirate npm paket `react-devtools`. Nakon instaliranja React DevTools i otvaranja React aplikacije u pretraživaču, otvorite alate za razvoj (developer tools) pregledača kako biste pristupili dve dodatne kartice namenjene debugovanju React-a – Components i Profiler.
- **Components Tab**: Ova kartica prikazuje svaku komponentu u formatu hijerarhije (tree view). Ovde možete uraditi sledeće:
  - pregledati hijerarhiju komponenti aplikacije
  - proveriti i modifikovati props, state i context vrednosti u realnom vremenu
  - proveriti izvorni kod za svaku odabranu komponentu
  - logovati podatke komponente u konzolu
  - inspekcionisati DOM elemente za komponentu
- **Profiler Tab**: Ova kartica vam pomaže da analizirate performanse komponenti. Možete snimiti performanse komponente kako biste identifikovali nepotrebne re-render-e, pregledali trajanje commit-a i naknadno optimizovali sporije komponente.

## React Komponente na serveru

- **Definicija**: React Server Components su React komponente koje se renderuju isključivo na serveru, slajući samo konačni HTML klijentu. To znači da te komponente mogu direktno pristupiti resursima sa strane servera i dramatično smanjiti količinu JavaScript-a poslaju na pretraživač.

## Razlike između stvarnih i percipiranih performansi

- **Percipirana performansa**: Ovo je kako korisnici percipiraju performanse veb sajta. To je način na koji ga procenjuju u smislu reaktivnosti i pouzdanosti. Ovo je subjektivna mera, pa je teško kvantifikovati, ali je veoma važna jer iskustvo korisnika određuje uspeh ili neuspeh veb sajta.
- **Stvarna performansa**: Ovo su objektivne i merljive performanse veb sajta. Meri se korišćenjem metrika kao što su vreme učitavanja stranice, vreme odgovora servera i vreme renderovanja. Ove mere pogođene su više faktora vezanih za mrežu i sam kod.

## Tehnike za poboljšanje percipirane performanse

- **Opuštanje učitavanja (Lazy Loading)**: Ova tehnika smanjuje početno vreme učitavanja koliko je moguće tako što neesencijalne resurse učitava u pozadini.
- **Minimiziranje kašnjenja fontova**: Ako vaš sajt koristi prilagođene fontove, trebalo bi da pokušate i da minimizirate kašnjenje učitavanja fonta, jer to može rezultirati treperenjem ili prikazivanjem rezervnog fonta dok se prilagođeni font učitava. Predlog za ovo je korišćenje rezervnog fonta koji je sličan prilagođenom fontu, tako da u slučaju toga, promena će biti manje primetna.
- **Korišćenje indikatora učitavanja**: Prikazivanje indikatora učitavanja za dugotrajan proces čim korisnik klikne na element može pomoći korisniku da se oseća povezano i angažovano sa procesom, što čini vreme čekanja kraćim.

## Osnovni koncepti performansi

- **Redosled izvora**: Ovo se odnosi na način na koji su HTML elementi strukturirani u dokumentu. To određuje šta se učitava prvo i može značajno uticati na performanse i pristupačnost.

Neke najbolje prakse za red izvora uključuju:

- Postavljanje kritičnog sadržaja kao što su naslovi, navigacija ili glavni tekst više u HTML strukturi.
- Odlaganje neesencijalnih skripti, kao što su one za analitiku ili widgete trećih strana, kako ne bi blokirale renderovanje.
- Korišćenje progresivnog poboljšanja (progressive enhancement), kako bi se osiguralo da osnovno iskustvo radi čak i pre nego što se učitaju stilovi i skripte. Progresivno poboljšanje je način izgradnje veb sajtova i aplikacija zasnovan na ideji da vaša stranica mora prvo da funkcioniše samo sa HTML-om.

Ovo je primer dobrog redosleda izvora, koristeći najbolje prakse koje smo upravo prošli:

```html
<h1>Welcome to FastSite!</h1>
<p>Critical information loads first.</p>
<script src="slow-script.js" defer></script>
```

- **Critical Rendering Path**: Ovo je sekvenca koraka koje pregledač prati kako bi pretvorio kod u piksele na ekranu.
- **Latency**: Ovo je vreme potrebno da zahtev putuje između pregledača i servera. Dakle, drugim rečima, visoko kašnjenje znači sporne stranice.

Neki načini za smanjenje latencije uključuju:

- Korišćenje CDN-ova, ili Content Delivery Networks, za serviranje fajlova sa bližih lokacija.
- Omogućavanje kompresije pomoću tehnologija kao što je Gzip radi smanjenja veličine fajlova.
- Optimizovanje slika i korišćenje lazy loading-a (lenog učitavanja).

```html
<img src="placeholder.jpg" data-src="real-image.jpg" loading="lazy">
```

## Poboljšanje INP-a

- **Definicija**: INP (Interaction to Next Paint) procenjuje ukupnu reaktivnost stranice merenjem vremena od trenutka kada korisnik interaguje, kao što je klik ili pritisak tastera, do sledećeg ažuriranja prikaza u pregledaču. Niži INP ukazuje na reaktivniju stranicu.

Evo nekoliko načina za poboljšanje INP:

- Smanjite opterećenje glavne niti razbijanjem dugih JavaScript zadataka.
- Koristite `requestIdleCallback()` za nekritične skripte. Ovo će staviti funkciju u red da bude pozvana tokom perioda neaktivnosti pretraživača.
- Odložite ili opušteno učitajte teške resurse o kojima je ranije rečeno.
- Optimizujte rukovače događaja (event handlers). Ako ovi rukovači rade prečesto ili obavljaju teške operacije, mogu usporiti stranicu i povećati INP. Rešenje za ovo je debouncing. Debouncing osigurava da funkcija radi samo nakon što korisnik prestane kucanje na kratak interval – recimo 300ms. Ovo sprečava nepotrebne proračune i poboljšava performanse.

## Kako funkcioniše renderovanje u pretraživaču

# --description--
Proces renderovanja u pretraživaču je složen i uključuje nekoliko koraka: parsiranje HTML-a, izgradnju DOM stabla, kreiranje CSSOM-a, a na kraju njihovo kombinovanje radi prikazivanja stranice. Ovaj vodič će vas kroz ove koncepte voditi.

# --hints--
Imajte na umu da JavaScript može manipulisati DOM-om nakon početne faze renderovanja. Razumevanje ovog razdvajanja je ključno za napredno web programiranje.
# --assignment--
Napišite jednostavan skript koji demonstrira kako promena imena klase elementa utiče na njegovo stilizovanje, bez korišćenja eksternih CSS fajlova. Koristite samo inline stilove ili interne ``<style>`` oznake.

```html
<div id="app">
  <!-- Sadržaj ide ovde -->
</div >
```

- **Kako funkcioniše renderovanje**: Prvo, pretraživač parsira ``HTML`` i gradi ``DOM``. Zatim, pretraživač obrađuje ``CSS``, konstruišući ``CSS Object Model`` ili ``CSSOM``. Ovo je još jedna stabljasta struktura koja određuje kako bi elementi trebalo da budu stilizovani. Konačno, pretraživač crta piksele na ekranu, renderujući svaki element na osnovu izračunatih stilova i rasporeda (layout). Na kompleksnim stranicama, ovo može uključivati više slojeva koji se kompoziraju zajedno kako bi se stvorio konačni vizuelni izlaz.

## Kako performanse utiču na održivost

- **Pozadinske informacije**: Internet čini oko 2% globalnih emisija ugljen-dioksida—što je isto koliko i avioindustrija! Svaki prebačen bajt zahteva električnu energiju, od centara podataka do korisničkih uređaja. Veći fajlovi i neefikasni skriptovi znače veće potrošnje energije. Veb sajt visokih performansi nije samo brži, već takođe smanjuje nepotrebno procesiranje i potrošnju energije.

## Načini za smanjenje vremena učitavanja stranice

- **Optimizujte medijske resurse**: Velike slike i video snimci su uobičajeni uzroci sporog učitavanja. Optimizovanjem ovih resursa, možete značajno ubrzati svoj sajt. Ovo uključuje stvari kao što su kompresovanje slika, korišćenje modernih formata poput WebP i upotreba opuštenog učitavanja (lazy loading) za resurse.
- **Iskoristite keširanje u pregledaču**: Keširanje omogućava pregledačima da lokalno skladište deo vašeg sajta, smanjujući vreme učitavanja za povratnike.
- **Minifikujte i komprimujte fajlove**: Smanjenje veličine vaših fajlova može dovesti do bržeg preuzimanja. Ovo uključuje smanjenje veličine prenetih fajlova i minifikovanje CSS i JavaScript fajlova.

## Poboljšanje „vremena do upotrebljivosti“

- **Definicija**: "time to usable" je interval od trenutka kada korisnik zahteva stranicu do trenutka kada može sa njom značajno da interaguje. Da biste poboljšali "time to usable", možete *lazy load*-ovati svoj resurs ili minimizirati resurse koji blokiraju renderovanje.

## Ključne metrike za merenje performansi

- **First Contentful Paint ili FCP**: Meri koliko brzo prvi deo sadržaja—tekst ili slika—pojavljuje se na ekranu. Dobar FCP se smatra vremenom ispod 1,8 sekundi, a loš FCP je iznad 3 sekunde. Možete proveriti svoj FCP koristeći Chrome DevTools i performansnu karticu (performance tab).
- **Total Blocking Time**: Ovo pokazuje koliko dugo je glavni *thread* blokiran teškim JavaScript zadacima. Ako je Total Blocking Time (TBT) visok, korisnici doživljavaju sporne interakcije. Da biste poboljšali TBT, razbijte dugačke zadatke i odložite neesencijalne skripte.
- **Bounce Rate**: Ovo je procenat posetilaca koji napuštaju sajt bez bilo kakve interakcije. Ako vaš sajt ima visok *bounce rate*, to može biti zato što je stranica previše spora.
- **Unique Users**: Ova metrika prati koliko pojedinačnih posetilaca dolazi na vaš sajt. Da biste videli *Bounce Rate* i *Unique Users*, možete koristiti Google Analytics. To će vam omogućiti da pratite ove metrike i poboljšate angažman.

## Uobičajeni alati za merenje performansi

- **Chrome DevTools**: Chrome DevTools je ugrađeni alat unutar Google Chromea koji vam omogućava da analizirate i debagujete performanse u realnom vremenu. DevTools će prikazati vremena učitavanja, korišćenje CPU-a i kašnjenja renderovanja. Posebno je koristan za merenje First Contentful Paint (FCP), što ukazuje koliko brzo korisnik vidi prvi vidljivi sadržaj. Ako vam se sajt čini sporim, DevTools će vam pomoći da locirate usko grlo performansi.
- **Lighthouse**: Ovo je automatizovani alat koji proverava performanse, SEO i pristupačnost.
- **WebPageTest**: Ovaj alat vam omogućava da testirate kako vaš sajt učitava iz različitih lokacija i uređaja. Alat daje detaljnu analizu Speed Index-a, Total Blocking Time-a i drugih ključnih metrika performansi vašeg sajta. Ako želite da znate kako stvarni korisnici doživljavaju vaš sajt globalno, WebPageTest je alat za to.
- **PageSpeed Insights**: Ovaj alat analizira vaš veb-sajt i predlaže brojne poboljšanja za mobilne i desktop uređaje. Pokazuje šta usporava vaš sajt i daje specifične preporuke kao što su optimizovanje slika, uklanjanje skripti koje blokiraju renderovanje i smanjenje vremena odgovora servera. PageSpeed Insights je brz i jednostavan način da proverite kako Google vidi performanse vašeg sajta.
- **Real User Monitoring**: Alati za RUM prače stvarni ponašanje korisnika, pokazujući kako pravi posetioci doživljavaju vaš sajt. Popularni RUM alati uključuju Google Analytics, koji prati vremena učitavanja stranice i stopu otpadanja (bounce rates), kao i New Relic ili Datadog, koji nadziru probleme performansi u realnom vremenu. Ako želite podatke od stvarnih korisnika, RUM alati su neophodni.

```markdown
## Rad sa Web API-jima za performanse
```

- **Definicija**: Performance Web API-ji omogućavaju programerima da prate koliko efikasno se web stranica učitava i reaguje, direktno u kodu. Ovi API-ji vam omogućavaju da merite vreme učitavanja stranice, pratite kašnjenja renderovanja i interakcije, te analizirate vreme izvršavanja JavaScript-a.
- **`performance.now()`**: Ovaj API vam pruža vremenske oznake visoke preciznosti (u milisekundama) za merenje koliko vremena je potrebno da se različiti delovi vaše stranice učitaju.

```js
const start = performance.now();  
// Run some code here  
const end = performance.now();  

console.log(`Execution time: ${end - start}ms`);
```

- **Performance Timing API**: Ovaj API pruža detaljnu analizu svake faze učitavanja stranice, od DNS pretrage do `DOMContentLoaded`.

```js
const timing = performance.timing;  

const pageLoadTime = timing.loadEventEnd - timing.navigationStart;  
console.log(`Page load time: ${pageLoadTime}ms`);
```

- **`PerformanceObserver`**: Ovaj API prati događaje performansi, kao što su pomeranja rasporeda (layout shifts), dugotrajni zadaci i interakcije korisnika.

```js
const observer = new PerformanceObserver((list) => {  
  list.getEntries().forEach((entry) => {  
    console.log(`Long task detected: ${entry.duration}ms`);  
  });  
});  

observer.observe({ type: "longtask", buffered: true });
```

## Tehnike za unaprjeđenje performansi CSS-a

- **CSS Animacije**: Animiranje određenih CSS svojstava, kao što su dimenzije, pozicija i raspored (*layout*), pokreće proces nazvan "reflow", tokom kojeg pregledač ponovo izračunava poziciju i geometriju određenih elemenata na stranici. Ovo zahteva *repaint*, što je računski skupo (ili: predstavlja visoko opterećenje za CPU). Stoga se preporučuje smanjiti broj CSS animacija koliko je moguće ili barem dati korisniku opciju da ih uključi/isključi (*toggle*).

## Tehnike za poboljšanje performansi JavaScript-a

- **Code Splitting**: Razdvajanje vašeg JavaScript koda u module koji izvršavaju kritične i nekritične zadatke je takođe korisno. Na ovaj način, moći ćete da preuzmete (preload) kritične module što je moguće ranije i odložiti nekritične kako bi se stranica renderovala što je moguće brže.
- **DOM Manipulation**: Imajte na umu da se DOM Manipulation odnosi na proces dinamičke promene sadržaja stranice pomoću JavaScript-a, interakcijom sa Document Object Model (DOM). Manipulacija DOM-om predstavlja računarski trošak/opterećenje. Smanjenjem količine DOM manipulacije u vašem JavaScript kodu poboljšaćete performanse.

## CSS Okvirnici

- **CSS frameworks**: CSS frameworkovi mogu ubrzati vaš radni proces, stvoriti uniforman vizuelni stil na celoj veb stranici, osigurati da je dizajn konzistentan u više pregledača i održati vaš CSS kod organizovanijim.
- **Popular CSS frameworks**: Neki od popularnih CSS frameworkova su Tailwind CSS, Bootstrap, Materialize i Foundation.
- **Potencijalne nedostatke**:
  - CSS koji pruža framework može kolidirati sa vašim prilagođenim (custom) CSS-om.
  - Vaša veb stranica može izgledati slično drugim veb stranicama koje koriste isti framework.
  - Veliki frameworkovi mogu uzrokovati probleme sa performansama.

```markdown
## Dve vrste CSS framework-ova
```

- **Utilitetni CSS frameworkovi**: Ovi frameworkovi sadrže male klase sa specifičnim namenama, kao što je postavljanje margine, `padding`-a ili boje pozadine. Ove male klase se mogu direktno primeniti na `HTML` elemente po potrebi. Tailwind CSS je klasifikovan kao utilitarni CSS framework.

Ovo je primer korišćenja Tailwind CSS za stilizovanje dugmeta.

```html
<button class="bg-blue-500 text-white font-bold py-2 px-4 rounded-full hover:bg-blue-700">
  Button
</button>
```

- **CSS okviri bazirani na komponentama**: Ovi okviri imaju gotovne komponente sa predodređenim stilovima koje možete dodati na svoj veb sajt. Komponente su dostupne u zvaničnoj dokumentaciji CSS okvira, a vi ih možete kopirati i zalepiti u svoj projekat. Bootstrap je kategorizovan kao CSS okvir baziran na komponentama.

Ovo je primer korišćenja Bootstrapa za kreiranje liste grupa. Umesto primene malih klasa na vaše HTML elemente, dodaćete celu komponentu, uključujući i njenu HTML strukturu.

```html
<div class="card" style="width: 25rem;">
  <ul class="list-group list-group-flush">
    <li class="list-group-item">HTML</li>
    <li class="list-group-item">CSS</li>
    <li class="list-group-item">JavaScript</li>
  </ul>
</div>
```

## Tailwind CSS

Tailwind je CSS framework koji koristi pristup zasnovan na utilitetima. Umesto da pišete prilagođene CSS pravila, kreirate dizajn kombinovanjem malih utilitarnih klasa direktno u vašem HTML-u.

### Utiliteti za responsivan dizajn

Tailwind koristi prefikse kao što su `sm:`, `md:`, i `lg:` za primenu stilova na različitim veličinama ekrana.

```html
<div class="w-full md:w-1/2 lg:flex-row">
  Responsive layout
</div>
```

### Utiliteti za Flexbox

Klasama poput `flex`, `flex-col`, `justify-around` i `items-center` olakšavaju kreiranje fleksibilnih rasporeda.

```html
<div class="flex flex-col md:flex-row justify-around items-center">
  <p>Column on small screens</p>
  <p>Row on medium and larger screens</p>
</div>
```

### Funkcije za Mreže

Tailwind uključuje utilitete za CSS Grid, kao što su `grid`, `grid-cols-1`, i `md:grid-cols-3`.

```html
<div class="grid grid-cols-1 md:grid-cols-3 gap-8">
  <div class="bg-gray-100 p-4">Column 1</div>
  <div class="bg-gray-100 p-4">Column 2</div>
  <div class="bg-gray-100 p-4">Column 3</div>
</div>
```

### Alati za Razmake

Alati poput `mt-8`, `mx-auto`, `p-4`, i `gap-4` pomažu u kreiranju konzistentnog razmaka bez pisanja CSS-a.

```html
<div class="mt-8 p-4 bg-indigo-600 text-white">
  Spaced content
</div>
```

### Tipografski Utiliteti

Alati poput ``uppercase``, ``font-bold``, ``font-semibold``, i ``text-4xl`` kontrolišu izgled teksta.

Možete postaviti veličine fonta koje se prilagođavaju na breakpointovima, kao što su `text-3xl` i `md:text-5xl`.

```html
<h1 class="text-3xl md:text-5xl font-semibold text-center">
  Responsive Heading
</h1>
```

### Boje i stanja prelaska mišem

Tailwind pruža široku paletu boja, kao što su `text-red-700`, `bg-indigo-600`, i `bg-gray-100`.

Klase poput `hover:bg-pink-600` čine interaktivne efekte jednostavnim.

```html
<a href="#" class="bg-pink-500 hover:bg-pink-600 text-white px-4 py-2 rounded-md">
  Hover Me
</a>
```

### Granice, Prstenovi i Efekti

- **Granice**: `border-2 border-red-300` dodaje granice sa definisanom debljinom i bojama.
- **Prstenje**: `ring-1 ring-gray-300` kreira efekte slične konturama, često korišćene za vizuelni fokus ili u okvirima karata/kartica.
- **Zaobnjeni uglovi i skaliranje**: Klase poput `rounded-md`, `rounded-xl`, i `scale-105` dodaju finalni vizuelni dodir.

```html
<div class="p-6 rounded-xl ring-2 ring-fuchsia-500 scale-105">
  Highlighted card
</div>
```

### Gradijenti

Tailwind podržava utilitete za gradijente kao što su `bg-gradient-to-r from-fuchsia-500 to-indigo-600`.

```html
<div class="p-4 text-white bg-gradient-to-r from-fuchsia-500 to-indigo-600">
  Gradient background
</div>
```

## CSS Preprocesori

- **CSS preprocessor**: CSS predprocesor je alat koji proširuje standardni CSS. Kompajlira kod sa proširenom sintaksom u nativni CSS fajl. Može biti koristan za pisanje čistijeg, ponovno upotrebljivijeg, manje repetitivnog i skalabilnog CSS-a za složene projekte.
- **Features**: Neke od funkcionalnosti koje mogu pružiti CSS predprocesori su varijable, mixini, ugnježđivanje (nesting) i nasleđivanje selektora (selector inheritance).
- **Popular CSS preprocessors**: Neki popularni CSS predprocesori su Sass, Less i Stylus.  
- **Potential disadvantages**:
  - Kompajliranje CSS pravila u standardni CSS može uzrokovati dodatno opterećenje (overhead).
  - Kompajlirani kod može biti težak za debugovanje.

## Sass

- **Sass**: To je jedan od najpopularnijih CSS predprocesora. Sass skraćeno je iz fraze "Syntactically Awesome Style Sheets."
- **Funkcionalnosti koje podržava Sass**: Sass podržava funkcionalnosti kao što su varijable, ugnježdena CSS pravila, moduli, mixini, nasleđivanje i operateri za osnovne matematičke operacije

## Dve sintakse koje Sass podržava

- **SCSS sintaksa**: SCSS (Sassy CSS) proširuje osnovnu sintaksu CSS-a. To je najčešće korišćena sintaksa za Sass. SCSS fajlovi imaju ekstenziju `.scss`.

Ovo je primer definisanja i korišćenja varijable u SCSS-u.

```scss
$primary-color: #3498eb;

header {
  background-color: $primary-color;
}
```

**Indentirana sintaksa**: Indentirana sintaksa je bila originalna sintaksa Sass-a i poznata je takođe kao „Sass sintaksa“.

Ovo je primer definisanja i korišćenja varijable u uvlačenoj sintaksi.

```sass
$primary-color: #3498eb

header
  background-color: $primary-color
```

```markdown
### Mikseri
```

- **Mixins**: Mixinsi vam omogućavaju da grupišete više CSS svojstava i njihove vrednosti pod jedinstvenim imenom, čime možete ponovo koristiti taj blok CSS koda u celoj vašoj stilskoj listi.

Ovo je primer definisanja miksina u SCSS sintaksi. U ovom slučaju, miksin se zove `center-flex`. Ima tri CSS svojstva za centriranje elemenata korišćenjem flexbox-a.

```scss
@mixin center-flex {
  display: flex;
  justify-content: center;
  align-items: center;
}
```

Ovo je primer korišćenja miksina koji ste definisali.

```scss
section {
  @include center-flex;
  height: 500px;
  background-color: #3289a8;
}
```

## Ručno i automatsko testiranje

- **Ručno testiranje**: U ručnom testiranju, tester fizički prolazi kroz svaki deo aplikacije i testira različite funkcionalnosti kako bi se uverio da sve radi ispravno. Ako se tokom procesa testiranja otkriju greške (bugovi), tester će ih prijaviti softverskom timu kako bi bile popravljene.
- **Automatizovano testiranje**: U automatizovanom testiranju, možete automatizovati svoje testove pisanjem zasebnog programa koji proverava da li vaša aplikacija radi očekivano i konzistentno.

## Jedinično testiranje

- **Testiranje jedinica** (Unit Testing): U testiranju jedinica, testirate svaku funkciju kako biste osigurali da sve radi onako kako je očekivano. Testovi jedinica takođe mogu služiti kao oblik dokumentacije za vašu aplikaciju jer su namenjeni da predstavljaju očekivano ponašanje vašeg koda.
- **Princip jedne odgovornosti** (Single Responsibility Principle): Princip jedne odgovornosti preporučuje održavanje svake funkcije malom i odgovornom samo za jednu stvar.
- **Uobičajeni JavaScript okviri za testiranje** (Common JavaScript Testing Frameworks): Neki uobičajeni okviri za testiranje uključuju Jest, Mocha i Vitest. Jest je popularan okvir za testiranje jedinica.

Ovo je primer jedinica testova korišćenjem Jest-a.

Prvo, možete kreirati funkciju koja je odgovorna za vraćanje novim formatiranim stringom:

```javascript
export function getFormattedWord(str) {
  if (!str) return '';
  return str.charAt(0).toUpperCase() + str.slice(1);
}
```

U posebnu datoteku `getFormattedWord.test.js`, možete napisati testove za proveru da li funkcija radi onako kako je očekivano. Datoteka `getFormattedWord.test.js` će izgledati ovako:

```javascript
import { getFormattedWord } from "./getFormattedWord.js";
import { test, expect } from "jest";

test('capitalizes the first letter of a word', () => {
  expect(getFormattedWord('hello')).toBe('Hello');
});
```

- **Funkcija**: Funkcija ``expect`` se koristi za testiranje vrednosti.
- **Matcher**: Matcher je funkcija koja proverava da li vrednost ponaša očekivano. U gornjem primeru, matcher je ``toBe()``. Jest ima raznolike matchere.

Da biste koristili Jest, prvo morate instalirati paket `jest` koristeći `npm i jest`. Takođe ćete morati dodati skriptu u vaš fajl `package.json` na sledeći način:

```js
"scripts": {
  "test": "jest"
},
```

Zatim, možete pokrenuti komandu ``npm run test`` da pokrenete svoje testove.

## Životni ciklus razvoja softvera

- **Različite faze životnog ciklusa razvoja softvera**:
  - **Faza planiranja**: Tim za razvoj prikuplja zahtjeve za predviđeni rad od zainteresovanih strana (stakeholders).
  - **Faza dizajna/projektovanja**: Tim za softver razlaže zahteve i odlučuje o najboljim pristupima rešavanja.
  - **Faza implementacije**: Tim za softver razlaže zahteve na upravljive zadatke i implementira ih.
  - **Faza testiranja**: Ovo uključuje ručno i automatizovano testiranje novog rada. Ponekad tim testira aplikaciju tokom cele faze razvoja kako bi uočio i popravio sve probleme koji se pojave.
  - **Faza postavljanja (Deployment)**: Tim postavlja nove promjene u okruženje za izgradnju ili testiranje.
  - **Faza održavanja**: Ovo uključuje popravljanje svih problema koji nastanu kod korisnika u produkcionoj aplikaciji.

- **Različiti modeli životnog ciklusa razvoja softvera**:
  - **Vodopad model (Waterfall Model)**: Vodopad model predviđa da mora biti završena svaka faza životnog ciklusa pre nego što počne sledeća faza.
  - **Agile model (Agile Model)**: Agile model se fokusira na iterativni razvoj razbijanjem rada na sprintove.

## BDD i TDD

- **TDD**: Test-driven development je metodologija koja naglašava pisanje testova kao prvog koraka. Pisanje testova pre razvoja funkcionalnosti pruža povratnu informaciju u realnom vremenu programerima tokom procesa razvoja.
- **BDD**: Behavior-driven development je pristup koji usklađuje niz testova sa poslovnim ciljevima. Test scenariji u BDD bi trebalo da budu napisani na jeziku koji mogu razumeti kako tehnički, tako i ne-tehnički pojedinci. Primer takvog sintaksnog jezika je Gherkin.
- **BDD Testing Frameworks**: Primjeri BDD testnih frameworkova uključuju Cucumber, JBehave i SpecFlow.

## Tvrdičine u testiranju jedinica

- **Tvrdišta (Assertions)**: Tvrdišta se koriste za testiranje da li kod radi onako kako je očekivano.
- **Biblioteke za tvrdišta (Assertion Libraries)**: Chai je često korišćena biblioteka za tvrdišta. Ostale uobičajene JavaScript biblioteke za tvrdišta su ``should.js`` i ``expect.js``.

Ovo je primer asertacije korišćenjem Chai-ja koja proverava da vrednost vraćena iz funkcije `addThreeAndFour` je jednaka broju 7:

```js
assert.equal(addThreeAndFour(), 7);
```

- **Najbolje prakse**: Bez obzira na biblioteku za asertacije koju koristite, trebalo bi da napišete jasne poruke za asertacije i poruke o neuspehu koje će vam pomoći da razumete koji testovi ne prolaze i zašto.

## Mokovanje, Špijanje i Stubovanje

- **Mockovanje (Mocking)**: Mockovanje je proces zamene stvarnih podataka lažnim podacima koji simuliraju ponašanje pravih komponenti. Na primer, umesto da izvršavate stalne API pozive radi preuzimanja podataka, možete mockovati odgovor API-ja tokom testiranja.
- **Stubovanje (Stubbing)**: Stubi su objekti koji vraćaju unapred definisane odgovore ili lažne podatke za očekivano ponašanje u aplikaciji. Na primer, možete stubovati ponašanje konekcije sa bazom podataka tokom testiranja bez oslanjanja na stvarnu bazu podataka.
- **Faking (Izmišljanje/Simulacija)**: Fakes su pojednostavljene verzije pravih komponenti, bez složenosti ili nuspojava. Na primer, možete napraviti "fake" bazu podataka skladištivanjem podataka u memoriji umesto komunikacije sa stvarnom bazom podataka. Ovo će vam omogućiti da simulirate operacije baze podataka direktno u memoriji, što je mnogo brže od rukovanja stvarnom bazom podataka.

## Funkcionalno testiranje

- **Funkcionalno testiranje**: Funkcionalno testiranje proverava da li karakteristike i funkcije aplikacije rade kako je očekivano. Cilj funkcionalnog testiranja je testiranje sistema u celini, pokrivajući više scenarija.
- **Nefunkcionalno testiranje**: Nefunkcionalno testiranje se fokusira na aspekte poput performansi i pouzdanosti.
- **Smoke testiranje**: Smoke testiranje je preliminarna provera sistema koja služi za identifikaciju osnovnih ili kritičnih grešaka pre nego što se započne detaljnije testiranje.

## End-to-End Testiranje

- **End-to-End Testing**: Testiranje od kraja do kraja, ili E2E, testira scenarije iz stvarnog sveta sa perspektive korisnika. Ovi testovi pomažu da se osigura da vaša aplikacija radi ispravno i da je predvidiva za krajnjeg korisnika. Međutim, zahteva mnogo vremena za podešavanje, dizajniranje i održavanje.
- **End-to-End Testing Frameworks**: Playwright je popularan okvir (framework) za E2E testiranje razvijen od strane Microsofta. Drugi primeri alata za end-to-end testiranje uključuju Cypress, Selenium i Puppeteer.

Ovo je primer E2E testova iz baze koda freeCodeCamp-a, korišćenjem Playwright-a. Hook ``beforeEach`` će se izvršiti pre svakog testa. Testovi proveravaju da donator ima link za sponzora u navigacionoj traci (menu bar), kao i specijalizovani stilizovani okvir oko njegovog avatara:

```js
test.beforeEach(async ({ page }) => {
  execSync("node ./tools/scripts/seed/seed-demo-user --set-true isDonating");
  await page.goto("/donate");
});

...

test("The menu should have a supporters link", async ({ page }) => {
  const menuButton = page.getByTestId("header-menu-button");
  const menu = page.getByTestId("header-menu");

  await expect(menuButton).toBeVisible();
  await menuButton.click();

  await expect(menu).toBeVisible();

  await expect(page.getByRole("link", { name: "Supporters" })).toBeVisible();
});

test("The Avatar should have a special border for donors", async ({ page }) => {
  const container = page.locator(".avatar-container");
  await expect(container).toHaveClass("avatar-container gold-border");
});
```

## Testiranje upotrebljivosti

- **Testiranje upotrebljivosti (Usability Testing)**: Testiranje upotrebljivosti je proces u kojem se omogućava interakcija pravih korisnika sa aplikacijom kako bi se otkrilo da li postoje problemi dizajna, korisničkog iskustva ili funkcionalnosti unutar aplikacije. Testiranje upotrebljivosti fokusira se na intuitivnost aplikacije za krajnje korisnike.
- **Četiri uobičajene vrste testiranja upotrebljivosti**:
  - **Eksplorativno (Explorative)**: Eksplorativno testiranje upotrebljivosti uključuje interakciju korisnika sa različitim funkcijama aplikacije kako bi se bolje razumelo kako te funkcije funkcionišu.
  - **Komparativno (Comparative)**: Komparativno testiranje je proces poređivanja iskustva korišćenja vaše aplikacije sa sličnim aplikacijama na tržištu.
  - **Procena (Assessment)**: Testiranje procene se sprovodi kada proučavate koliko je aplikacija intuitivna za upotrebu.
  - **Validacija (Validation)**: Testiranje validacije služi za identifikovanje bilo kakvih ozbiljnih problema koji bi mogli da spreče korisnika da efikasno koristi aplikaciju.
- **Alati za testiranje upotrebljivosti**: Primeri alata za testiranje upotrebljivosti uključuju Loop11, Maze, Userbrain, UserTesting i UXTweak.

## Testiranje Kompatibilnosti

- **Testiranje kompatibilnosti**: Cilj testiranja kompatibilnosti je osigurati da vaša aplikacija radi u različitim računarskim okruženjima.
- **Različiti tipovi testiranja kompatibilnosti**:
  - **Kompatibilnost sa starijim verzijama (Backwards Compatibility)**: Kompatibilnost sa starijim verzijama odnosi se na to kada je softver kompatibilan sa ranijim verzijama.
  - **Kompatibilnost sa budućim verzijama (Forwards Compatibility)**: Kompatibilnost sa budućim verzijama odnosi se na to kako će softver i sistemi raditi sa budućim verzijama.
  - **Kompatibilnost hardvera**: Kompatibilnost hardvera je sposobnost softvera da radi ispravno u različitim konfiguracijama hardvera.
  - **Kompatibilnost operativnih sistema**: Kompatibilnost operativnih sistema je sposobnost softvera da radi na različitim operativnim sistemima, kao što su macOS, Windows i distribucije Linuxa poput Ubuntu i Fedora.
  - **Mrežna kompatibilnost (Network Compatibility)**: Mrežna kompatibilnost znači da softver može raditi u različitim mrežnim uslovima, kao što su različite brzine mreže, protokoli, sigurnosne postavke itd.
  - **Kompatibilnost pretraživača**: Kompatibilnost pretraživača znači da web aplikacija može dosledno raditi u različitim pregledačima, kao što su Google Chrome, Safari, Firefox itd.
  - **Mobilna kompatibilnost (Mobile Compatibility)**: Važno je osigurati da vaše softverske aplikacije rade na raznovrsnim Android i iOS uređajima, uključujući telefone i tablete.

## Testiranje performansi

- **Performance Testing**: U performans testiranju, testira se brzina, reaktivnost, skalabilnost i stabilnost aplikacije pod različitim opterećenjima. Cilj je eliminisanje bilo kakvog uskog grla performansi.
- **Različite vrste performance testiranja**:
  - **Load Testing**: Load testiranje određuje kako sistem funkcioniše tokom normalnih i vrhunskih perioda opterećenja.
  - **Stress Testing**: Stres testiranje je proces u kojem se aplikacija testira pod ekstremnim opterećenjima radi proveravanja koliko dobro sistem reaguje na veće opterećenje.
  - **Soak Testing (Endurance Testing)**: Soak testiranje ili testiranje izdržljivosti je vrsta load testiranja gde se sistem testira sa konstantnim, većim opterećenjem tokom produženog vremenskog perioda.
  - **Spike Testing**: Spike testiranje je proces u kojem se drastično povećavaju i smanjuju opterećenja, analizirajući reakcije sistema na te promene.
  - **Breakpoint Testing (Capacity Testing)**: Breakpoint testiranje ili testiranje kapaciteta je proces postepenog povećavanja opterećenja tokom vremena do tačke gde sistem počinje da ne funkcioniše ispravno ili da se degradira.

## Testiranje bezbednosti

- **Testiranje bezbednosti**: Testiranje bezbednosti pomaže u identifikovanju ranjivosti i slabosti.
- **Principi bezbednosti**:
  - **Pouzdanost (Confidentiality)**: Štiti od otkrivanja osetljivih informacija drugim primaocima koji nisu namenski primatelj.
  - **Integritet (Integrity)**: Uključuje sprečavanje da zlonamerni korisnici modifikuju podatke korisnika.
  - **Autentikacija (Authentication)**: Podrazumeva proveru identiteta korisnika kako bi se osiguralo da mu je dozvoljeno korišćenje tog sistema.
  - **Autorizacija (Authorization)**: Ovo je proces određivanja koje akcije su autentifikovanim korisnicima dozvoljene ili koji delovi sistema im je dozvoljeno pristupiti.
  - **Dostupnost (Availability)**: Osigurava da su informacije i usluge dostupne ovlašćenim korisnicima kada su im potrebne.
  - **Non-Repudiation**: Osigurava da i pošiljalac i primaoc imaju dokaz o dostavi i verifikaciju identiteta pošiljaoca. Štiti od toga da pošiljalac negira slanje informacija.
- **Uobičajene pretnje bezbednosti**:
  - **Cross-Site Scripting (XSS)**: XSS napadi se dešavaju kada napadač ubaci zlonamenske skripte u web stranicu, a zatim ih izvrši u kontekstu pregledača žrtve.
  - **SQL Injection**: SQL injection omogućava zlonamernim korisnicima da ubace zlonamenski kod u bazu podataka.
  - **Denial-of-Service (DoS) Attack**: DoS napad se dešava kada zlonamerni korisnici preplave veb stranicu velikim brojem zahteva ili saobraćajem, što uzrokuje usporavanje i eventualni pad servera, čineći sajt nedostupnim za korisnike.
- **Kategorije alata za testiranje bezbednosti**:
  - **Static Application Security Testing**: Ovi alati procenjuju izvorni kod aplikacije radi identifikovanja sigurnosnih ranjivosti.
  - **Dynamic Application Security Testing**: Ovi alati komuniciraju sa front-endom aplikacije kako bi otkrili potencijalne bezbednosne slabosti. DAST alati nemaju pristup izvornom kodu.
- **Penetration Testing (pentest)**: Penetracioni testiranje je vrsta testiranja bezbednosti koja uključuje kreiranje simuliranih kibernetičkih napada na aplikaciju radi identifikovanja bilo kakvih ranjivosti u sistemu.

## A/B Testiranje

- **A/B Testiranje**: A/B testiranje podrazumeva poređenje dve verzije stranice ili aplikacije i proučavanje koja od tih verzija bolje funkcioniše. Takođe se naziva i *bucket* ili *split* testiranje. A/B testiranje vam omogućava da donosite više odluka zasnovanih na podacima i kontinuirano poboljšavate korisničko iskustvo.
- **Alati za A/B Testiranje**: Primeri alata koje možete koristiti za A/B testiranje uključuju GrowthBook i LaunchDarkly.

## Alfa i Beta Testiranje

Nakon što su završene faze inicijalnog razvoja i softverskog testiranja, ključno je da aplikacija bude testirana od strane profesionalnih testera i krajnjih korisnika. Ovde dolaze u igru alfa i beta testiranje.

- **Alfa testiranje**: Alfa testiranje se obavlja od strane izabrane grupe testera koji prolaze kroz aplikaciju kako bi osigurali da nema grešaka (bugova) pre nego što bude pušteno na tržište. Alfa testiranje je deo testiranja prihvatanja i koristi tehnike white box i black box testiranja.
- **Beta testiranje**: Beta testiranje je faza u kojoj se aplikacija stavlja na raspolaganje stvarnim korisnicima. Korisnici mogu da interaguju sa aplikacijom i pruže povratne informacije. Beta testiranje je takođe oblik testiranja prihvatanja od strane korisnika (user acceptance testing).
- **Testiranje prihvatanja**: Testiranje prihvatanja osigurava da softverska aplikacija ispunjava poslovne zahteve i potrebe korisnika pre nego što bude puštena.
- **Black Box testiranje**: Black box testiranje se fokusira samo na očekivano ponašanje aplikacije.
- **White Box testiranje**: White box testiranje uključuje testere koji znaju interne komponente i sprovode testove na njima.

## Regresiono testiranje

- **Regresija**: Regresija se odnosi na situacije u kojima nove promene nehotično narušavaju postojeće funkcionalnosti.
- **Regresiono testiranje (Regression Testing)**: Regresiono testiranje pomaže u otkrivanju problema regresije. U regresionom testiranju, ponovo pokrećete funkcionalne testove na delovima vaše aplikacije kako biste osigurali da sve i dalje radi onako kako je očekivano.
- **Alati za regresiono testiranje (Tools for Regression Testing)**: Alati koje možete koristiti za obavljanje regresionog testiranja uključuju Puppeteer, Playwright, Selenium i Cypress.
- **Tehnike za regresiono testiranje (Techniques for Regression Testing)**: 
  - **Jedinično regresiono testiranje (Unit Regression Testing)**: Ovo je situacija u kojoj imate listu stavki koje moraju biti testirane svaki put kada se implementiraju velike promene ili popravci u aplikaciju.
  - **Parcijalno regresiono testiranje (Partial Regression Testing)**: Ovo uključuje ciljane pristupe kako bi se osiguralo da nove promene nisu narušile specifične aspekte aplikacije.
  - **Potpuno regresiono testiranje (Complete Regression Testing)**: Ovo pokreće testove protiv svih funkcionalnosti u bazi koda. Ovo je najzahtevnije i detaljnije opcija.
- **Ponovno testiranje (Retesting)**: Ponovno testiranje se koristi za proveru poznatih problema i osiguravanje da su rešeni. Naspram toga, regresiono testiranje traži nepoznate probleme koji su možda nastali usled nedavnih promena u bazi koda.

## Šta je TypeScript

- **JavaScript**: JavaScript je jezik sa dinamičkom tipizacijom. To znači da varijable mogu prihvatiti bilo koje vrednosti tokom izvršavanja. Izazov dinamički tipizovanih jezika jeste to što nedostatak tipovne sigurnosti može dovesti do grešaka.

Na primer, čak i ako vaša JavaScript funkcija očekuje niz, možete je i dalje pozvati sa brojem:

```javascript
const getRandomValue = (array) => {
  return array[Math.floor(Math.random() * array.length)];
}

console.log(getRandomValue(10));
```

Izlaz za gornji primer biće ``undefined`` koji je rezultat od ``console``.

- **TypeScript**: TypeScript proširuje JavaScript jezik tako što uključuje statično tipiziranje, što pomaže u otkrivanju grešaka uzrokovanih neusklađenim tipovima pre nego što pokrenete svoj kod.

Na primer, možete definisati tip za parametar `array` na sledeći način:

```typescript
const getRandomValue = (array: string[]) => {
  return array[Math.floor(Math.random() * array.length)];
}
```

Ova tip definicija govori TypeScript-u da ako mu prosledite broj kao argument za parametar ``array`` mora biti niz (array) stringova. Zatim, kada pozovete ``getRandomValue``, dobićete grešku koja se zove compiler error.

- **Komajler**: Prvo morate kompajlirati TypeScript kod u običan JavaScript. Kada pokrenete komajler, TypeScript će evaluirati vaš kod i baciti grešku za bilo kakve probleme gde tipovi ne odgovaraju.

## Tipovi Podataka u TypeScriptu
# --description--
TypeScript je super-set jezika JavaScript, koji dodaje statičko tipiziranje. Ovo znači da možete definisati tipove promenljivih i funkcija tokom razvoja, što pomaže u otkrivanju grešaka pre nego što kod uopšte pokrene. U ovom zadatku ćete raditi sa osnovnim tipovima podataka u TypeScriptu: `string`, `number`, `boolean` i `any`.

# --hints--
Zapamtite da je TypeScript dizajniran da poboljša pouzdanost koda, posebno za velike aplikacije. Kada koristite funkcije ili promenljive, pokušajte da eksplicitno definišete očekivani tip podataka kako biste osigurali konzistentnost i smanjili mogućnost grešaka pri izvršavanju.
# --assignment--
Implementirajte TypeScript klasu pod nazivom `Book` koja predstavlja knjigu. Ova klasa mora imati sledeća polja:
1.  `title`: Naslov knjige (tip `string`).
2.  `author`: Autor knjige (tip `string`).
3.  `publicationYear`: Godina izdavanja (tip `number`).

Takođe, implementirajte konstruktor koji automatski inicijalizuje sva tri polja kada se kreira novi objekat klase. Napokon, dodajte metodu pod nazivom `getSummary()` koja vraća string sa formatiranim rezimeom knjige u sledećem obliku: "Naslov je napisan od strane Autora i objavljen je godine [Godina]".

```js
// Vaš kod ovde
```

- **Primitivni tipovi podataka u TypeScriptu**: Za primitivne tipove podataka ``string``, ``null``, ``undefined``, ``number``, ``boolean``, i ``bigint``, TypeScript nudi odgovarajuće ključne reči za tipove.

```typescript
let str: string = "Naomi";
let num: number = 42;
let bool: boolean = true;
let nope: null = null;
let nada: undefined = undefined;
```

- **Array**: Možete definisati niz specifičnog tipa sa dve različite sintakse.

```typescript
const arrOne: string[] = ["Naomi"];
const arrTwo: Array<string> = ["Camperchan"];
```

- **Objekti**: Možete definisati tačnu strukturu objekta.

```typescript
const obj: { a: string, b: number } = { a: "Naomi", b: 42 };
```

Ako želite objekat sa bilo kojim ključevima, ali gde sve vrednosti moraju biti stringovi, postoji dva načina da ga definišete:

```typescript
const objOne: Record<string, string> = {};
const objTwo: { [key: string]: string } = {};
```

- **Ostali korisni tipovi u TypeScriptu**:
  - **`any`**: `any` ukazuje da vrednost može imati bilo koji tip. Govori kompajleru da prestane da brine o tipu te varijable.
  - **`unknown`**: `unknown` govori TypeScriptu da *znaš* za tip vrednosti, ali zapravo ne znaš tačan tip. Preporučuje se korišćenje `unknown` umesto `any`.
  - **`void`**: Ovo je poseban tip koji ćete obično koristiti samo prilikom definisanja funkcija. Za funkcije koje ne vraćaju vrednost, koristi se tip povratne vrednosti `void`.
  - **`never`**: On predstavlja tip koji nikada neće postojati.

- **`type` Keyword**: Ova ključna reč je slična `const`, ali umesto deklarisanja varijable, možete deklarisati tip.

Korisno je za deklarisanje prilagođenih tipova, kao što su unioni tipovi ili tipovi koji sadrže samo specifične vrednosti:

```typescript
type stringOrNumber = string | number;
type bot = "camperchan" | "camperbot" | "naomi";
```

- **`interface`**: Interfejsi su poput klasa za tipove. Oni mogu da implementiraju ili prošire druge interfejse, predstavljajući specifične objekte i generalno se preferiraju osim ako vam nije potrebna određena funkcionalnost koju deklaracija pruža.

```typescript
interface wowie {
  zowie: boolean;
  method: () => void;
}
```

- **Definisati tip povratne vrednosti**: Takođe možete definisati _tip povratne vrednosti_ funkcije.

Primer ispod definiše povratnu vrednost kao string. Ako pokušate da vratite nešto drugo, TypeScript će generisati grešku kompajlera.

```typescript
const getRandomValue = (array: string[]): string => {
  return array[Math.floor(Math.random() * array.length)];
}
```

## Generički tipovi

**Definisanje generičkog tipa**: Možete definisati generički tip i referirati se na njega u svojoj funkciji. To se može smatrati kao poseban parametar koji pružate funkciji kako biste kontrolisali ponašanje definicije tipa i funkcije.

Ovo je primer definisanja generičkog tipa za funkciju:

```typescript
const getRandomValue = <T>(array: T[]): T => {
  return array[Math.floor(Math.random() * array.length)];
}
const val = getRandomValue([1, 2, 4])
```

Sintaksa ``<T>`` govori TypeScript-u da definišete generički tip ``T`` za funkciju. ``T`` je uobičajeno ime za generičke tipove, ali možete koristiti bilo koje ime.

Zatim, kažete TypeScriptu da je parametar ``array`` niz vrednosti koje odgovaraju generičkom tipu, i da vraćena vrednost predstavlja pojedinačan element tog istog tipa.

- **Navođenje Tipskog Argumenta u Pozivu Funkcije**: Možete proslediti tipski argument pozivu funkcije koristeći uglaste zagrade između imena funkcije i njenih parametara.

Evo primera kako se prosljeđuje tipični argument funkcijskom pozivu:

```typescript
const email = document.querySelector<HTMLInputElement>("#email");
console.log(email.value);
```

Ovo govori TypeScriptu da se element koji očekujete da pronađete tretira kao `<input>` element.

## Sužavanje tipa

**Sužavanje po istinitosti**: U primeru ispod, dobijate grešku kompajlera pokušavajući da pristupite svojstvu ``value`` od ``email`` jer ``email`` _može_ biti ``null``.

```typescript
const email = document.querySelector<HTMLInputElement>("#email");
console.log(email.value);
```

Možete koristiti uslovnu izjavu da potvrdite da je ``email`` *truthy* pre nego što pristupite svojstvu:

```typescript
const email = document.querySelector<HTMLInputElement>("#email");
if (email) {
  console.log(email.value);
}
```

Provere istinitosti takođe mogu da rade i u suprotnom smeru:

```typescript
const email = document.querySelector<HTMLInputElement>("#email");
if (!email) {
  throw new ReferenceError("Could not find email element!")
}
console.log(email.value);
```

Bacanje greške prekida logičko izvršavanje ovog koda, što znači da kada dođete do poziva ``console.log()``, TypeScript zna da ``email`` _ne može biti_ ``null``.

**Optionalni lančev**: Optional chaining ``?.`` je takođe oblik sužavljanja tipa (type narrowing), na pretpostavci da pristup svojstvu (property access) neće biti moguć ako vrednost ``email`` jeste ``null``.

```typescript
const email = document.querySelector<HTMLInputElement>("#email");
console.log(email?.value);
```

- **`typeof` Operater**: Možete koristiti uslovnu strukturu za proveru tipa varijable korišćenjem operatera `typeof`.

```typescript
const myVal = Math.random() > 0.5 ? 222 : "222";
if (typeof myVal === "number") {
  console.log(myVal / 10);
}
```

- **`instanceof` Keyword**: Ako objekat dolazi iz klase, možete koristiti ``instanceof`` keyword da suzite tip.

```typescript
const email = document.querySelector("#email");

if (email instanceof HTMLInputElement) {
    console.log(email.value);
}
```

**Kasting**: Kada TypeScript ne može automatski da utvrdi tip vrednosti, kao što je rezultat metode ``request.json()`` u primeru ispod, dobićete grešku kompajlera. Jedan način za rešavanje ovoga je kasting (izričito konvertovanje) tipa, ali to oslabjava sposobnost TypeScript-a da uhvati potencijalne greške.

```typescript
interface User {
    name: string;
    age: number;
}

const printAge = (user: User) => 
  console.log(`${user.name} is ${user.age} years old!`)

const request = await fetch("url")
const myUser = await request.json() as User;
printAge(myUser);
```

- **Type Guard**: Umesto konverzije tipova, možete napisati Type Guard:

```typescript
interface User {
    name: string;
    age: number;
}

const isValidUser = (user: unknown): user is User => {
  return !!user && 
    typeof user === "object" &&
    "name" in user &&
    "age" in user;
}
```

Sintaksa `user is User`` ukazuje da vaša funkcija vraća booleanski vrednost, što kada je tačno znači da vrednost ``user`` ispunjava interfejs ``User`

## `tsconfig` Fajl

- **`tsconfig.json`**: Podešavanja kompajlera TypeScript-a se nalaze u fajlu ``tsconfig.json`` u korenom direktorijumu vašeg projekta.

```json
{
  "compilerOptions": {
    "rootDir": "./src",
    "outDir": "./prod",
    "lib": ["ES2023"],
    "target": "ES2023",
    "module": "ES2022",
    "moduleResolution": "Node",
    "esModuleInterop": true,
    "skipLibCheck": true,
    "strict": true
  },
  "exclude": ["test/"]
}
```

Evo opisa opcija kompilatora korišćene u gornjem primeru:

- **`compilerOptions`**: Svojstvo ``compilerOptions`` je mesto gde kontrolišete kako se TypeScript kompajler ponaša.
- **`rootDir`** i **`outDir`**: ``rootDir`` i ``outDir`` govore TypeScript-u koji direktorijum sadrži vaše izvorni fajlove, a koji bi trebalo da sadrži transpilovani JavaScript kod.
- **`lib`**: Svojstvo ``lib`` određuje koje tip definicije kompajler koristi, i omogućava vam da uključite podršku za specifične ES verzije, DOM i više.
- **`module`** i **`moduleResolution`**: ``module`` i ``moduleResolution`` rade zajedno da upravljaju načinom na koji vaš paket koristi module – bilo CommonJS ili ECMAScript.
- **`esModuleInterop`**: ``esModuleInterop`` omogućava bolju interoperebilnost između CommonJS i ES modula automatskim kreiranjem objekata prostora imena (namespace) za uvoz.
- **`skipLibCheck`**: Opcija ``skipLibCheck`` preskače validaciju fajlova ``.d.ts`` koji nisu referencirani uvozima u vašem kodu.
- **`strict`**: Flag ``strict`` omogućava nekoliko provera, kao što je osiguravanje pravilne obrade nullabilnih tipova i upozorenje kada TypeScript se vraća na ``any``.
- **`exclude`**: Top-level svojstvo ``exclude`` govori kompajleru da ignoriše ove TypeScript fajlove tokom kompilacije.


# --assignment--
Pregledajte teme i koncepte Biblioteka za front-end.