---
id: 683ec95f2f8781355556ff7e
title: Korak 6
challengeType: 20
dashedName: step-6
---

# --description--

Sada kada ste izračunali indeks `mid`, morate zapravo dohvatiti vrednost sa `search_list` na tom indeksu. Ova vrednost je onaj parametar koji ćete uporediti sa vašim `value`, a to je cilj koji tražite.

Kreirajte varijablu `value_at_middle` i dodelite joj element iz `search_list`, koji se nalazi na indeksu `mid`.

# --hints--

You should create a `value_at_middle` variable set to `search_list[mid]`.

```js
({
    test: () =>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2 \\nvalue_at_middle = search_list[mid]")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
--fcc-editable-region--
        
--fcc-editable-region--
```
