---
id: 683ec95f2f8781355556ff7c
title: Korak 4
challengeType: 20
dashedName: step-4
---

# --description--

Sada kada ste definisali granice svog pretrage, vreme je da kreirate petlju koja će izvršiti binarni pretraživački algoritam.

Binarna pretraga radi ponavljajućim sužavanjem traženog prostora. Ovaj proces se nastavlja sve dok postoji validan opseg elemenata za proveru.

Ovo možete izraziti sa ``while`` petljom koja će se nastaviti sve dok je vaš pokazivač ``low`` manji ili jednak vašem pokazivaču ``high``.

Ako ``low`` ikada postane veći od ``high``, to znači da je traženi prostor postao prazan, i da ``value`` nije u listi.

Unutar funkcije ``binary_search``, kreirajte ``while`` petlju sa uslovom koji proverava da li je ``low`` manji ili jednak `high`. Unutar petlje, dodajte ključnu reč ``pass`` kao privremeni zamena.

# --hints--

You should create a `while` loop that runs while `low` is less than or equal to `high`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_conditions()[0]
assert _cond.is_equivalent("low <= high") or _cond.is_equivalent("high >= low")`
) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
