---
id: 683ec95f2f8781355556ff81
title: Korak 9
challengeType: 20
dashedName: step-9
---

# --description--

Moraš da isprobaš stvari kako bi razumeo tok algoritma u ovoj početnoj fazi.

Da bi to uradio, prvo moraš da izađeš iz petlje. To je zato što trenutna implementacija će dozvoliti samo jednu iteraciju, pa ako uslov nije ispunjen, biće beskonačna petlja.

Neposredno nakon naredbe ``if``, koristi ključnu reč ``break`` da izađeš iz petlje ``while``. Zatim, nakon petlje ``while``, vrati prazan spisak kako bi se signalizovalo da vrednost nije pronađena.

# --hints--

You should use the `break` keyword to break out of the `while` loop after the `if` statement.

```js
({
    test: () => assert(runPython(
      `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().has_stmt("break")`)
    )
})
```

You should return an empty list after your `while` loop.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").has_return("[]")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
        
        if value == value_at_middle:
            return path_to_target
--fcc-editable-region--
        
--fcc-editable-region--
```
