---
id: 683ec95f2f8781355556ff7a
title: Korak 2
challengeType: 20
dashedName: step-2
---

# --description--

Algoritam binarnog pretraživanja koji gradite neće samo vratiti `True` ili `False` da bi ukazao da li je vrednost pronađena ili ne. Takođe će vratiti putanju koju koristi za pronalaženje te vrednosti. I za to, imaćete varijablu `path_to_target` da pratite putanju.

Unutar vaše funkcije `binary_search`, zamenite ključnu reč `pass` varijablom `path_to_target` i inicijalizujte je praznim listom, za početak.

# --hints--

You should create a `path_to_target` variable set to an empty list within your `binary_search` function. Don't forget to remove the `pass` keyword.

```js
({ 
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_body().is_equivalent("path_to_target = []") and not _Node(_code).find_function("binary_search").has_pass()`
      )
    ) 
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
