---
id: 683ec95f2f8781355556ff7b
title: Korak 3
challengeType: 20
dashedName: step-3
---

# --description--

U binarnom pretraživanju, postoji najniži mogući indeks i najviši mogući indeks.

Najniži mogući indeks predstavlja najleviju granicu u trenutnom delu vrednosti koje se pretražuju, u ovom slučaju, `search_list`. Najviši mogući indeks, s druge strane, predstavlja najdesniju granicu u vrednostima koje se pretražuju.

Da biste uzela u obzir ta dva, morate razmotriti celu listu koja se pretražuje. Stoga, unutar funkcije `binary_search`, definišite varijablu `low` sa vrednošću `0`, i varijablu `high` sa vrednošću koja uzima u obzir poslednji indeks u listi koja se pretražuje.

# --hints--

You should have a variable named `low` in your `binary_search` function.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("low")`)) 
})
```

Your `low` variable should be set to `0`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_variable("low").is_equivalent("low = 0")`
      )
    )
})
```

You should have a variable named `high` in your `binary_search` function.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("high")`)) 
})
```

Your `high` variable should be set to `len(search_list) - 1`.

```js
({
  test: () => assert(
    runPython(
    `_Node(_code).find_function("binary_search").find_variable("high").is_equivalent("high = len(search_list) - 1")`
    ))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
--fcc-editable-region--
    
--fcc-editable-region--
```
