---
id: 683ec39565dbc7e3375263b1
title: Korak 1
challengeType: 20
dashedName: step-1
---

# --description--

U prethodnoj lekciji ste naučili o binarnom pretraživanju i kako se razlikuje od linearnog pretraživanja.

U ovom radionici ćete to primeniti kreiranjem algoritma binarnog pretraživanja.

Da biste počeli, kreirajte funkciju `binary_search` sa parametrima `search_list` i `value`.

`search_list` je lista elemenata unutar kojih će funkcija pretraživati, a `value` je cilj koji će funkcija pokušati da pronađe unutar `search_list`.

Dodajte ključnu reč `pass` unutar funkcije za sada kako bi test prošao.

# --hints--

You should create a `binary_search` function.

```js
({ 
  test: () => assert(runPython(`_Node(_code).has_function("binary_search")`)) 
})
```

Your `binary_search` function should have a `search_list` and `value` parameters.

```js
({ 
  test: () => assert(
    runPython(
      `_Node(_code).find_function("binary_search").has_args("search_list, value") or _Node(_code).find_function("binary_search").has_args("value, search_list")`
    )
  ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
