---
id: 683ec95f2f8781355556ff7d
title: Korak 5
challengeType: 20
dashedName: step-5
---

# --description--

Da biste počeli sa petljom, morate pronaći srednji indeks trenutnog opsega pretrage.

Srednji indeks možete dobiti izračunavanjem proseka varijabli `low` i `high`. Moraćete da koristite celobrojno deljenje (floor division) kako biste dobili prosek nakon dodavanja dve vrednosti, tako da će odgovor uvek biti ceo broj zaokružen nadole, jer indeksi moraju biti cele brojeve.

Uklonite ključnu reč `pass` iz vaše `while` petlje i kreirajte varijablu `mid` postavljenu na prosek `low` i `high`.

# --hints--

You should remove the `pass` keyword from the while loop and create a `mid` variable set to `(low + high) // 2`.

```js
({
    test: ()=>
        assert(
            runPython(
                `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (low + high) // 2") or _Node(_code).find_function("binary_search").find_whiles()[0].find_body().is_equivalent("mid = (high + low) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
--fcc-editable-region--
        pass
--fcc-editable-region--
```
