---
id: 683ec95f2f8781355556ff83
title: Korak 11
challengeType: 20
dashedName: step-11
---

# --description--

Možete videti da prvi poziv vraća `[3]`. To je zato što je `3` srednja vrednost u `[1, 2, 3, 4, 5]`.

S druge strane, drugi poziv vraća prazan spisak. Ovo se dešava zato što petlja trenutno nema logiku za sužavanje opsega pretrage ažuriranjem promenljivih `low` ili `high`. Deluje samo kao provera jedne središnje tačke i zatim se završava bez daljeg pretraživanja.

Da bi se binarna pretraga nastavila sa sužavanjem pretrage, dodajte blok `elif` koji proverava da li je `value` veće od `value_at_middle`. Za sada dodajte ključnu reč `pass` unutar `elif`.

# --hints--

You should add an `elif` block for when `value` is greater than `value_at_middle`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("value > value_at_middle") or _cond.is_equivalent("value_at_middle < value")
`) })
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
        
        if value == value_at_middle:
            return path_to_target
--fcc-editable-region--
        
--fcc-editable-region--
        break
    
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
```
