---
id: 684169b00a073c9b9f0a703c
title: Korak 14
challengeType: 20
dashedName: step-14
---

# --description--

Da li ste primetili da drugi poziv funkcije sada ima vrednosti `3, 5, 4` u listi? To ukazuje da je binarna pretraga funkcioniše kako što je predviđeno.

Ovo se dogodilo ovako: Algoritam je prvo proverio `3` kao sredinu početne liste. Pošto je `4` veće od `3`, pretraga se pomerila na desnu polovinu. Zatim je pregledao `5` kao novu sredinu. Budući da je `4` manje od `5`, pretraga se pomerila ulevo, na kraju identifikujući `4` kao sredinu konačnog raspona.

Da biste ponovo testirali funkciju, pozovite je sa `[1, 3, 5, 9, 14, 22], 10` i odmah ispišite poziv. Ovo je situacija u kojoj vrednost neće biti pronađena.

# --hints--

You should call the function with `binary_search([1, 3, 5, 9, 14, 22], 10)` and print it right away.

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_call("print(binary_search([1, 3, 5, 9, 14, 22], 10))")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
        
        if value == value_at_middle:
            return path_to_target
        elif value > value_at_middle:
            low = mid + 1
        else:
            high = mid - 1
    
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
--fcc-editable-region--

--fcc-editable-region--
```
