---
id: 594810f028c0303b75339acd
title: "'Uainishaji wa nambari kamili chanya: nyingi, duni na kamilifu'"
challengeType: 1
forumTopicId: 302221
dashedName: abundant-deficient-and-perfect-number-classifications
---

# --description--

Ove definišu tri tipa pozitivnih celih brojeva u zavisnosti od njihovih prostih faktora.

Neka je $P(n)$ zbir svih pravilnih delitelja od `n`, gde su svi pravilni delitelji pozitivni ceo brojevi `n`, osim samog `n`.

Kama `P(n) < n` basi `n` huainishwa kama `deficient`

Ako `P(n) === n` onda `n` klasifikuje se kao `perfect`

Ako `P(n) > n`, onda `n` se klasifikuje kao `abundant`

**Mfano**: `6` ina vigawanyo halisi vya `1`, `2`, na `3`. `1 + 2 + 3 = 6`, kwa hiyo `6` huainishwa kama nambari kamilifu.

# --instructions--

Implementiraj funkciju koja broji koliko je pozitivnih celih broja od `1` do `num` (uključujući) nalazi se u svakoj od tri klase. Vrati rezultate kao niz podataka sa formatom `[deficient, perfect, abundant]`.

# --hints--

`getDPA` Trebalo bi da bude funkcija.

```js
assert(typeof getDPA === 'function');
```

`getDPA(5000)` inapaswa kurudisha safu ya data.

```js
assert(Array.isArray(getDPA(5000)));
```

`getDPA(5000)` Vraćeni red mora imati dužinu od `3`.

```js
assert(getDPA(5000).length === 3);
```

`getDPA(5000)` treba da vrati `[3758, 3, 1239]`.

```js
assert.deepEqual(getDPA(5000), [3758, 3, 1239]);
```

`getDPA(10000)` treba vratiti `[7508, 4, 2488]`.

```js
assert.deepEqual(getDPA(10000), [7508, 4, 2488]);
```

`getDPA(20000)` treba da vrati `[15043, 4, 4953]`.

```js
assert.deepEqual(getDPA(20000), [15043, 4, 4953]);
```

# --seed--

## --seed-contents--

```js
function getDPA(num) {

}
```

# --solutions--

```js
function getDPA(num) {
  const dpa = [1, 0, 0];
  for (let n = 2; n <= num; n += 1) {
    let ds = 1;
    const e = Math.sqrt(n);
    for (let d = 2; d < e; d += 1) {
      if (n % d === 0) {
        ds += d + (n / d);
      }
    }
    if (n % e === 0) {
      ds += e;
    }
    dpa[ds < n ? 0 : ds === n ? 1 : 2] += 1;
  }
  return dpa;
}
```
