---
id: 594810f028c0303b75339ad1
title: Brojevi sreće
challengeType: 1
forumTopicId: 302280
dashedName: happy-numbers
---

# --description--

Broj sreće je definisan sledećim procesom:

Kuanzia na nambari yoyote kamili chanya, badilisha nambari hiyo kwa jumla ya mraba wa tarakimu zake, kisha rudia mchakato huu hadi nambari iwe sawa na `1` (ambayo itabaki hapo), au itazunguka bila mwisho katika mzunguko usiojumuisha `1`. Nambari ambazo mchakato huu unamalizika kwa `1` ni nambari za furaha, wakati zile ambazo hazimaliziki kwa `1` ni nambari zisizofuraha.

# --instructions--

Implementiraj funkciju koja vraća true ako je broj srećan, ili false ako nije.

# --hints--

`happy` Trebalo bi da bude funkcija.

```js
assert(typeof happy === 'function');
```

`happy(1)` Trebalo bi da vrati vrednost istina ili laž.

```js
assert(typeof happy(1) === 'boolean');
```

`happy(1)` inapaswa kurejesha `true`.

```js
assert(happy(1));
```

`happy(2)` treba da resetuje `false`.

```js
assert(!happy(2));
```

`happy(7)` treba da vrati `true`.

```js
assert(happy(7));
```

`happy(10)` trebalo bi da vrati `true`.

```js
assert(happy(10));
```

`happy(13)` treba da vrati `true`.

```js
assert(happy(13));
```

`happy(19)` treba da restaurira `true`.

```js
assert(happy(19));
```

`happy(23)` treba da vrati `true`.

```js
assert(happy(23));
```

`happy(28)` treba da obnovi `true`.

```js
assert(happy(28));
```

`happy(31)` trebalo bi da vrati `true`.

```js
assert(happy(31));
```

`happy(32)` treba da vrati `true`.

```js
assert(happy(32));
```

`happy(33)` trebalo bi da vrati `false`.

```js
assert(!happy(33));
```

# --seed--

## --seed-contents--

```js
function happy(number) {

}
```

# --solutions--

```js
function happy (number) {
  let m;
  let digit;
  const cycle = [];

  while (number !== 1 && cycle[number] !== true) {
    cycle[number] = true;
    m = 0;
    while (number > 0) {
      digit = number % 10;
      m += Math.pow(digit, 2);
      number = (number - digit) / 10;
    }
    number = m;
  }
  return (number === 1);
}
```
