---
id: 59667989bf71cf555dd5d2ff
title: Izgovor S
challengeType: 1
forumTopicId: 302303
dashedName: s-expressions
---

# --description--

<a href="https://rosettacode.org/wiki/S-expressions" target="_blank" rel="noopener noreferrer nofollow">Usemi wa S</a> ni njia rahisi ya kuchambua na kuhifadhi data.

# --instructions--

Implementirajte jednostavan čitač/analizator S Formata koji obrađuje sekvence znakova sa i bez navodnika, celi brojeve i decimalne brojeve.

Kitendakazi kinapaswa kusoma Usemi wa S mmoja tu lakini ulio pangwa ndani kutoka kwenye mfuatano wa herufi na kurudisha kama safu ya data (zilizopangwa ndani).

Mstari mpya na nafasi nyingine za tupu zinaweza kupuuzwa isipokuwa ziko ndani ya mfuatano wa herufi ulio na nukuu.

"`()`" unutar sekvence karaktera sa navodnicima smatra se delom sekvence karaktera, a ne prevodi se.

Kushughulikia nukuu zilizokimbizwa ndani ya mfuatano wa herufi ni hiari; kwa hivyo "`(foo"bar)`" inaweza kuchukuliwa kama mfuatano wa herufi "`foo"bar`", au kama kosa.

Kwa hili, msomaji haahitaji kutambua `\` kwa ajili ya kukimbiza, lakini pia anapaswa kutambua nambari ikiwa lugha ina aina zinazofaa za data.

Zapamtite da, osim `()"` (`\` ako trčanje pomaže) i praznih mesta, nema posebnih slova. Ništa drugo nije dozvoljeno bez navodnika.

Msomaji anapaswa kuwa na uwezo wa kusoma ingizo lifuatalo

<pre>((data "quoted data" 123 4.5)
(data (!@# (4.5) "(more" "data)")))
</pre>

na kulibadilisha kuwa muundo wa data wa asili.

# --before-each--

```js
const simpleSExpr = '(data1 data2 data3)';
const simpleSolution = ['data1', 'data2', 'data3'];

const basicSExpr = '((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))';
const basicSolution = [["data","\"quoted data\"",123,4.5],["data",["!@#",[4.5],"\"(more\"","\"data)\""]]];
```

# --hints--

`parseSexpr` Trebalo bi da bude funkcija.

```js
assert(typeof parseSexpr === 'function');
```

`parseSexpr('(data1 data2 data3)')` treba da vrati `['data1', 'data2', 'data3']`

```js
assert.deepEqual(parseSexpr(simpleSExpr), simpleSolution);
```

`parseSexpr('((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))')` treba da vrati `[['data', '"quoted data"', 123, 4.5], ['data', ['!@#', [4.5], '"(more"', '"data)"']]]`.

```js
assert.deepEqual(parseSexpr(basicSExpr), basicSolution);
```

# --seed--

## --seed-contents--

```js
function parseSexpr(str) {

  return true;
}
```

# --solutions--

```js
function parseSexpr(str) {
  const t = str.match(/\s*("[^"]*"|\(|\)|"|[^\s()"]+)/g);
  for (var o, c = 0, i = t.length - 1; i >= 0; i--) {
    var n,
      ti = t[i].trim();
    if (ti == '"') return;
    else if (ti == '(') t[i] = '[', c += 1;
    else if (ti == ')') t[i] = ']', c -= 1;
    else if ((n = +ti) == ti) t[i] = n;
    else t[i] = `'${ti.replace('\'', '\\\'')}'`;
    if (i > 0 && ti != ']' && t[i - 1].trim() != '(') t.splice(i, 0, ',');
    if (!c) if (!o) o = true; else return;
  }
  return c ? undefined : eval(t.join(''));
}
```
