---
id: 5a23c84252665b21eecc7ec5
title: Problem Josephusa
challengeType: 1
forumTopicId: 302294
dashedName: josephus-problem
---

# --description--

Tatizo la Josephus ni fumbo la hisabati lenye maelezo mabaya: wafungwa $n$ wamesimama mduara, wakiwa wamepangwa kwa mfululizo kutoka $0$ hadi $n-1$.

Izvršilac trči oko tog kruga, počevši od zatvorenika broja $0$, uklanjajući svakog zatvorenika sa $k$ i ubijajući ga.

Mchakato huu unaendelea, mduara unakuwa mdogo zaidi na zaidi, hadi mfungwa mmoja tu abaki, ambaye kisha hurejeshwa uhuru.

Na primer, ako ima zatvorenika $n=5$ i $k=2$, sekvenca ubijenih zatvorenika (nazovimo je „sekvenca ubijanja“) biće 1, 3, 0, i 4, a preživeli će biti #2.

Ako je za bilo koje $n, k > 0$, pronađi koji element će na kraju biti jedan od njih.

Katika tukio moja, kulikuwa na wafungwa 41 na kila mfungwa wa tatu alikuwa akiuwawa ($k=3$).

Miongoni mwao alikuwa jamaa mwerevu aitwaye Josephus ambaye alitatua tatizo hilo, akasimama katika nafasi ya kuishi, na akaishi kuendelea kusimulia hadithi.

Koja je tačka bila njega?

# --instructions--

Napiši funkciju koja uzima početni broj zatvorenika i `k` kao parametre i vraća broj zatvorenika koji je najverovatnije preživeli.

# --hints--

`josephus` To bi trebalo da bude funkcija.

```js
assert(typeof josephus == 'function');
```

`josephus(30,3)` inapaswa kurudisha nambari.

```js
assert(typeof josephus(30, 3) == 'number');
```

`josephus(30,3)` treba da vrati `28`.

```js
assert.equal(josephus(30, 3), 28);
```

`josephus(30,5)` treba da vrati `2`.

```js
assert.equal(josephus(30, 5), 2);
```

`josephus(20,2)` treba da vrati `8`.

```js
assert.equal(josephus(20, 2), 8);
```

`josephus(17,6)` treba da vrati `1`.

```js
assert.equal(josephus(17, 6), 1);
```

`josephus(29,4)` treba da vrati `1`.

```js
assert.equal(josephus(29, 4), 1);
```

# --seed--

## --seed-contents--

```js
function josephus(init, kill) {

}
```

# --solutions--

```js
function josephus(init, kill) {
  const arr = Array.from(Array(init).keys());
  let curr = -1
  while (arr.length > 1) {
    curr = (curr + kill) % arr.length;
    arr.splice(curr, 1);
    curr--;
  }
  return arr[0];
}
```
