---
id: 673b567e3ba535dda140d278
title: Izgradite sistem glasanja
challengeType: 26
dashedName: build-a-voting-system
saveSubmissionToDB: true
---

# --description--

U ovom laboratorijumu, izgradit ćete sistem glasanja koji koristi `Map` za kreiranje ankete i `Set` za sprečavanje duplog glasanja.

**Cilj:** Ispunite priče korisnika ispod i učinite da svi testovi prođu kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebali biste inicijalizovati varijablu `poll` novim objektom `Map`.

2. Trebali biste imati funkciju `addOption` koja prima parametar `option`.

3. U funkciji `addOption`:

   - Ako `option` već ne postoji u anketi, treba ga dodati u anketu sa praznim `Set` kao vrednostju za praćenje glasača. Takođe trebate vratiti poruku `Option "<option>" added to the poll.`

   - Ako `option` već postoji, treba da vrati poruku `Option "<option>" already exists.`

   - Ako pokušate dodati praznu opciju, funkcija treba da vrati poruku `Option cannot be empty.`

4. Trebali biste imati funkciju `vote` koja prima dva parametra, `option` (opcija za koju se glasa) i `voterId` (jedinstveni ID za glasača).

5. U funkciji `vote`:

   - Ako `option` ne postoji u anketi, funkcija treba da vrati poruku `Option "<option>" does not exist.`

   - Ako `option` postoji, funkcija treba da proveri da li je `voterId` već glasao za ovu `option`.

   - Ako glasač već glasao, funkcija treba da vrati poruku `Voter <voterId> has already voted for "<option>".`

   - Ako glasač nije glasao, `voterId` treba dodati u `Set` za `voters` za ovu opciju. Funkcija treba da vrati poruku `Voter <voterId> voted for "<option>".`

6. Trebali biste imati najmanje tri opcije u vašem `poll`.

7. Vaš `poll` treba da ima najmanje tri glasa.

8. Trebali biste imati funkciju `displayResults` koja vraća rezultate ankete u sledećem formatu:

```js
Poll Results:
OptionA: N votes
OptionB: N votes
.
.

/*
sample output

Poll Results:
Turkey: 2 votes
Morocco: 1 votes
*/
```

# --hints--

You should have a `poll` variable initialized to a new `Map` object.

```js
assert.exists(poll);
assert.instanceOf(poll, Map);
```

You should define a function `addOption` that accepts a parameter `option`.

```js
assert.isFunction(addOption);
assert.match(
  addOption.toString(),
  /^function\s+addOption\s*\(\s*option\s*\)\s*\{|\(?\s*option\s*\)?\s*=>/
);
```

You should define a function `vote` that accepts two parameters, `option` and `voterId`.

```js
assert.isFunction(vote);
assert.match(vote.toString(), /\(\s*option\s*,\s*voterId\s*\)/);
```

You should define a function `displayResults` to display the poll results.

```js
assert.isFunction(displayResults);
```

You should have at least three options in your `poll`.

```js
assert.isAtLeast(poll.size, 3);
```

Your `poll` should have at least three votes.

```js
let totalVotes = 0;
for (let voters of poll.values()) {
  totalVotes += voters.size;
}
assert.isAtLeast(totalVotes, 3);
```

You should ensure each voting option maps to a `Set` object.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();

  addOption('Turkey');
  assert.instanceOf(poll.get('Turkey'), Set);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`addOption("Egypt")` should return `Option "Egypt" added to the poll.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(addOption('Egypt'), 'Option "Egypt" added to the poll.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Adding an empty option should return `"Option cannot be empty."`

```js
assert.equal(addOption(''), 'Option cannot be empty.');
```

When `Turkey` is already added, `addOption("Turkey")` should return `Option "Turkey" already exists.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Turkey');
  assert.equal(addOption('Turkey'), 'Option "Turkey" already exists.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

When `Malaysia` exists in the voting options, `vote("Malaysia", "traveler1")` should return `Voter traveler1 voted for "Malaysia".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Malaysia');
  assert.equal(
    vote('Malaysia', 'traveler1'),
    'Voter traveler1 voted for "Malaysia".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`vote` should update the `Set` of voters for an option.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Alaska');
  vote('Alaska', 'seal');
  assert.isTrue(poll.get('Alaska').has('seal'));
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

When `traveler1` tries to vote for `Algeria` again, `vote("Algeria", "traveler1")` should return `Voter traveler1 has already voted for "Algeria".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Algeria');
  vote('Algeria', 'traveler1');
  assert.equal(
    vote('Algeria', 'traveler1'),
    'Voter traveler1 has already voted for "Algeria".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Duplicate votes should not increase the size of the `Set`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  vote('Turkey', 'traveler1');
  const voterSet = poll.get('Turkey');
  const initialSize = voterSet.size;

  vote('Turkey', 'traveler1');
  assert.equal(voterSet.size, initialSize);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

When `Nigeria` is not in the voting options, `vote("Nigeria", "traveler2")` should return `Option "Nigeria" does not exist.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(
    vote('Nigeria', 'traveler2'),
    'Option "Nigeria" does not exist.'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

A unique option should be able to receive multiple votes.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Bali');
  assert.equal(vote('Bali', 'traveler1'), 'Voter traveler1 voted for "Bali".');
  assert.equal(vote('Bali', 'traveler2'), 'Voter traveler2 voted for "Bali".');
  assert.equal(poll.get('Bali').size, 2);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

When option `Turkey` gets two votes, option `Morocco` gets one vote, and option `Spain` doesn't get any vote, `displayResults()` should return `"Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes"`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');
  addOption('Spain');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`displayResults()` should return the results in the correct format.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const poll = new Map();

function addOption(option) {
  if (!option || option.trim() === '') {
    return 'Option cannot be empty.';
  }
  if (!poll.has(option)) {
    poll.set(option, new Set());
    return `Option "${option}" added to the poll.`;
  } else {
    return `Option "${option}" already exists.`;
  }
}

function vote(option, voterId) {
  if (!poll.has(option)) {
    return `Option "${option}" does not exist.`;
  }
  const voters = poll.get(option);
  if (voters.has(voterId)) {
    return `Voter ${voterId} has already voted for "${option}".`;
  } else {
    voters.add(voterId);
    return `Voter ${voterId} voted for "${option}".`;
  }
}

function displayResults() {
  let results = 'Poll Results:\n';
  for (let [option, voters] of poll.entries()) {
    results += `${option}: ${voters.size} votes\n`;
  }
  return results.trim();
}

addOption('Turkey');
addOption('Morocco');
addOption('Spain');

vote('Turkey', 'traveler1');
vote('Turkey', 'traveler2');
vote('Morocco', 'traveler3');
```
