---
id: 68ed661e9f463c155bb677a1
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovom radionici ćete vežbati rad sa CSS flexbox-om dizajniranjem niza kolor kutija.

Da biste počeli, unutar elementa `body`, dodajte element `header`. Ugnježđujte element `h1` unutar elementa `header` sa tekstom `Colored Boxes Layout`.

# --hints--

You should have a `header` element inside the `body` element.

```js
const headerEl = document.querySelector('body > header');
assert.exists(headerEl);
```

Your `header` element should have an `h1` element inside it.

```js
const h1El = document.querySelector('body > header > h1');
assert.exists(h1El);
```

Your `h1` element should have the text `Colored Boxes Layout` inside it.

```js
const h1El = document.querySelector('body > header > h1');
assert.equal(h1El?.textContent.trim(), 'Colored Boxes Layout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
</head>
<body>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```
