---
id: 68f91df5776e693049fb8942
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Sada je vreme da dodate neke CSS pravila u `styles.css` datoteku. Počnite kreiranjem selektora za element `h1`.

Centrirajte sadržaj elementa `h1` postavljanjem njegove svojstva `text-align` na `center`. Takođe kreirajte svojstvo `margin-bottom` sa vrednošću `10px` da dodate margina između dna elementa `h1` i bilo kog HTML elementa koji sledi.

# --hints--

You should have an `h1` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

Your `h1` selector should set the `text-align` property to `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('text-align'), 'center');
```

Your `h1` selector should set the `margin-bottom` property to `10px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h1')?.getPropVal('margin-bottom'), '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```

