---
id: 66f686b8ebdb982fa8e14330
title: Napravite program za konverziju datuma
challengeType: 26
dashedName: lab-date-conversion
---

# --description--

U ovom laboratorijumu, kreiraćete program za konverziju datuma koji konvertuje dati datum u različite formate. Na primer, trenutni datum `Fri Sep 27 2024 16:04:43 GMT+0500 (Pakistan Standard Time)` bi bio formatiran na sledeće 2 načina:

- `(MM/DD/YYYY): 9/27/2024`.
- `(Month Day, Year): September 27, 2024`.

**Cilj:** Ispunite priče korisnika ispod i osigurajte da svi testovi prođu kako biste završili laboratorijum. 

**Priče korisnika:**

1. Trebalo bi da kreirate varijablu pod nazivom `currentDate` i dodelite joj trenutni datum i vreme koristeći objekat `Date`.
2. Trebalo bi da kreirate varijablu pod nazivom `currentDateFormat` i dodelite joj string `Current Date and Time: [current date]`. Zamenite `[current date]` rezultatom varijable `currentDate`.
3. Trebalo bi da logujete vrednost `currentDateFormat` u konzolu.
4. Trebalo bi da kreirate funkciju pod nazivom `formatDateMMDDYYYY` koja prima objekat datuma kao parametar. Ovaj parametar možete nazvati kako god želite.
5. Vaša funkcija `formatDateMMDDYYYY` bi trebalo da konvertuje trenutni datum u format `month/day/year` i vrati string `Formatted Date (MM/DD/YYYY): [month]/[day]/[year]`.
6. Trebalo bi da kreirate funkciju pod nazivom `formatDateLong` koja prima objekat datuma kao parametar.
7. Vaša funkcija `formatDateLong` bi trebalo da konvertuje trenutni datum u format `Month Day, Year` i vrati string `Formatted Date (Month Day, Year): [formatted date]`.

**Napomena:** Da bi testovi prošli, uverite se da koristite `en-US` za lokaciju (locale) prilikom formatiranja datuma.

# --hints--

You should have a variable named `currentDate` that holds the current date and time using the `Date` object.

```js
// excluded timestamps
const expectedDate = new Date();
const expectedDateOnly = new Date(expectedDate.getFullYear(), expectedDate.getMonth(), expectedDate.getDate());

const currentDateOnly = new Date(currentDate.getFullYear(), currentDate.getMonth(), currentDate.getDate());

assert.equal(currentDateOnly.toString(), expectedDateOnly.toString());
```

You should have a variable named `currentDateFormat` that holds the current date in the format `Current Date and Time: <ddd> <MMM> <dd> <yyyy> <HH>:<mm>:<ss> <TIMEZONE>`.

```js
const expectedDateString = `Current Date and Time: `;
assert.include(currentDateFormat, expectedDateString);

const currentTimestamp = new Date(currentDateFormat.replace(expectedDateString, '')).getTime();
const expectedTimestamp = new Date().getTime();
assert.approximately(currentTimestamp, expectedTimestamp, 2000);
```

You should log the value of `currentDateFormat` to the console.

```js
assert.match(code, /console\.log\(\s*currentDateFormat\s*\)/);
```

You should have a function `formatDateMMDDYYYY` 

```js
assert.isFunction(formatDateMMDDYYYY);
```

The function `formatDateMMDDYYYY` should take a single parameter.

```js
assert.lengthOf(formatDateMMDDYYYY, 1);
```

When the date object holds `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`, the function `formatDateMMDDYYYY` should return `Formatted Date (MM/DD/YYYY): 9/27/2024`.

```js
let dateString = "Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)";
let dateObject = new Date(dateString);

assert.equal(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/27/2024");

// prevent hardcoding

dateString = "Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)";
dateObject = new Date(dateString);

assert.strictEqual(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/28/2024");
```

You should have a function `formatDateLong` 

```js
assert.isFunction(formatDateLong);
```

The function `formatDateLong` should take a single a parameter.

```js
assert.lengthOf(formatDateLong, 1);
```

When the date object holds `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`, the function `formatDateLong` should return `Formatted Date (Month Day, Year): September 27, 2024`.

```js
let testDate = new Date('Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)');
let result = formatDateLong(testDate);

assert.strictEqual(result,'Formatted Date (Month Day, Year): September 27, 2024', `Test failed: ${result}`); 

// prevent hardcoding

testDate = new Date('Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)');
result = formatDateLong(testDate);
assert.strictEqual(result, 'Formatted Date (Month Day, Year): September 28, 2024'); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const currentDate = new Date();
const currentDateFormat = `Current Date and Time: ${currentDate}`; 
console.log(currentDateFormat);

function formatDateMMDDYYYY(date) {
  const month = date.getMonth() + 1;
  const day = date.getDate();
  const year = date.getFullYear();
  return `Formatted Date (MM/DD/YYYY): ${month}/${day}/${year}`;
}

function formatDateLong(date) {
  const longDateFormat =  date.toLocaleDateString('en-US', {
    year: 'numeric',
    month: 'long',
    day: 'numeric'
  });
  return `Formatted Date (Month Day, Year): ${longDateFormat}`;
}

// usage
console.log(formatDateMMDDYYYY(currentDate));
console.log(formatDateLong(currentDate));

```

