---
id: 68dbc6b7e8f794a2ae69bc74
title: Implementacija N-Kraljevskog algoritma
challengeType: 27
dashedName: implement-the-n-queens-algorithm
---

# --description--

Problem N Kraljica traži da postavite N kraljica na NxN šahovnicu tako da se nijedna dve kraljice ne napadaju međusobno (da dve ne dele red, kolonu ili dijagonalu).

Na primer, ako postoji šahovnica 4x4, jedna validna raspored je:

```md
[1, 3, 0, 2]
```

To znači da je u redu 0 kraljica postavljena u kolonu 1; u redu 1, kraljica je postavljena u kolonu 3; u redu 2, kraljica je postavljena u kolonu 0; a u redu 3, kraljica je postavljena u kolonu 2.

Vizuelno, ovaj raspored izgleda ovako:

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

Gde `Q` predstavlja kraljicu, a `.` praznu polju.

U ovom laboratorijumu, implementiraćete rešivač problema N Kraljica koristeći pristup pretrazi dubine (depth-first search).


**Cilj**: Ispunite priče korisnika ispod i prođite sve testove kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da imate funkciju pod nazivom `dfs_n_queens`.
2. Funkcija treba da prihvati tačno jedan argument: ceo broj `n`.
3. Ako je `n` manji od `1`, funkcija bi trebalo da vrati prazan spisak (`[]`).
4. Funkcija treba da vrati spisak rešenja; svako rešenje je sam po sebi spisak dužine `n`, gde element na indeksu `i` predstavlja indeks kolone (0-baziran) kraljice u redu `i`.

# --hints--

You should have a function named `dfs_n_queens` that takes one argument.

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs_n_queens)
  sig = inspect.signature(dfs_n_queens)
  assert len(sig.parameters) == 1
`) })
```

If `n` is less than `1`, the function should return an empty list.

```js
({ test: () => runPython(`
  assert dfs_n_queens(0) == []
  assert dfs_n_queens(-1) == []
  assert dfs_n_queens(-5) == []
`) })
```

The function should return a list of solutions, where each solution is a list of length `n`.

```js
({ test: () => runPython(`
  result = dfs_n_queens(4)
  assert isinstance(result, list)
  for solution in result:
    assert isinstance(solution, list)
    assert len(solution) == 4
`) })
```

`dfs_n_queens(1)` should return `[[0]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(1) == [[0]]
`) })
```

`dfs_n_queens(2)` should return `[]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(2) == []
`) })
```

`dfs_n_queens(3)` should return `[]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(3) == []
`) })
```

`dfs_n_queens(4)` should return `[[1, 3, 0, 2], [2, 0, 3, 1]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(4) == [[1, 3, 0, 2], [2, 0, 3, 1]]
`) })
```

`dfs_n_queens(5)` should return `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(5) == [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]
`) })
```

`len(dfs_n_queens(5))` should be `10`.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(5)) == 10
`) })
```

`len(dfs_n_queens(8))` should be `92`.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(8)) == 92
`) })
```

`dfs_n_queens` should return the correct result.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(7)) == 40
  assert dfs_n_queens(6) == [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs_n_queens(n):
    if n < 1:
        return []

    results = []

    def is_safe(queens, row, col):
        for r, c in enumerate(queens):
            if c == col or abs(row - r) == abs(col - c):
                return False
        return True

    def dfs(queens):
        row = len(queens)
        if row == n:
            results.append(queens[:])
            return
        for col in range(n):
            if is_safe(queens, row, col):
                queens.append(col)
                dfs(queens)
                queens.pop()

    dfs([])
    return results
```
