---
id: a5229172f011153519423690
title: Izgraditi kalkulator zbira neparnih Fibonaccijevih brojeva
challengeType: 26
dashedName: build-an-odd-fibonacci-sum-calculator
---

# --description--

U ovom laboratorijskom zadatku ćete izgraditi kalkulator zbira neparnih Fibonaccijevih brojeva koji računa zbir svih neparnih Fibonaccijevih brojeva koji su manji ili jednaki datoj pozitivnom celom broju.

**Cilj**: Ispunite priče korisnika ispod i prođite sve testove kako biste završili laboratorijski zadatak.

**Priče korisnika:**

1. Trebalo bi da imate funkciju ``sumFibs`` koja prima broj kao argument.
1. Funkcija ``sumFibs`` bi trebalo da vrati zbir svih neparnih Fibonaccijevih brojeva koji su manji ili jednaki datoj broju.
1. Fibonaccijanska sekvenca počinje sa ``0`` i ``1``, a svaki sledeći broj je zbir prethodnih dva.
1. U zbir se treba dodati samo neparni Fibonaccijevski brojevi.

# --hints--

You should have a `sumFibs` function.

```js
assert.isFunction(sumFibs);
```

`sumFibs(1)` should return a number.

```js
assert.isNumber(sumFibs(1));
```

`sumFibs(1000)` should return `1785`.

```js
assert.strictEqual(sumFibs(1000), 1785);
```

`sumFibs(4000000)` should return `4613732`.

```js
assert.strictEqual(sumFibs(4000000), 4613732);
```

`sumFibs(4)` should return `5`.

```js
assert.strictEqual(sumFibs(4), 5);
```

`sumFibs(75024)` should return `60696`.

```js
assert.strictEqual(sumFibs(75024), 60696);
```

`sumFibs(75025)` should return `135721`.

```js
assert.strictEqual(sumFibs(75025), 135721);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumFibs(num) {
  let a = 0, b = 1;
  let sum = 0;

  while (b <= num) {
    if (b % 2 !== 0) {
      sum += b;
    }
    [a, b] = [b, a + b];
  }

  return sum;
}
```
