---
id: 68f890b73bbfc9efa7379e7b
title: Pronađite prisiljene tipove grešaka u aplikaciji sa greškama
challengeType: 26
dashedName: debug-type-coercion-errors-in-a-buggy-app
---

# --description--

Tek ste se pridružili lokalnoj web razvojnoj radnici/platformi, i vaš prvi zadatak je da očistite kod sa greškama koji su ostavili prethodni programeri.

Pokušavali su da izvrše jednostavne aritmetičke operacije, ali nešto je pošlo naopako, rezultati nemaju smisla. Vaš zadatak je pregledati, pronaći greške i ispraviti kod tako da radi kako treba i da bude lako čitljiv.

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da imaš parametar nazvan `firstResult` koji precizno dodaje brojeve `5` i `10` kako bi proizveo vrednost `15`.
2. Trebalo bi da imaš parametar nazvan `secondResult` koji precizno oduzima `5` od `8` kako bi proizveo vrednost `3`.
3. Trebalo bi da imaš parametar nazvan `thirdResult` koji vraća vrednost `6` dodavanjem dva broja. Izmeni trenutnu istinu ili lažnu vrednost koja se koristi u izrazu za proračun i broj.
4. Trebalo bi da imaš parametar nazvan `fourthResult` koji vraća vrednost `8` dodavanjem dva broja. Izmeni trenutnu istinu ili lažnu vrednost koja se koristi u izrazu za proračun i broj.
5. Trebalo bi da imaš parametar nazvan `fifthResult` koji precizno množi dva broja kako bi dobio proizvod `20`. Izmeni trenutni niz znakova koja se koristi u izrazu za proračun i broj.
6. Trebalo bi da imaš parametar nazvan `sixthResult` koji precizno dodaje dva broja kako bi proizveo vrednost `22`. Izmeni `null` koja se trenutno koristi u izrazu za proračun i broj.

# --hints--

Trebalo bi da imaš parametar nazvan `firstResult`.

```js
assert.match(code, /const\s+firstResult\s*=/);
```

`firstResult` mora imati numeričku vrednost `15`.

```js
assert.strictEqual(firstResult, 15);
```

Trebalo bi da imaš parametar koji se zove `secondResult`.

```js
assert.match(code, /const\s+secondResult\s*=/);
```

`secondResult` treba da ima numeričku vrednost `3`.

```js
assert.strictEqual(secondResult, 3);
```

Trebalo bi da imate parametar koji se zove `thirdResult`.

```js
assert.match(code, /const\s+thirdResult\s*=/);
```

`thirdResult` treba da ima numeričku vrednost `6`.

```js
assert.strictEqual(thirdResult, 6);
```

`thirdResult` ne bi trebalo da sadrži vrednost tačno ili netačno (`true` ili `false`).

```js
const match = code.match(/const\s+thirdResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

Trebalo bi da imaš parametar nazvan `fourthResult`.

```js
assert.match(code, /const\s+fourthResult\s*=/);
```

`fourthResult` mora imati numeričku vrednost `8`.

```js
assert.strictEqual(fourthResult, 8);
```

`fourthResult` ne bi smelo da sadrži vrednost istinito ili lažno (`true` ili `false`).

```js
const match = code.match(/const\s+fourthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

Trebalo bi da imaš parametar nazvan `fifthResult`.

```js
assert.match(code, /const\s+fifthResult\s*=/);
```

`fifthResult` mora imati numeričku vrednost `20`.

```js
assert.strictEqual(fifthResult, 20);
```

`fifthResult` Ne bi trebalo da uključuje množenje sekvence karaktera (primer "10" * 2).

```js
const match = code.match(/const\s+fifthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /["'`].*["'`]\s*\*/);
```

Trebalo bi da imaš parametar nazvan `sixthResult`.

```js
assert.match(code, /const\s+sixthResult\s*=/);
```

`sixthResult` mora imati numeričku vrednost `22`.

```js
assert.strictEqual(sixthResult, 22);
```

`sixthResult` ne bi trebalo da uključuje `null` u svoj matematički izraz.

```js
const match = code.match(/const\s+sixthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\bnull\b/);
```

# --seed--

## --seed-contents--

```js
const firstResult = 5 + "10";
console.log(`5 + 10 = ${firstResult}`);

const secondResult = "Eight" - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = true + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = false + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = "10" * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = null + 22;
console.log(`0 + 22 = ${sixthResult}`);
```

# --solutions--

```js
const firstResult = 5 + 10;
console.log(`5 + 10 = ${firstResult}`);

const secondResult = 8 - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = 1 + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = 0 + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = 10 * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = 0 + 22;
console.log(`0 + 22 = ${sixthResult}`);
```
