---
id: 68f890b73bbfc9efa7379e7b
title: Dijagnostikovanje grešaka tip konverzije u bug-ovoj aplikaciji
challengeType: 26
dashedName: debug-type-coercion-errors-in-a-buggy-app
---

# --description--

Upravo ste pristupili lokalnoj radnji za web razvoj, i vaš prvi zadatak je da očistite neki neispravan kod koji su ostavili prethodni programeri.

Pokušavali su jednostavne aritmetičke operacije, ali nešto je pošlo naopako, rezultati nemaju smisla. Vaš zadatak je da pregledate, otklonite greške (debug) i ispravite kod tako da radi kako treba i da je lakši za čitanje.

**Cilj:** Ispunite korisničke priče ispod i učinite da svi testovi prođu kako biste završili laboratoriju. 

**Korisničke priče:**

1. Trebalo bi da imate varijablu pod nazivom `firstResult` koja ispravno sabira brojeve `5` i `10` kako bi proizvela vrednost `15`.
2. Trebalo bi da imate varijablu pod nazivom `secondResult` koja ispravno oduzima `5` od `8` kako bi proizvela vrednost `3`.
3. Trebalo bi da imate varijablu pod nazivom `thirdResult` koja proizvede vrednost `6` sabiranjem dva broja. Zamenite boolean koji se trenutno koristi u izrazu brojem.
4. Trebalo bi da imate varijablu pod nazivom `fourthResult` koja proizvede vrednost `8` sabiranjem dva broja. Zamenite boolean koji se trenutno koristi u izrazu brojem.
5. Trebalo bi da imate varijablu pod nazivom `fifthResult` koja ispravno množi dva broja kako bi dobila proizvod `20`. Zamenite string koji se trenutno koristi u izrazu brojem.
6. Trebalo bi da imate varijablu pod nazivom `sixthResult` koja ispravno sabira dva broja kako bi proizvela vrednost `22`. Zamenite `null` koji se trenutno koristi u izrazu brojem.

# --hints--

You should have a variable called `firstResult`.

```js
assert.match(code, /const\s+firstResult\s*=/);
```

`firstResult` should have the numeric value `15`.

```js
assert.strictEqual(firstResult, 15);
```

You should have a variable called `secondResult`.

```js
assert.match(code, /const\s+secondResult\s*=/);
```

`secondResult` should have the numeric value `3`.

```js
assert.strictEqual(secondResult, 3);
```

You should have a variable called `thirdResult`.

```js
assert.match(code, /const\s+thirdResult\s*=/);
```

`thirdResult` should have the numeric value `6`.

```js
assert.strictEqual(thirdResult, 6);
```

`thirdResult` should not include a boolean value (`true` or `false`).

```js
const match = code.match(/const\s+thirdResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

You should have a variable called `fourthResult`.

```js
assert.match(code, /const\s+fourthResult\s*=/);
```

`fourthResult` should have the numeric value `8`.

```js
assert.strictEqual(fourthResult, 8);
```

`fourthResult` should not include a boolean value (`true` or `false`).

```js
const match = code.match(/const\s+fourthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

You should have a variable called `fifthResult`.

```js
assert.match(code, /const\s+fifthResult\s*=/);
```

`fifthResult` should have the numeric value `20`.

```js
assert.strictEqual(fifthResult, 20);
```

`fifthResult` should not include a string multiplication (e.g. "10" * 2).

```js
const match = code.match(/const\s+fifthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /["'`].*["'`]\s*\*/);
```

You should have a variable called `sixthResult`.

```js
assert.match(code, /const\s+sixthResult\s*=/);
```

`sixthResult` should have the numeric value `22`.

```js
assert.strictEqual(sixthResult, 22);
```

`sixthResult` should not include `null` in its expression.

```js
const match = code.match(/const\s+sixthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\bnull\b/);
```

# --seed--

## --seed-contents--

```js
const firstResult = 5 + "10";
console.log(`5 + 10 = ${firstResult}`);

const secondResult = "Eight" - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = true + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = false + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = "10" * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = null + 22;
console.log(`0 + 22 = ${sixthResult}`);
```

# --solutions--

```js
const firstResult = 5 + 10;
console.log(`5 + 10 = ${firstResult}`);

const secondResult = 8 - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = 1 + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = 0 + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = 10 * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = 0 + 22;
console.log(`0 + 22 = ${sixthResult}`);
```
