---
id: 5dc174fcf86c76b9248c6eb2
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovoj radionici, nastavićete da radite sa osnovnim HTML elementima kao što su naslovi, pasusi i liste, gradeći aplikaciju za fotografije mačaka.

Počnite radionicu dodavanjem elementa `h1` sa tekstom `CatPhotoApp`.

# --hints--

The text `CatPhotoApp` should be present in the code. You may want to check your spelling.

```js
assert.match(code, /catphotoapp/i);
```

Your `h1` element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Your `h1` tags should be in lowercase. By convention, all HTML tags are written in lowercase.

```js
assert.notMatch(code, /<\/?H1>/);
```

Your `h1` element should have a closing tag. Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Your `h1` element's text should be `CatPhotoApp`. You have either omitted the text, have a typo, or it is not between the `h1` element's opening and closing tags.

```js
assert.equal(document.querySelector('h1')?.innerText?.trim().toLowerCase(), 'catphotoapp');
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
