---
id: 5dc23f9bf86c76b9248c6eba
title: Korak 7
challengeType: 0
dashedName: step-7
---

# --description--

Možete dodati slike na svoj sajt koristeći element `img`. Elementi `img` imaju otvarajući tag bez zatvarajućeg taga. Element bez zatvarajućeg taga poznat je kao <dfn>void element</dfn>.

Dodajte element `img` ispod elementa `p`. U ovom trenutku, nijedna slika neće biti prikazana u pretraživaču.

# --hints--

Your `img` element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector('img'));
```

Your `img` element should not have a closing tag `</img>`.

```js
assert.notMatch(code, /<\/img\>/);
```

You should only have one `img` element. Remove any extras.

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

Your `img` element should be below the `p` element. You have them in the wrong order.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
