---
id: 5dc2385ff86c76b9248c6eb7
title: Korak 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 ima neke elemente koji identifikuju različite sadržajne oblasti. Ovi elementi čine vaš HTML lakšim za čitanje i pomažu u optimizaciji pretraživača (SEO) i pristupačnosti.

Element ``main`` se koristi za predstavljanje glavnog sadržaja tela HTML dokumenta. Sadržaj unutar elementa ``main`` treba da bude jedinstven za dokument i ne bi trebalo da se ponavlja u drugim delovima dokumenta.

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Identifikujte glavni deo ove stranice dodavanjem otvarajućeg taga ``<main>`` pre elementa ``h1``, i zatvarajućeg taga ``</main>`` nakon elementa ``p``.

# --hints--

Your `main` element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector('main'));
```

Your `main` tags should be in lowercase. By convention, all HTML tags are written in lowercase.

```js
const mainTagMatches = code.matchAll(/<\/?main>/gi);
for (const match of mainTagMatches) {
  const tag = match[0]
  assert.strictEqual(
    tag,
    tag.toLowerCase()
  );
}
```

Your `main` element should have a closing tag. Closing tags have a `/` just after the `<` character.

```js
assert.match(code, /<\/main\>/);
```

Your `main` element's opening tag should be below the `body` element's opening tag. You have them in the wrong order.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

Your `main` element's opening tag should be above the `h1` element. You have them in the wrong order.

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('MAIN'), collection.indexOf('H1'));
```

Your `main` element's closing tag should be below the `p` element. You have them in the wrong order.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode?.contains(pNode));
assert.match(pNode?.textContent.toLowerCase(), /everyone loves cute cats online/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```
