---
id: 5dc24165f86c76b9248c6ebc
title: Korak 9
challengeType: 0
dashedName: step-9
---

# --description--

Svi `img` elementi bi trebalo da imaju atribut `alt`. Tekst atributa `alt` koristi se za čitače ekrana radi poboljšanja pristupačnosti i prikazuje se ako slika ne uspe da se učita.

Evo primera `img` elementa sa atributom `alt`:

```html
<img src="cat.jpg" alt="A cat">
```

Unutar `img` elementa, dodajte atribut `alt` sa ovim tekstom:

`A cute orange cat lying on its back`

# --hints--

Your code should have an `img` element. You removed the `img` element from an earlier step.

```js
assert.exists(document.querySelector('img'));
```

Your `img` element does not have an `alt` attribute. Check that there is a space after the opening tag's name and/or there are spaces before all attribute names.

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

Your `img` element's `alt` attribute value is set to something other than `A cute orange cat lying on its back`. Make sure the `alt` attribute's value is surrounded with quotation marks.

```js
const altText = document
  .querySelector('img')
  ?.alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert.match(altText, /A cute orange cat lying on its back\.?$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```
