---
id: 5dc23991f86c76b9248c6eb8
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

U prethodnom koraku, stavili ste elemente `h1`, `h2`, komentar i `p` unutar elementa `main`. Ovo se zove *gnježdenje*. Gnježđeni elementi treba da budu postavljeni dva razmaka dalje udesno od elementa u koji su gnježđeni. Ovaj razmak se zove uvlačenje i koristi se kako bi bilo lakše čitati HTML.

Evo primera gnježđenja i uvlačenja:

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Element `h1`, element `h2` i komentar su uvlačeni za dva razmaka više od elementa `main` u kodu ispod. Koristite razmaknicu na tastaturi da dodate još dva razmaka ispred elementa `p` kako bi i on bio pravilno uvlačen.

# --hints--

Your code should have an `h2` element with text `Cat Photos`. You may have accidentally deleted it, it is missing both opening and closing tags, or the text has changed.

```js
assert.exists(document.querySelector('h2'));
assert.match(code, /<\/h2\>/);
assert.equal(document.querySelector('h2')?.innerText?.trim().replace(/\s+/g, ' ').toLowerCase(),'cat photos')
```

You should not add the `ul` or `li` elements from the example.

```js
assert.notExists(document.querySelector('ul'));
assert.notExists(document.querySelector('li'));
```

You should not change the indentation on the line with your `h2` element. Its opening tag should start 6 spaces over from the start of the line. You can restart the step to restore the original indentation.

```js
assert.match(code.toLowerCase(), /<\/h1\>\s*\n\s{6}<h2>/);
```

Your code should have a comment. You removed the comment from an earlier step.

```js
assert.match(code, /<!--.*-->/);
```

The comment's text should be `TODO: Add link to cat photos`. Do not change the text or spacing of the comment.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i);
```

You should not change the indentation on the line with your comment element. Its opening tag should start 6 spaces over from the start of the line. You can restart the step to restore the original indentation.

```js
assert.match(
  code
    .toLowerCase(),
    /<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/
);
```

Your code should have a `p` element. You have removed the `p` element from an earlier step.

```js
assert.exists(document.querySelector('p'));
```

The text of the `p` element should be `Everyone loves cute cats online!` Do not change the text, spacing, or punctuation of the `p` element.

```js
assert.match(
  document
    .querySelector('p')
    ?.innerText?.trim().toLowerCase(),
    /^everyone\s+loves\s+cute\s+cats\s+online!$/
);
```

The opening `p` tag should have indentation that matches your `h2` and comment elements. Its opening tag should start 6 spaces over from the start of the line.

```js
assert.match(code.toLowerCase(), /-->\s*\n\s{6}<p>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
    <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
    </main>
  </body>
</html>
```
