---
id: 5dfa2407b521be39a3de7be1
title: Korak 15
challengeType: 0
dashedName: step-15
---

# --description--

Da biste otvorili linkove u novoj karti, možete koristiti atribut `target` na elementi jangkoru (anchor) (`a`).

Atribut `target` određuje gde će se otvoriti povezani dokument. `target="_blank"` otvara povezani dokument u novoj karti ili prozoru.

Evo osnovne sintakse za element `a` sa atributom `target`:

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

Dodajte atribut `target` sa vrednošću `_blank` u otvarajući tag jangkoru (anchor) (`a`) elementa `cat photos`, kako bi se link otvorio u novoj karti.

# --hints--

Your `p` element should have a nested anchor (`a`) element with the text `cat photos`. You may have deleted it or have a typo.

```js
const anchor = document.querySelectorAll('p > a')[1];
assert.equal(
    anchor?.innerText?.toLowerCase().trim().replace(/\s+/g, ' '), 'cat photos'
);
```

Your anchor (`a`) element does not have a `target` attribute. Check that there is a space after the opening tag's name and/or there are spaces before all attribute names.

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

The value of the `target` attribute should be `_blank`. You have either omitted the value or have a typo. Remember that attribute values should be surrounded with quotation marks.

```js
assert.equal(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
