---
id: 5ef9b03c81a63668521804ea
title: Korak 38
challengeType: 0
dashedName: step-38
---

# --description--

Primećujte da je sve što ste do sada dodali na stranicu unutar elementa `body`. Svi elementi sadržaja stranice koji bi trebalo da se prikažu na stranici idu unutar elementa `body`. Međutim, druge važne informacije idu unutar elementa `head`.

Element `head` se koristi za sadržaj metapodataka o dokumentu, kao što su njegov naslov, veze ka listovima stilova i skripti. Metapodaci su informacije o stranici koje nisu direktno prikazane na stranici.

Dodajte element `head` iznad elementa `body`.

# --hints--

You have either deleted the `body` element or it is missing an opening tag or closing tag.

```js
assert.exists(document.querySelector('body'));
assert.match(code, /<\/body>/);
```

Your `head` element should have an opening tag. Opening tags have the following syntax: `<elementName>`.

```js
assert.match(code, /\<head\>/);
```

Your `head` element should have a closing tag. Closing tags have a `/` just after the `<` character.

```js
assert.match(code, /\<\/head\>/);
```

Your `head` element should be above the opening `body` element tag. You have put it somewhere else.

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<\/head\>\<body\>/);
```

# --seed--

## --seed-contents--

```html
<html>
--fcc-editable-region--
  
--fcc-editable-region--
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
