---
id: 68e4276483b0c71658aab1d2
title: Korak 13
challengeType: 1
dashedName: step-13
---

# --description--

Biće trenutaka kada ćete morati da koristite više metoda za kreiranje novog niza karaktera. Tokom sledećih nekoliko koraka, koristićete neke od metoda koje ste do sada naučili kako biste kreirali novi niz karaktera koji koristi konvenciju imenovanja camelCase.

Počnite kreiranjem varijable nazvane `lowercaseWord` i dodelite joj niz karaktera `"camelcase"`. 

Ispod varijable `lowercaseWord`, kreirajte još jednu varijablu nazvanu `camelCasedVersion` i dodelite joj prazan niz karaktera za sada.

Ispod varijable `camelCasedVersion`, dodajte `console.log()` sa nizom karaktera `"Camel cased version:"` praćen drugim `console.log()` za varijablu `camelCasedVersion`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should have a variable called `lowercaseWord`.

```js
assert.exists(lowercaseWord);
```

Your `lowercaseWord` should be a string.

```js
assert.isString(lowercaseWord);
```

You should assign the string `"camelcase"` to your `lowercaseWord` variable.

```js
assert.equal(lowercaseWord, "camelcase");
```

You should have a `camelCasedVersion` variable.

```js
assert.exists(camelCasedVersion);
```

Your `camelCasedVersion` variable should be a string.

```js
assert.isString(camelCasedVersion);
```

You should assign an empty string to your `camelCasedVersion` variable.

```js
assert.isEmpty(camelCasedVersion);
```

You should log the string `"Camel cased version:"` to the console.

```js
assert.equal(getLogs()[12], "Camel cased version:");
```

You should log the `camelCasedVersion` variable to the console.

```js
assert.equal(getLogs()[13], camelCasedVersion);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
