---
id: 68e4276483b0c71658aab1ca
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, veštaćete se sa radom sa raznim metodama za stringove koje se koriste za formatiranje stringova.

Počnite kreiranjem varijable nazvane `userInput` i dodelite joj sledeću vrednost: `"   Hello World!   "`.

# --hints--

You should have a `userInput` variable.

```js
assert.exists(userInput);
```

Your `userInput` variable should be a string.

```js
assert.isString(userInput);
```

You should assign `"   Hello World!   "` to your `userInput` variable. Make sure there are three spaces before the `H` and three spaces after the `!`.

```js
assert.equal(userInput, "   Hello World!   ");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
