---
id: 68e4276483b0c71658aab1d0
title: Korak 10
challengeType: 1
dashedName: step-10
---

# --description--

Biće trenutaka kada ćete morati da kreirate novi niz sa drugačijim velikim i malim slovima.

Da biste kreirali novi niz čije su sve slova konvertovane u VELIKA SLOVA, možete koristiti metod `toUpperCase()`.

Evo primera:

```js
const sentence = "I love to code!";
sentence.toUpperCase(); // "I LOVE TO CODE!"
```

Kreirajte varijablu nazvanu `upperCaseInput` i dodelite joj rezultat korišćenjem metode `toUpperCase()` na varijabli `cleanedInput`.

# --hints--

You should have a `upperCaseInput` variable.

```js
assert.exists(upperCaseInput);
```

Your `upperCaseInput` variable should be a string.

```js
assert.isString(upperCaseInput);
```

You should assign `cleanedInput.toUpperCase()` to your `upperCaseInput` variable.

```js
assert.equal(upperCaseInput, cleanedInput.toUpperCase());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

--fcc-editable-region--

--fcc-editable-region--
```
