---
id: 68e4276483b0c71658aab1cc
title: Korak 4
challengeType: 1
dashedName: step-4
---

# --description--

Vrednost vašeg `userInput` ima neke dodatne razmake na oba kraja niza karaktera. Da biste kreirali novi niz karaktera bez tog belog razmaka, možete koristiti metod `trim()`.  

Evo primera:

```js
const sentence = "  Learning about the trim method.  ";
sentence.trim(); // "Learning about the trim method."
```

Kreirajte varijablu pod imenom `cleanedInput` i dodelite joj rezultat korišćenjem metode `trim()` na varijabli `userInput`.

# --hints--

You should have a `cleanedInput` variable.

```js
assert.exists(cleanedInput);
```

Your `cleanedInput` variable should be a string.

```js
assert.isString(cleanedInput);
```

You should assign `userInput.trim()` to your `cleanedInput` variable.

```js
assert.equal(cleanedInput, userInput.trim());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);
--fcc-editable-region--

--fcc-editable-region--
```
