---
id: 68e4276483b0c71658aab1ce
title: Korak 6
challengeType: 1
dashedName: step-6
---

# --description--

Postojeće će trenutke kada ćete želeti da uklonite samo razmake sa početka niza karaktera (stringa). Tu je metoda `trimStart()` korisna.

Evo primera:

```js
const sentence = "  I love learning.  ";
sentence.trimStart(); // "I love learning.  "
```

Zapamtite da su nizovi karaktera (stringovi) nemenjivi (immutable). Ova metoda će kreirati novi string sa rezultatom. Ne modifikuje originalni.

Kreirajte varijablu `trimmedStart` i dodelite joj rezultat korišćenjem metode `trimStart()` na varijabli `userInput`.

# --hints--

You should have a `trimmedStart` variable.

```js
assert.exists(trimmedStart);
```

Your `trimmedStart` variable should be a string.

```js
assert.isString(trimmedStart);
```

You should assign `userInput.trimStart()` to your `trimmedStart` variable.

```js
assert.equal(trimmedStart, userInput.trimStart());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);
--fcc-editable-region--

--fcc-editable-region--
```
