---
id: 68e4276483b0c71658aab1d1
title: Korak 12
challengeType: 1
dashedName: step-12
---

# --description--

Ako je potrebno kreirati novi niz znakova gde su sva slova u malim slovima, možete koristiti metodu `toLowerCase()`.

Počnite kreiranjem varijable nazvane `lowerCaseInput` i dodelite joj rezultat korišćenjem metode `toLowerCase()` na varijabli `cleanedInput`.

Zatim logujte niz znakova `"Result of using the toLowerCase() method:"`.

Konačno, dodajte još jedan `console.log()` za varijablu `lowerCaseInput`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should have a `lowerCaseInput` variable.

```js
assert.exists(lowerCaseInput);
```

Your `lowerCaseInput` variable should be a string.

```js
assert.isString(lowerCaseInput);
```

You should assign `cleanedInput.toLowerCase()` to your `lowerCaseInput` variable.

```js
assert.equal(lowerCaseInput, cleanedInput.toLowerCase());
```

You should log the string `"Result of using the toLowerCase() method:"` to the console.

```js
assert.equal(getLogs()[10], "Result of using the toLowerCase() method:");
```

You should log the `lowerCaseInput` variable to the console.

```js
assert.equal(getLogs()[11], lowerCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);


const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
