---
id: 68e4276483b0c71658aab1cd
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Sada je vreme da vidimo rezultat.

Počnite dodavanjem `console.log()` sa stringom `"Result of trimming whitespace from both ends:"`.

Zatim ispod toga, ispod `console.log()`, dodajte još jedan `console.log()` sa varijablom `cleanedInput`. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should log the string `"Result of trimming whitespace from both ends:"` to the console.

```js
assert.equal(getLogs()[2], "Result of trimming whitespace from both ends:");
```

You should log the `cleanedInput` variable to the console.

```js
assert.equal(getLogs()[3], cleanedInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
--fcc-editable-region--

--fcc-editable-region--
```
