---
id: 68f1179eaea3bb59ba434386
title: Korak 8
challengeType: 1
dashedName: step-8
---

# --description--

Da biste uklonili razmake na kraju niza karaktera, možete koristiti metod `trimEnd()`.

Evo primera:

```js
const sentence = "  I love learning.  ";
sentence.trimEnd(); // "  I love learning."
```

Kreirajte varijablu nazvanu `trimmedEnd` i dodelite joj rezultat korišćenjem metode `trimEnd()` na varijabli `userInput`.

# --hints--

You should have a `trimmedEnd` variable.

```js
assert.exists(trimmedEnd);
```

Your `trimmedEnd` variable should be a string.

```js
assert.isString(trimmedEnd);
```

You should assign `userInput.trimEnd()` to your `trimmedEnd` variable.

```js
assert.equal(trimmedEnd, userInput.trimEnd());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--

--fcc-editable-region--
```
