---
id: 68f3f3d65d51d80ac0b412fc
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Sada je vreme da vidimo rezultate uklanjanja razmaka sa početka niza karaktera.

Ispod vaše varijable `trimmedStart`, dodajte `console.log()` sa nizom karaktera (stringom) `"After using the trimStart() method, leading spaces removed:"`. 

Zatim ispod tog `console.log()`, dodajte još jedno `console.log()` sa varijablom `trimmedStart`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should log the string `"After using the trimStart() method, leading spaces removed:"` to the console.

```js
assert.equal(getLogs()[4], "After using the trimStart() method, leading spaces removed:");
```

You should log the `trimmedStart` variable to the console.

```js
assert.equal(getLogs()[5], trimmedStart);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
--fcc-editable-region--

--fcc-editable-region--
```
