---
id: 68f40925d810ae14b2d5a169
title: Korak 16
challengeType: 1
dashedName: step-16
---

# --description--

Trenutno vaﾅ｡a promenljiva `camelCasedVersion` sadrﾅｾi vrednost `camelC`. Ali ﾅｾeljeni rezultat je `camelCase`.

Da biste dobili preostale `ase` slova, moﾅｾete koristiti metod `slice()` sa negativnom indeksnom vrednoﾅ｡ﾄ㎡. Zapamtite da koriﾅ｡ﾄ㌃nje negativnih indeksnih vrednosti govori metodi `slice()` da ekstrahu od kraja.

Evo primera za podsetnik:

```js
const str = "abcd";
str.slice(-2); // cd
```

Aﾅｾurirajte dodelu `camelCasedVersion` tako da takoﾄ粗 konkateniﾅ｡e `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` sa rezultatom koriﾅ｡ﾄ㌃nja metode `slice()` na `lowercaseWord`. Proﾄ訴te broj `-3` za metodu `slice()`.

I sa tom poslednjom promenom, vaﾅ｡ radionica je gotova!

# --hints--

You should concatenate `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` with `lowercaseWord.slice(-3)`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase() + lowercaseWord.slice(-3));
```

You should NOT assign the literal string `"camelCase"` to the `camelCasedVersion`. Please re read through the instructions for the correct approach.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelCase(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

# --solutions--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";

const camelCasedVersion =
  lowercaseWord.slice(0, 5) +
  lowercaseWord[5].toUpperCase() +
  lowercaseWord.slice(-3);

console.log("Camel cased version:");
console.log(camelCasedVersion);
```
