---
id: 68f405c8c9df44132c8fe1bb
title: Korak 15
challengeType: 1
dashedName: step-15
---

# --description--

Kao što ste se setili iz ranijih lekcija, konvencija imenovanja camelCase počinje malim slovom za prvu reč, a svaka naknadna reč počinje velikim slovom praćen malim slovima.

Evo nekoliko primera za podsetnik: 

```md
freeCodeCamp
camelCase
learningIsFun
```

Druga reč u promenljivoj `lowercaseWord` je `"case"`. Da biste pristupili `c` u toj reči, možete koristiti `lowercaseWord[5]`. 

Ažurirajte dodelu `camelCasedVersion` da koristite operator `+` za konkatenaciju `lowercaseWord.slice(0, 5)` sa rezultatom korišćenjem ispravne metode za pretvaranje stringova u velika slova na `lowercaseWord[5]`.

Sada bi trebalo da vidite `camelC` u konzoli.

# --hints--

You should concatenate `lowercaseWord.slice(0, 5)` with `lowercaseWord[5].toUpperCase()`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase());
```

You should NOT assign the literal string `"camelC"` to the `camelCasedVersion`. Please re read through the instructions for the correct approach.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelC(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5)
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

