---
id: 61329d52e5010e08d9b9d66b
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Još jedan važan `meta` element za pristupačnost i SEO je definicija `description`. Vrednost atributa `content` koristi se od strane pretraživača radi pružanja opisa vaše stranice.

Dodajte element `meta` sa atributom `name` postavljenim na `description`, i dodelite mu koristan atribut `content`.

# --hints--

You should add a new `meta` element to the `head`.

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

You should give the `meta` a `name` attribute of `description`.

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

You should give the `meta` a `content` attribute.

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

The `content` attribute value should not be more than 165 characters. _This is Google's maximum description length._

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
