---
id: 61329b210dac0b08047fd6ab
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Nastavljajući sa elementima `meta`, definicija `viewport` govori pretraživaču kako da prikaže stranicu. Uključivanje poboljšava vizuelnu pristupačnost na mobilnim uređajima i unapređuje _SEO_ (optimizaciju za pretraživače).

Dodajte definiciju `viewport` sa atributom `content`, koji detaljno opisuje `width` i `initial-scale` stranice.

# --hints--

You should create another `meta` element in the `head`.

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

You should give the `meta` a `name` attribute of `viewport`.

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

You should give the `meta` a `content` attribute of `width=device-width, initial-scale=1`.

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
