---
id: 66bb3e20d3dc5b6d0a21f5dd
title: Izgradite real-time brojač
challengeType: 25
dashedName: build-a-real-time-counter
demoType: onClick
---

# --description--

U ovom laboratorijskom zadatku, kreiraćete brojač karaktera u realnom vremenu koristeći JavaScript. Brojač će prikazivati broj karaktera unetih u element `textarea`. Brojač će se ažurirati u realnom vremenu dok korisnik kuca u `textarea`.

**Cilj:** Ispunite priče korisnika ispod i prođite sve testove kako biste završili laboratorijski zadatak.

**Priče korisnika:**

1. Trebalo bi da imate element `textarea` sa atributom `id` od `text-input`.
2. Trebalo bi da postoji element `p` sa atributom `id` od `char-count`. Inicijalno treba da sadrži tekst `Character Count: 0/50`. Ovaj tekst naslovni deo treba zameniti trenutnim brojem karaktera u elementu `textarea`.
3. Kada element `#text-input` sadrži tekst `hello coder`, element `#char-count` bi trebalo da sadrži tekst `"Character Count: 11/50"`
4. Kada element `#text-input` sadrži tekst `i am a programmer`, element `#char-count` bi trebalo da sadrži tekst `"Character Count: 17/50"`
5. Kada element `#text-input` sadrži tekst `hello world`, element `#char-count` bi trebalo da sadrži tekst `"Character Count: 11/50"`
6. Kada element `#text-input` sadrži tekst `I am learning a new language and it is called c++.`, element `#char-count` bi trebalo da sadrži tekst `"Character Count: 50/50"`. **NAPOMENA**: Iako bi atribut `maxlength` postigao ispravnu funkcionalnost, ovaj laboratorijski zadatak zahteva rešenje samo sa JavaScript-om.
7. Korisnik ne bi trebalo da unese više od `50` karaktera u element `textarea`. Kada broj karaktera dostigne `50`, bilo kakav dodatni unos treba automatski obrezati, a tekst `Character Count: 50/50` bi trebalo da se prikaže u `red`.

**Napomena:** Obavezno povežite svoj stilski list i JavaScript fajl u vašem HTML-u.

# --hints--

You should have a `textarea` element with the `id` of `text-input`. 

```js
const textInput = document.querySelector('textarea');
assert.equal(textInput.id, 'text-input');
```

You should have a `p` element with the `id` of `char-count`.

```js
assert.exists("p#char-count");
```

The `#char-count` element should initially contain the text `Character Count: 0/50`.

```js
const charCount = document.querySelector('p#char-count')?.innerText;
assert.equal(charCount, 'Character Count: 0/50');
```

When the `#text-input` element contains the text `hello coder` the `#char-count` element should contain the text `"Character Count: 11/50"`

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "hello coder";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(charCount.innerText.trim(), 'Character Count: 11/50');
```

When the character count is `50`, the text should be displayed in `red`.

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "I am learning a new language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
const computedStyle = window.getComputedStyle(charCount);
assert.oneOf(computedStyle.color, [
  "red",
  "rgb(255, 0, 0)",
  "rgb(100%, 0%, 0%)",
  "rgba(255, 0, 0, 1)",
  "rgba(100%, 0%, 0%, 1)",
  "hsl(0, 100%, 50%)",
  "hsla(0, 100%, 50%, 1)"
]);
```

If character count is greater than or equal to `50`, the extra input should be trimmed.

```js
const textInput = document.getElementById('text-input');
textInput.value = "I am learning a new programming language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(textInput.value.length, 50);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real Time Counter</title>
</head>

<body>

</body>


</html>
```

```css

```

```js
    
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real-time Character Counter</title>
    <link rel="stylesheet" href="styles.css">

<body>
    <div class="container">
        <h1>Real-Time Character Counter</h1>
        <textarea id="text-input" placeholder="Type something..."></textarea>
        <p id="char-count">Character Count: 0/50</p>
    </div>

<script src="script.js"></script>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    font-family: Arial, sans-serif;
    background-color: #c0c1ec;
    margin: 0;
}

h1{
    background-color: #f2f48b;
    padding-left: 10px;
    padding-right: 10px;
}
.container {
    text-align: center;
}

textarea {
    width: 300px;
    height: 100px;
    font-size: 18px;
    padding: 10px;
    border: 1px solid #ccc;
    border-radius: 5px;
    resize: none;
    display: block;
    margin: 0 auto;
    border: 1px solid slategray;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

textarea:focus {
    outline: none;
    border-color: #f2f48b;
}

#char-count {
    margin-top: 10px;
    font-size: 18px;
    color: darkslategray;
    font-weight: bold;
}
```

```js
document.getElementById("text-input").addEventListener("input", function () {
  let textInput = document.getElementById("text-input");
  let charCount = textInput.value.length;

  if (charCount > 50) {
    textInput.value = textInput.value.substring(0, 50);
    charCount = 50;
  }

  let charCountDisplay = document.getElementById("char-count");
  charCountDisplay.innerHTML = `Character Count: ${charCount}/50`;

  if (charCount === 50) {
    charCountDisplay.style.color = "red";
  } 

});
```
