---
id: 685239f0cad43c7a909cbd98
title: Debugovanje stranice za usvajanje ljubimaca
challengeType: 25
dashedName: lab-debug-pet-adoption-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Sali, vlasnica prodavnice za usvajanje kućnih ljubimaca, napravila je svoju prvu veb stranicu, ali postoje neki problemi.

Vaš zadatak je da popravite sve greške kako bi Sali nastavila sa izgradnjom svoje stranice. 

**Cilj:** Ispunite priče korisnika ispod i prođite sve testove kako biste završili laboratoriju.

**Priče korisnika:**

1. Sali želi da koristi sliku mačaka, ali se ona ne prikazuje pravilno. Moraćete da popravite sledeće u elementu `img`:
   - Zamenite atribut `href` ispravnim atributom za izvor slike. 
   - Zamenite atribut `att` ispravnim atributom koji predstavlja kratak, opisni tekst za slike.
   - Uklonite zatvarajući tag `</img>` jer su elementi `img` prazni (void) elementi i nemaju zatvarajuće tagove.
2. Sali želi da koristi neke linkove kako bi usmerila korisnike na stranice za pse i mačke. Ali linkovi ne rade pravilno. Moraćete da popravite sledeće u elementima `a`:
   - Zamenite oba atributa `src` ispravnim atributima koji se koriste za navođenje URL adresa.

# --hints--

Your `img` element should have a `src` attribute instead of the `href` attribute.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("src"));
```

Your `img` element should have an `alt` attribute instead of the non-existent `att` attribute.

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("alt"));
```

Your `img` element should not have a `</img>` closing tag.

```js
assert.notMatch(code, /<\/img>/);
```

Your `a` element with the text `Visit cats page` needs to have an `href` attribute instead of a `src` attribute.

```js
const anchors = document.querySelectorAll("a");
const catLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit cats page");
assert.isDefined(catLink);
assert.isTrue(catLink.hasAttribute("href"));
```

Your `a` element with the text `Visit dogs page` needs to have an `href` attribute instead of a `src` attribute.

```js
const anchors = document.querySelectorAll("a");
const dogLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit dogs page");
assert.isDefined(dogLink);
assert.isTrue(dogLink.hasAttribute("href"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" att="Two tabby kittens sleeping together on a couch."></img>

<h2>Adopt a cat!</h2>
<a src="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a src="/dogs">Visit dogs page</a>
```

# --solutions--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">

<h2>Adopt a cat!</h2>
<a href="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a href="/dogs">Visit dogs page</a>
```
