---
id: 61fd67a656743144844941cb
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Čitači ekrana najavljuju HTML elemente na osnovu toka dokumenta. Na kraju ćemo želeti da bilansni list ima naslov "Balance Sheet" i podnaslov "AcmeWidgetCorp". Međutim, ovaj redosled nema smisla ako ga najavi čitač ekrana.

Dodajte vašem postojećem `span` atribut `class` podešen na `flex`, i dodajte unutar njega dva `span` elementa. Prvom dajte tekst `AcmeWidgetCorp`. Drugom dajte tekst `Balance Sheet`. Koristićete CSS da obrnete redosled teksta na stranici, ali će HTML redosled imati više smisla za čitač ekrana.

# --hints--

Your existing `span` element should have the `class` attribute set to `flex`.

```js
assert.isTrue(
  document.querySelector('h1')?.children?.[0]?.classList?.contains('flex')
);
```

Your existing `span` element should have two new `span` elements within it.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.localName,
  'span'
);
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.localName,
  'span'
);
```

Your new `span` elements should not have a `class` attribute.

```js
assert.isEmpty(document.querySelector('.flex')?.children?.[0]?.classList);
assert.isEmpty(document.querySelector('.flex')?.children?.[1]?.classList);
```

Your first new `span` element should have the text `AcmeWidgetCorp`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.textContent.trim(),
  'AcmeWidgetCorp'
);
```

Your second new `span` element should have the text `Balance Sheet`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.textContent.trim(),
  'Balance Sheet'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
