id: 5a858944d96184f06fd60d61
title: Kreiraj svoj prvi CSS Grid kontejner
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cqwREC4'
forumTopicId: 301129
dashedName: create-your-first-css-grid

# --description--

Promijeni bilo koji HTML element u grid kontejnerom postavljanjem njegove `display` osobine na `grid`. To vam omogućava korištenje svih ostalih svojstava povezanih sa CSS Grid-om.

**Napomena:** U CSS Gridu, roditeljski element naziva se `<dfn>kontejner</dfn>`, a njegovi potomci nazivaju se `<dfn>stavke</dfn>`.

# --instructions--

Promijeni prikaz `div`a sa klasom `container` u `grid`.

# --hints--

Klasa `container` treba imati svojstvo `display` s vrijednošću `grid`.

```js
assert.match(code, /.container\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    display: grid;
  }
</style>
```