id: 5a90374338fddaf9a66b5d3a
title: Poravnajte element horizontalno koristeći justify-self
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpKHq'
forumTopicId: 301122
dashedName: align-an-item-horizontally-using-justify-self
---

# --description--

U CSS Gridu, sadržaj svakog elementa nalazi se unutar kutije nazvane <dfn>ćelica</dfn>. Možete poravnati raspored sadržaja unutar svoje ćelije koristeći svojstvo `justify-self` na grid elementu. Po defaultu, ovo svojstvo ima vrednost `stretch`, što će naterati sadržaj da popuni celu širinu ćelije. Ovo CSS Grid svojstvo prima i druge vrednosti:

`start`: poravnava sadržaj na levoj strani ćelije,

`center`: poravnava sadržaj u sredini ćelije,

`end`: poravnava sadržaj na desnoj strani ćelije.

# --instructions--

Koristite svojstvo `justify-self` da poravnate element sa klasom `item2` u sredinu.

# --hints--

Klasa `item2` treba da ima svojstvo `justify-self` sa vrednošću `center`.

```js
assert.match(
  code,
  /.item2\s*?{[\s\S]*justify-self\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }

  .item2 {
    background: LightSalmon;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item2 {
    justify-self: center;
  }
</style>
```