id: 5a90376038fddaf9a66b5d3c
title: Poravnajte sve stavke horizontalno koristeći justify-items
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpECn'
forumTopicId: 301120
dashedName: align-all-items-horizontally-using-justify-items
---

# --description--

Ponekad želite da sve stavke u vašem CSS mrežnom sistemu dele isto poravnanje. Možete koristiti svojstva koja ste ranije naučili i poravnati ih po jednu, ili možete poravnati sve zajedno horizontalno koristeći `justify-items` na vašoj mrežnoj kontejneru. Ovo svojstvo može primiti sve vrednosti koje ste naučili u prethodnih dve izazova, razlika je u tome što će pomeriti **sve stavke** u našem mrežnom sistemu do željenog poravnanja.

# --instructions--

Koristite ovo svojstvo da centrirate sve stavke horizontalno.

# --hints--

Klasa `container` bi trebalo da ima svojstvo `justify-items` sa vrednošću `center`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*justify-items\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    justify-items: center;
  }
</style>
```