id: 5a9036ee38fddaf9a66b5d34
title: Koristite jedinice CSS Grid za menjanje veličine redova i kolona
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cvE8phd'
forumTopicId: 301134
dashedName: use-css-grid-units-to-change-the-size-of-columns-and-rows
---

# --description--

Možete koristiti fiksne i relativne jedinice kao što su `px` i `em` u CSS Grid za određivanje veličine redova i kolona. Takođe možete koristiti:

`fr`: čini kolonu ili red delom raspoloživog prostora,

`auto`: postavlja kolonu ili red na direktnu širinu ili visinu njihovog sadržaja,

`%`: postavlja kolonu ili red na procenat širine svog kontejnera.

Ovaj kod generiše rezultate na prethodnom prikazu:

```css
grid-template-columns: auto 50px 10% 2fr 1fr;
```

Ovaj deo stvara pet kolona. Prva kolona je široka kao njen sadržaj, druga kolona je 50px, treća kolona je 10% od svog kontejnera, a za poslednje dve kolone; preostali prostor se deli na tri dela, dva su namenjena četvrtoj koloni, a jedan petoj.

# --instructions--

Napravite grid sa tri kolone širine kao što sledi: 1fr, 100px i 2fr.

# --hints--

Vaša klasa `container` treba da ima svojstvo `grid-template-columns` koje postavlja три kolone sa sledećim širinama: `1fr`, `100px` i `2fr`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?1fr\s*?100px\s*?2fr\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: auto 50px 10% 2fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 50px 50px;
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: 1fr 100px 2fr;
  }
</style>
```