id: 5a9036d038fddaf9a66b5d32
title: Dodajte kolone sa grid-template-columns
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c7NzDHv'
forumTopicId: 301117
dashedName: add-columns-with-grid-template-columns
---

# --description--

Creiranje samo mrežnog elementa nije dovoljno. Morate definisati i strukturu mreže. Da biste dodali kolone u mrežu, koristite svojstvo `grid-template-columns` na grid kontejneru, kao što je prikazano ispod:

```css
.container {
  display: grid;
  grid-template-columns: 50px 50px;
}
```

Ovo će stvoriti dve kolone u vašoj mrežnoj strukturi, svaka široka 50px. Broj parametara navedenih za svojstvo `grid-template-columns` pokazuje broj kolona u mreži, a vrednost svakog parametra ukazuje širinu svake kolone.

# --instructions--

Dodajte tri kolone u grid kontejneru, svaka široka `100px`.

# --hints--

Klasa `container` bi trebalo da ima svojstvo `grid-template-columns` sa tri jedinice od `100px`.

```js
const templateColumns = new __helpers.CSSHelp(document).getStyle(
  '.container'
)?.gridTemplateColumns;
assert.strictEqual(templateColumns, '100px 100px 100px');
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: 100px 100px 100px;
  }
</style>
```