id: 5a94fe2669fb03452672e45e
title: Koristite grid-area bez kreiranja šablona područja
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c6N7VhK'
forumTopicId: 301135
dashedName: use-grid-area-without-creating-an-areas-template
---

# --description--

Svojstvo `grid-area` koje ste naučili u prethodnom izazovu može se koristiti i na drugi način. Ako vaša mreža nema referencijalni šablon područja, možete brzo kreirati područje za element koristeći sledeći kod:

```css
item1 {
  grid-area: 1/1/2/4;
}
```

Ovo koristi brojeve linija koje ste ranije naučili da definišete gde će se nalaziti područje tog elementa. Brojevi u gornjem primeru predstavljaju ove vrednosti:

```css
grid-area: horizontalna linija od / vertikalna linija od / horizontalna
  linija do / vertikalna linija do;
```

Dakle, element u primeru će zauzeti red između linija 1 i 2, i kolonu između linija 1 i 4.

# --instructions--

Koristeći svojstvo `grid-area`, postavite element sa klasom `item5` između treće i četvrte horizontalne linije i između prve i četvrte vertikalne linije.

# --hints--

Klasa `item5` treba da ima svojstvo `grid-area` kako bi popunila celo područje između treće i četvrte horizontalne linije, i između prve i četvrte vertikalne linije.

```js
assert.match(
  code,
  /.item5\s*?{[\s\S]*grid-area\s*?:\s*?3\s*?\/\s*?1\s*?\/\s*?4\s*?\/\s*?4\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }

  .item5 {
    background: PaleGreen;
    /* Only change code below this line */


    /* Only change code above this line */
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item5 {
    grid-area: 3/1/4/4;
  }
</style>
```