id: 5a94fe4469fb03452672e460
title: Smanjite veličinu elementa korišćenjem funkcije minmax
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cD97RTv'
forumTopicId: 301131
dashedName: limit-item-size-using-the-minmax-function
---

# --description--

Postoji još jedna funkcija koju možete koristiti zajedno sa `grid-template-columns` i `grid-template-rows`, a ona se zove `minmax`. Koristi se za prilagođavanje veličine elemenata kada se kontejner mreže menja veličina. Da biste to uradili, morate definisati raspon veličina koje vaš element može prihvatiti. Evo primera:

```css
grid-template-columns: 100px minmax(50px, 200px);
```

U gornom kodu, `grid-template-columns` je podešen da kreira dve kolone; prva ima širinu od 100px, a druga ima minimalnu širinu od 50px i maksimalnu širinu od 200px.

# --instructions--

Koristeći funkciju `minmax`, promenite `1fr` u funkciji `repeat` veličinom kolone koja ima minimalnu širinu od `90px` i maksimalnu širinu od `1fr`, a zatim promenite veličinu prikaznog panela da biste videli efekat.

# --hints--

Klasa `container` treba da ima svojstvo `grid-template-columns` postavljeno tako da ponavlja 3 kolone sa minimalnom širinom od `90px` i maksimalnom širinom od `1fr`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?3\s*?,\s*?minmax\s*?\(\s*?90px\s*?,\s*?1fr\s*?\)\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(3, 1fr);

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(3, minmax(90px, 1fr));
  }
</style>
```