id: 5a94fe3669fb03452672e45f
title: Smanjite ponavljanje korišćenjem funkcije `repeat`
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cQvqyHR'
forumTopicId: 301133
dashedName: reduce-repetition-using-the-repeat-function
---

# --description--

Kada ste koristili `grid-template-columns` i `grid-template-rows` za kreiranje mrežnog rasporeda (grid layout), uneli ste vrednost za svaki red ili stub koji ste stvorili.

Recimo da želite mrežni raspored sa 100 redova iste visine. Nije lako unositi 100 vrednosti jednu po jednu. Srećom, postoji mnogo bolji način - korišćenjem funkcije `repeat` koja pokazuje koliko puta želite da se ponovi red ili stub, praćen zarezom i vrednošću koju želite da ponovite.

Evo primera koji bi kreirao mrežni raspored od 100 redova, svaki sa visinom od 50px.

```css
grid-template-rows: repeat(100, 50px);
```

Takođe možete ponoviti više vrednosti koristeći funkciju `repeat` i umetnuti tu funkciju između drugih vrednosti prilikom kreiranja mrežnog rasporeda. To izgleda ovako:

```css
grid-template-columns: repeat(2, 1fr 50px) 20px;
```

Ovo se interpretira kao:

```css
grid-template-columns: 1fr 50px 1fr 50px 20px;
```

**Napomena:** `1fr 50px` se ponavlja dva puta, praćeno sa 20px.

# --instructions--

Koristite `repeat` da uklonite ponavljanje iz svojstva `grid-template-columns`.

# --hints--

Klasa `container` treba da ima svojstvo `grid-template-columns` postavljeno tako da ponovi 3 kolone širine `1fr`.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?3\s*?,\s*?1fr\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: 1fr 1fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(3, 1fr);
  }
</style>
```