id: 5a94fe8569fb03452672e464
title: Kreiranje mreža unutar mreža
challengeType: 0
forumTopicId: 301128
dashedName: create-grids-within-grids

# --description--

Pretvaranje elementa u mrežu utiče samo na ponašanje njegovog direktnog roditelja. Stoga, pretvaranjem direktnog roditelja u mrežu, dobijate mrežu unutar mreže.

Na primer, postavljanjem svojstava `display` i `grid-template-columns` za element sa klasom `item3`, kreirate mrežu unutar vaše mreže.

# --instructions--

Pretvorite element sa klasom `item3` u mrežu koja ima dva reda širine `auto` i `1fr`, koristeći `display` i `grid-template-columns`.

# --hints--

Klasa `item3` mora imati svojstvo `grid-template-columns` sa vrednostima `auto` i `1fr`.

```js
assert.match(
  code,
  /.item3\s*?{[\s\S]*grid-template-columns\s*?:\s*?auto\s*?1fr\s*?;[\s\S]*}/gi
);
```

Klasa `item3` mora imati svojstvo `display` sa vrednošću `grid`.

```js
assert.match(code, /.item3\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .container {
    font-size: 1.5em;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: auto 1fr;
    grid-template-rows: auto 1fr auto;
    grid-gap: 10px;
    grid-template-areas:
      'advert header'
      'advert content'
      'advert footer';
  }
  .item1 {
    background: LightSkyBlue;
    grid-area: header;
  }

  .item2 {
    background: LightSalmon;
    grid-area: advert;
  }

  .item3 {
    background: PaleTurquoise;
    grid-area: content;
    /* Only change code below this line */


    /* Only change code above this line */
  }

  .item4 {
    background: lightpink;
    grid-area: footer;
  }

  .itemOne {
    background: PaleGreen;
  }

  .itemTwo {
    background: BlanchedAlmond;
  }
</style>

<div class="container">
  <div class="item1">header</div>
  <div class="item2">advert</div>
  <div class="item3">
    <div class="itemOne">paragraph1</div>
    <div class="itemTwo">paragraph2</div>
  </div>
  <div class="item4">footer</div>
</div>
```

# --solutions--

```html
<style>
  .item3 {
    grid-template-columns: auto 1fr;
    display: grid;
  }
</style>
```