---
id: 699d1f8e51fcdc8230f779fb
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, izgradit ćete katalog sistem za biblioteku nasleđa. Parsiraćete sirove podatke o knjigama, pretraživati i grupisati unose, validirati kvalitet podataka i izvoziti rezultate u različite formate.

Svaka katalog kartica je sačuvana kao pojedinačni niz u formatu `"Title | Author | Year | Location"`, sa poljima razdvojenim znakom pipe (|) (`|`).

Koristićete niz ovih niza tokom radionice za izgradnju i testiranje svojih katalog funkcija. Kreirajte niz pod nazivom `rawCatalogCards` i dodajte mu sledeće dve linije:

```js
"From a Buick 8 | King, Stephen | 2002 | Shelf K7",
"The Shining | King, Stephen | 1977 | Shelf K1"
```

# --hints--

You should declare a variable named `rawCatalogCards`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+rawCatalogCards\s*=/);
```

`rawCatalogCards` should be an array.

```js
assert.isArray(rawCatalogCards);
```

The first element of `rawCatalogCards` should be `"From a Buick 8 | King, Stephen | 2002 | Shelf K7"`.

```js
assert.strictEqual(rawCatalogCards[0], "From a Buick 8 | King, Stephen | 2002 | Shelf K7");
```

The second element of `rawCatalogCards` should be `"The Shining | King, Stephen | 1977 | Shelf K1"`.

```js
assert.strictEqual(rawCatalogCards[1], "The Shining | King, Stephen | 1977 | Shelf K1");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
