id: bad87fee1348bd9aec908847
title: Podeli svoj Bootstrap red
challengeType: 0
forumTopicId: 18306
dashedName: split-your-bootstrap-row
---

# --description--

Sada kada imamo Bootstrap red, podelićemo ga na dve kolone da postavimo naše stavke.

Tengeneza vipengele viwili vya `div` ndani ya safu yako, vyote viwe na darasa la vitu la `col-xs-6`.

# --hints--

Dve elementa `div class="col-xs-6"` treba da budu postavljene unutar tvog elementa `div class="row"`.

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

Vipengele vyote vya `div` vinapaswa kuwa na lebo za kufunga.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g).length , code.match(/<div/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">


  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6"></div>
    <div class="col-xs-6"></div>
  </div>
</div>
```