id: bad87fee1348bd9aec908847
title: Podeli svoj Bootstrap redak
challengeType: 0
forumTopicId: 18306
dashedName: split-your-bootstrap-row
---

# --description--

Kada imamo Bootstrap redak, podelimo ga na dve kolone kako bismo postavili svoje elemente.

Kreiraj dva `div` elementa unutar svog reda, i da oba imaju klasu elemenata `col-xs-6`.

# --hints--

Dva `div` elementa sa klasom `"col-xs-6"` treba da budu postavljena unutar tvog `div` elementa sa klasom `"row"`.

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

Svi `div` elementi bi trebalo da imaju zatvarajuće oznake.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g).length , code.match(/<div/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">


  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6"></div>
    <div class="col-xs-6"></div>
  </div>
</div>
```