---
id: bad87fee1348bd9bec908846
title: Napravi Bootstrap red
challengeType: 0
forumTopicId: 16813
dashedName: create-a-bootstrap-row
---

# --description--

Sada cemo napraviti Bootstrap red za nase inline elemente.

Napravi `div` element ispod `h3` taga, sa klasom `row`.

# --hints--

Treba da dodas `div` element ispod svog `h3` elementa.

```js
assert.lengthOf(document.querySelectorAll('div'),2);
assert.lengthOf(document.querySelectorAll('div.row h3.text-primary'),0);
assert.lengthOf(document.querySelectorAll('div.row + h3.text-primary'),0)
assert.lengthOf(document.querySelectorAll('h3.text-primary + div.row'),1);
```

Tvoj `div` element treba da ima klasu `row`

```js
const newDiv = document.querySelectorAll('div')?.[1]; 
assert.isTrue(newDiv?.classList?.contains('row'));
```

Tvoj `row div` treba da bude smesten unutar `container-fluid div` elementa

```js
assert.lengthOf(document.querySelectorAll('div.container-fluid div.row'),1);
```

Tvoj `div` element treba da ima zatvarajuci tag.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g).length ,code.match(/<div/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>

</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row"></div>
</div>
```
