id: bad87fee1348bd9bec908846
title: Kreiranje redova Bootstrap-a
challengeType: 0
forumTopicId: 16813
dashedName: create-a-bootstrap-row
---

# --description--

Sada ćemo kreirati red Bootstrap-a za naše inline elemente.

Kreirajte element `div` ispod oznake `h3`, sa klasom elemenata `row`.

# --hints--

Trebalo bi da dodate element `div` ispod vašeg elementa `h3`.

```js
assert.lengthOf(document.querySelectorAll('div'),2);
assert.lengthOf(document.querySelectorAll('div.row h3.text-primary'),0);
assert.lengthOf(document.querySelectorAll('div.row + h3.text-primary'),0)
assert.lengthOf(document.querySelectorAll('h3.text-primary + div.row'),1);
```

Vaš element `div` treba da ima klasu elemenata `row`.

```js
const newDiv = document.querySelectorAll('div')?.[1]; 
assert.isTrue(newDiv?.classList?.contains('row'));
```

Vaš `row div` bi trebalo da bude unutar `container-fluid div`.

```js
assert.lengthOf(document.querySelectorAll('div.container-fluid div.row'),1);
```

Vaš element `div` treba da ima zatvarajući tag.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g).length ,code.match(/<div/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>

</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row"></div>
</div>
```