---
id: bad87fee1348bd9aec908846
title: Napravi Bootstrap naslov
challengeType: 0
forumTopicId: 16812
dashedName: create-a-bootstrap-headline
---

# --description--

Sada hajde da napravimo nesto od nule kako bismo uvezbali nase HTML, CSS i Bootstrap vestine.

Napravicemo jQuery playground, koji cemo uskoro koristiti u nasim jQuery izazovima.

Za pocetak, napravi `h3` element sa tekstom `jQuery Playground`.

Oboji svoj `h3` element Bootstrap klasom `text-primary`, a zatim ga centriraj Bootstrap klasom `text-center`.

# --hints--

Treba da dodas `h3` element na stranicu.

```js
assert.lengthOf(document.querySelectorAll('h3'),1);
```

Tvoj `h3` element treba da ima zatvarajuci tag.

```js
assert.match(code,/<\/h3>/g);
assert.match(code,/<h3/g);
assert.equal( code.match(/<\/h3>/g).length , code.match(/<h3/g).length);
```

Tvoj `h3` element treba da bude obojen pomocu klase `text-primary`

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-primary'));
```

Tvoj `h3` element treba da bude centriran pomocu klase `text-center`

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-center'));
```

Tvoj `h3` element treba da sadrzi tekst `jQuery Playground`.

```js
assert.match(document.querySelector('h3')?.textContent, /jquery(\s)+playground/gi);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```
