---
id: bad87fee1348bd9aec908746
title: Smesti nasu stranicu unutar Bootstrap container-fluid div-a
challengeType: 0
forumTopicId: 18198
dashedName: house-our-page-within-a-bootstrap-container-fluid-div
---

# --description--

Sada se postaraj da sav sadrzaj na stranici bude prilagodljiv mobilnim uredjajima.

Smesti svoj `h3` element unutar `div` elementa sa klasom `container-fluid`.

# --hints--

Tvoj `div` element treba da ima klasu `container-fluid`.

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container-fluid'));
```

Svaki `div` element treba da ima zatvarajuci tag.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);

assert.equal(code.match(/<\/div>/g).length ,code.match(/<div/g).length);
```

Tvoj `h3` element treba da bude smesten unutar `div` elementa.

```js
const divElement = document.querySelector('div');
const divChildren = divElement?.querySelectorAll(`:scope ${'h3'}`)
assert.lengthOf(divChildren,1);
```

# --seed--

## --seed-contents--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```

# --solutions--

```html
<div class="container-fluid">
    <h3 class="text-primary text-center">jQuery Playground</h3>
</div>
```
