id: bad87fee1348bd9aec908850
title: Primeni podrazumevani stil dugmeta Bootstrap-a
challengeType: 0
forumTopicId: 16657
dashedName: apply-the-default-bootstrap-button-style
---

# --description--

Bootstrap ima klasu za elemente dugmeta koja se zove `btn-default`.

Koristi klase elemenata `btn` i `btn-default` za svaki tvoj element `<button>`.

# --hints--

Trebaš koristiti klasu elemenata `btn` za svaki tvoj element `<button>`.

```js
assert.lengthOf(document.querySelectorAll('.btn'),6);
```

Trebaš koristiti klasu elemenata `btn-default` za svaki tvoj element `<button>`.

```js
assert.lengthOf(document.querySelectorAll('.btn-default'), 6);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">
        <button></button>
        <button></button>
        <button></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">
        <button></button>
        <button></button>
        <button></button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">
        <button class="btn btn-default"></button>
        <button class="btn btn-default"></button>
        <button class="btn btn-default"></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">
        <button class="btn btn-default"></button>
        <button class="btn btn-default"></button>
        <button class="btn btn-default"></button>
      </div>
    </div>
  </div>
</div>
```