id: bad87fee1348bd9aec908848
title: Kreiranje Bootstrap well-ova
challengeType: 0
forumTopicId: 16825
dashedName: create-bootstrap-wells
---

# --description--

Bootstrap ima klasu elemenata nazvanu `well` koja može stvoriti osećaj dubine za vaše redove.

Ubacite jedan element `div` sa klasom `well` unutar svakog vašeg elementa `div` sa klasom `col-xs-6`.

# --hints--

Treba da dodate element `div` sa klasom `well` unutar svakog vašeg elementa `div` sa klasom `col-xs-6`

```js
const wells = document.querySelectorAll('div.col-xs-6 > div.well');
assert.lengthOf( wells,2 ); 
```

Oba elementa `div` sa klasom `col-xs-6` treba da budu postavljena unutar vašeg elementa `div` sa klasom `row`.

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

Svi elementi `div` bi trebalo da imaju zatvarajuće oznake.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g)?.length , code.match(/<div/g)?.length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">

    </div>
    <div class="col-xs-6">

    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
  </div>
</div>
```