---
id: bad87fee1348bd9aec908848
title: Napravi Bootstrap well elemente
challengeType: 0
forumTopicId: 16825
dashedName: create-bootstrap-wells
---

# --description--

Bootstrap ima klasu `well` koja moze da stvori utisak dubine u tvojim kolonama.

Smesti po jedan `div` element sa klasom `well` unutar svakog `col-xs-6` `div` elementa.

# --hints--

Treba da dodas `div` element sa klasom `well` unutar svakog `div` elementa sa klasom `col-xs-6`

```js
const wells = document.querySelectorAll('div.col-xs-6 > div.well');
assert.lengthOf( wells,2 ); 
```

Oba `div` elementa sa klasom `col-xs-6` treba da budu smestena unutar tvog `div` elementa sa klasom `row`.

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

Svi `div` elementi treba da imaju zatvarajuce tagove.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g)?.length , code.match(/<div/g)?.length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">

    </div>
    <div class="col-xs-6">

    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
  </div>
</div>
```
