---
id: bad87fee1348bd9aec908847
title: Podeli svoj Bootstrap red
challengeType: 0
forumTopicId: 18306
dashedName: split-your-bootstrap-row
---

# --description--

Sada kada imamo Bootstrap red, podelimo ga na dve kolone kako bismo rasporedili nase elemente.

Napravi dva `div` elementa unutar svog reda, oba sa klasom `col-xs-6`.

# --hints--

Dva `div class="col-xs-6"` elementa treba da budu smestena unutar tvog `div class="row"` elementa.

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

Svi `div` elementi treba da imaju zatvarajuce tagove.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g).length , code.match(/<div/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">


  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6"></div>
    <div class="col-xs-6"></div>
  </div>
</div>
```
