id: bad87fee1348bd9aec908846
title: Kreiraj naslov novinskog članka sa Bootstrap-om
challengeType: 0
forumTopicId: 16812
dashedName: create-a-bootstrap-headline
---

# --description--

Sada ćemo nešto izgraditi od nule kako bismo navikli naše znanje o HTML, CSS i Bootstrap-u.

Izgradićemo jQuery Playground, koji ćemo uskoro koristiti u našim jQuery izazovima.

Počnite tako što ćete kreirati element `h3` sa tekstom `jQuery Playground`.

Boja vašeg `h3` elementa pomoću Bootstrap klase `text-primary`, i centrirajte ga koristeći Bootstrap klasu `text-center`.

# --hints--

Treba da dodate `h3` element na svoju stranicu.

```js
assert.lengthOf(document.querySelectorAll('h3'),1);
```

Vaš `h3` element treba da ima zatvarajući tag.

```js
assert.match(code,/<\/h3>/g);
assert.match(code,/<h3/g);
assert.equal( code.match(/<\/h3>/g).length , code.match(/<h3/g).length);
```

Ваш `h3` element treba da bude obojen korišćenjem klase `text-primary`

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-primary'));
```

Ваш `h3` element treba da bude centriran korišćenjem klase `text-center`

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-center'));
```

Ваш `h3` element treba da ima tekst `jQuery Playground`.

```js
assert.match(document.querySelector('h3')?.textContent, /jquery(\s)+playground/gi);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```