---
id: bad87fee1347bd9aedf08845
title: Odbaci prilagodjeni CSS i koristi Bootstrap
challengeType: 0
forumTopicId: 17565
dashedName: ditch-custom-css-for-bootstrap
---

# --description--

Mozemo da ocistimo nas kod i ucinimo da nasa aplikacija CatPhotoApp izgleda urednije koristeci Bootstrap-ove ugradjene stilove umesto prilagodjenih stilova koje smo ranije napravili.

Ne brini, kasnije cemo imati dovoljno vremena da prilagodjavamo nas CSS.

Obrisi CSS deklaracije za `.red-text`, `p` i `.smaller-image` iz svog `style` elementa, tako da u njemu ostanu samo deklaracije za `h2` i `thick-green-border`.

Zatim obrisi `p` element sa neispravnim linkom. Posle toga ukloni klasu `red-text` sa svog `h2` elementa i zameni je Bootstrap klasom `text-primary`.

Na kraju ukloni klasu `smaller-image` sa prvog `img` elementa i zameni je klasom `img-responsive`.

# --hints--

Tvoj `h2` element vise ne bi trebalo da ima klasu `red-text`.

```js
assert.isFalse(document.querySelector('h2')?.classList?.contains('red-text'));
```

Tvoj `h2` element sada treba da ima klasu `text-primary`.

```js
assert.isTrue(document.querySelector('h2')?.classList?.contains('text-primary'));
```

Tvoji pasusi vise ne bi trebalo da koriste font `Monospace`.

```js
const paragraphElement = document.querySelector('p');
const paragraphFontFamily = window.getComputedStyle(paragraphElement)["font-family"]; 
assert.notMatch(paragraphFontFamily,/monospace/i);
```

Klasu `smaller-image` treba ukloniti sa tvoje gornje slike.

```js
assert.isFalse(document.querySelector('img')?.classList?.contains('smaller-image'));
```

Treba da dodas klasu `img-responsive` na svoju gornju sliku.

```js
assert.lengthOf(document.querySelectorAll('.img-responsive'),2);
```

# --seed--

## --seed-contents--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  .red-text {
    color: red;
  }

  h2 {
    font-family: Lobster, Monospace;
  }

  p {
    font-size: 16px;
    font-family: Monospace;
  }

  .thick-green-border {
    border-color: green;
    border-width: 10px;
    border-style: solid;
    border-radius: 50%;
  }

  .smaller-image {
    width: 100px;
  }
</style>

<div class="container-fluid">
  <h2 class="red-text text-center">CatPhotoApp</h2>

  <p>Click here for <a href="#">cat photos</a>.</p>

  <a href="#"><img class="smaller-image thick-green-border" src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg" class="img-responsive" alt="Tuxedo cats running on dirt ground.">
  <div class="row">
    <div class="col-xs-4">
      <button class="btn btn-block btn-primary">Like</button>
    </div>
    <div class="col-xs-4">
      <button class="btn btn-block btn-info">Info</button>
    </div>
    <div class="col-xs-4">
      <button class="btn btn-block btn-danger">Delete</button>
    </div>
  </div>
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor"> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label>
    <label><input type="checkbox" name="personality"> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Crazy</label>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</div>
```

# --solutions--

```html
<link href="https://fonts.googleapis.com/css?family=Lobster" rel="stylesheet" type="text/css">
<style>
  h2 {
    font-family: Lobster, Monospace;
  }

  .thick-green-border {
    border-color: green;
    border-width: 10px;
    border-style: solid;
    border-radius: 50%;
  }
</style>

<div class="container-fluid">
  <h2 class="text-primary text-center">CatPhotoApp</h2>

  <a href="#"><img class="img-responsive thick-green-border" src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg" class="img-responsive" alt="Tuxedo cats running on dirt ground.">
  <div class="row">
    <div class="col-xs-4">
      <button class="btn btn-block btn-primary">Like</button>
    </div>
    <div class="col-xs-4">
      <button class="btn btn-block btn-info">Info</button>
    </div>
    <div class="col-xs-4">
      <button class="btn btn-block btn-danger">Delete</button>
    </div>
  </div>
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://freecatphotoapp.com/submit-cat-photo">
    <label><input type="radio" name="indoor-outdoor"> Indoor</label>
    <label><input type="radio" name="indoor-outdoor"> Outdoor</label>
    <label><input type="checkbox" name="personality"> Loving</label>
    <label><input type="checkbox" name="personality"> Lazy</label>
    <label><input type="checkbox" name="personality"> Crazy</label>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</div>
```
