---
id: 66ebd4ae2812430bb883c787
title: Izgradite Event Hub
challengeType: 25
dashedName: lab-event-hub
demoType: onClick
---

# --description--

U ovom laboratorijumu ćete koristiti semantičke HTML elemente za kreiranje strukture veb stranice. Dodaćete sadržaj i slike kako bi izgledalo kao pravi centar događaja.

**Cilj:** Ispunite priče korisnika ispod i prođite sve testove da biste završili laboratorijum.

**Priče korisnika:**

1. Treba vam `header` element.
1. Unutar `header` elementa, treba vam `h1` element koji sadrži tekst `Event Hub`, i `nav` element.
1. Unutar `nav` elementa, treba vam neuređeni spisak od dva stavke koje sadrže linkove ka različitim sekcijama stranice. Prva stavka treba da ima tekst `Upcoming Events`, a druga stavka treba da ima tekst `Past Events`.
1. Svaki link bi trebalo da bude predstavljen `a` elementom sa atributom `href` koji vodi ka odgovarajućoj sekciji stranice, `#upcoming-events` i `#past-events` respektivno.
1. Treba vam `main` element koji sadrži različite sekcije stranice.
1. Unutar `main` elementa, treba vam dva `section` elementa.
1. Prvi `section` element bi trebalo da ima atribut `id` sa vrednošću `upcoming-events`
1. Unutar sekcije `#upcoming-events`, treba vam:
   
   - `h2` element sa tekstom `Upcoming Events`.
   - Dva `article` elementa. Svaki članak bi trebalo da predstavlja događaj, i treba da ima:
      - `h3` element za naslov događaja.
      - `p` element za opis događaja. Možete dodati datum na dno ako želite.

1. Drugi `section` element bi trebalo da ima atribut `id` sa vrednošću `past-events`.
1. Unutar sekcije `#past-events`, treba vam:

   - `h2` element sa tekstom `Past Events`.
   - Dva `article` elementa. Svaki članak bi trebalo da predstavlja prošli događaj, i treba da ima:
     - `h3` element za naslov događaja, 
     - `p` element za opis događaja. Možete dodati datum na dno ako želite.
     - Element slike sa atributom `src` koji ukazuje na fajl slike i atributom `alt` sa opisom slike.

**Napomena:** Možete koristiti bilo koji tekst za opise događaja i datume. Ako želite, možete koristiti sledeće URL adrese slika za slike:

- `https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg`.
- `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

# --hints--

You should have a `header` element.

```js
assert.isNotNull(document.querySelector("header"));
```

Your `header` element should come after the opening `body` tag.

```js
assert.equal(document.querySelector("body")?.firstElementChild?.tagName, "HEADER");
```

Inside the `header` element, you should have an `h1` element that contains the text `Event Hub`.

```js
const h1Element = document.querySelector('header h1');
assert.strictEqual(h1Element?.innerText.trim(), "Event Hub");
```

Inside the `header` element, after the `h1` element, you should have a `nav` element.

```js
assert.isNotNull(document.querySelector("header>h1+nav"));
```

Your `nav` element should contain an unordered list of two items.

```js
const liElements = document.querySelectorAll('header nav>ul>li');

assert.isNotNull('header nav>ul');
assert.strictEqual(liElements.length, 2);
```

The first item in the unordered list should be a link.

```js
const firstLink = document.querySelectorAll('header nav ul li a')[0];
assert.exists(firstLink);
```

The second item in the unordered list should be a link.

```js
const secondLink = document.querySelectorAll('header nav ul li a')[1];
assert.exists(secondLink);
```

The text of the first item in the unordered list should be `Upcoming Events`.

```js
const firstLink = document.querySelectorAll('header nav>ul>li>a')[0];
assert.strictEqual(firstLink.innerText.trim(), "Upcoming Events");
```

The first item in the unordered list should have the `href` set to `#upcoming-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[0];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#upcoming-events");
```

The text of the second item in the unordered list should be `Past Events`.

```js
const secondLink = document.querySelectorAll('header nav>ul>li>a')[1];
assert.strictEqual(secondLink.innerText.trim(), "Past Events");
```

The second item in the unordered list should have the `href` set to `#past-events`.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[1];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#past-events");
```

You should have a `main` element after the `header` element closing tag.

```js
const mainElement = document.querySelector("body>header+main");
assert.isNotNull(mainElement);
```

Inside the `main` element, you should have two `section` elements.

```js
const sectionElements = document.querySelectorAll('body>header+main>section');
assert.strictEqual(sectionElements.length, 2);
```

Your first `section` element should have an `id` attribute with the value `upcoming-events`.

```js
const firstSection = document.querySelectorAll('body>header+main>section')[0];
const idAttribute = firstSection?.getAttribute("id");
assert.strictEqual(idAttribute, "upcoming-events");
```

Your second `section` element should have an `id` attribute with the value `past-events`.

```js
const secondSection = document.querySelectorAll('body>header+main>section')[1];
const idAttribute = secondSection?.getAttribute("id");
assert.strictEqual(idAttribute, "past-events");
```

Inside the `#upcoming-events` section, you should have an `h2` element with the text `Upcoming Events`.

```js
const h2Element = document.querySelector('#upcoming-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Upcoming Events");
```

Inside the `#upcoming-events` section, you should have two `article` elements below the `h2` element.

```js
const articleElements = document.querySelectorAll('#upcoming-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Both of the `article` elements inside the `#upcoming-events` section should have an `h3` element for the event title.

```js
const h3Elements = document.querySelectorAll('#upcoming-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Both of the `article` elements inside the `#upcoming-events` section should have a paragraph element for the event description.

```js
const articles = document.querySelectorAll('#upcoming-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

Inside the `#past-events` section, you should have an `h2` element with the text `Past Events`.

```js
const h2Element = document.querySelector('#past-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Past Events");
```

Inside the `#past-events` section, you should have two `article` elements below the `h2` element.

```js
const articleElements = document.querySelectorAll('#past-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Both of the `article` elements inside the `#past-events` section should have an `h3` element for the event title.

```js
const h3Elements = document.querySelectorAll('#past-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Both of the `article` elements inside the `#past-events` section should have a paragraph element for the event description.

```js
const articles = document.querySelectorAll('#past-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

Both of the `article` elements inside the `#past-events` section should have an image element.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);
```

Both of the image elements inside the `#past-events` section should have the `src` attribute pointing to an image file.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("src"));
}
```

Both of the image elements inside the `#past-events` section should have the `alt` attribute with a description of the image.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("alt"));
}
```

Each `h3` element should have the event title.

```js
const eventTitles = document.querySelectorAll('h3');
assert.isNotEmpty(eventTitles);
eventTitles.forEach((eventTitle => assert.isNotEmpty(eventTitle.innerText.trim())));
```

Each `p` element should have the event description.

```js
const eventDescriptions = document.querySelectorAll('p');
assert.isNotEmpty(eventDescriptions);
eventDescriptions.forEach((eventDescription => assert.isNotEmpty(eventDescription.innerText.trim())));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Event Hub</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Event Hub</title>
</head>
<body>
    <header>
        <h1>Event Hub</h1>
        <nav>
            <ul>
                <li><a href="#upcoming-events">Upcoming Events</a></li>
                <li><a href="#past-events">Past Events</a></li>
            </ul>
        </nav>
    </header>
    <main>
        <section id="upcoming-events">
            <h2>Upcoming Events</h2>
            <article>
                <h3>AI & Machine Learning Conference 2024</h3>
                <p>Join us for a deep dive into the latest advancements in artificial intelligence and machine learning. Industry leaders will share insights and case studies on how AI is transforming various sectors.</p>
                <p>Date: August 10, 2024</p>
            </article>
            <article>
                <h3>Web Development Bootcamp</h3>
                <p>A hands-on workshop designed for developers looking to enhance their skills in modern web technologies including React, Node.js, and GraphQL. Perfect for both beginners and experienced developers.</p>
                <p>Date: September 5, 2024</p>
            </article>
        </section>
        <section id="past-events">
            <h2>Past Events</h2>
            <article>
                <h3>Cybersecurity Summit 2024</h3>
                <p>An event focusing on the latest trends and threats in cybersecurity. Experts discussed strategies for protecting data and ensuring privacy in an increasingly digital world.</p>
                <p>Date: June 15, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg" alt="Image from Cybersecurity Summit 2024">
            </article>
            <article>
                <h3>Blockchain Expo 2024</h3>
                <p>A comprehensive event covering the future of blockchain technology. Topics included decentralized finance (DeFi), smart contracts, and the impact of blockchain on various industries.</p>
                <p>Date: July 20, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg" alt="Image from Blockchain Expo 2024">
            </article>
        </section>
    </main>
</body>
</html>
```

