---
id: 587d7dae367417b2b2512b79
title: Proširite komponentu ili primite unos
challengeType: 1
forumTopicId: 18235
dashedName: extend-constructors-to-receive-arguments
---

# --description--

Komponenti `Bird` i `Dog` iz prethodnog izazova radile su dobro. Međutim, primetite da sve `Birds` kreirane od komponente `Bird` dobijaju ime Albert automatski, imaju plavu boju i dva para nogu. Da li biste želeli ptice sa različitim atributima imena i boje? Moguće je promeniti atribute svake ptice direktno, ali to će biti mnogo posla:

```js
let swan = new Bird();
swan.name = "Carlos";
swan.color = "white";
```

Zamislite da kreirate program za praćenje stotina ili čak hiljada različitih ptica u aviariju. To oduzima mnogo vremena da se kreiraju sve ptice, pa da im se mijenjaju karakteristike sa različitim vrijednostima za svaku. Da lako kreirate različite stvari od `Bird`, možete dizajnirati svoju komponentu Ptica da prima parametre:

```js
function Bird(name, color) {
  this.name = name;
  this.color = color;
  this.numLegs = 2;
}
```

Zatim prosledite te vrednosti kao argumente kako biste definisali svaku jedinstvenu pticu u komponenti `Bird`: `let cardinal = new Bird("Bruce", "red");`. Ovo pruža novi primer za `Bird` sa svojstvima od `name` i `color` postavljena na `Bruce` i `red`, respektivno. Svojstvo `numLegs` je i dalje postavljeno na 2. `cardinal` ima ova svojstva:

```js
cardinal.name
cardinal.color
cardinal.numLegs
```

Komponente su lakše. Sada je moguće definisati svojstva svakog `Bird` prilikom kreiranja, što je jedan od načina na koji su komponente JavaScript veoma važne. Oni grupišu stvari zajedno prema sličnim karakteristikama i definišu mapu koja pojednostavljuje njihovo kreiranje.

# --instructions--

Unda kijenzi kingine cha `Dog`. Mara hii, kiandae kupokea vigezo vya `name` na `color`, na kuwa na sifa ya `numLegs` iliyowekwa kuwa 4. Kisha unda `Dog` mpya iliyohifadhiwa katika kigezo cha `terrier`. Pitia vigezo viwili vya mfuatano wa herufi kama hoja kwa sifa za `name` na `color`.

# --hints--

`Dog` inapaswa kupokea hoja kwa ajili ya `name`.

```js
assert(new Dog('Clifford').name === 'Clifford');
```

`Dog` inapaswa kupokea hoja kwa ajili ya `color`.

```js
assert(new Dog('Clifford', 'yellow').color === 'yellow');
```

`Dog` inapaswa kuwa na sifa ya `numLegs` iliyowekwa kuwa 4.

```js
assert(new Dog('Clifford').numLegs === 4);
```

`terrier` inapaswa kuundwa kwa kutumia kijenzi cha `Dog`.

```js
assert(terrier instanceof Dog);
```

# --seed--

## --seed-contents--

```js
function Dog() {

}
```

# --solutions--

```js
function Dog (name, color) {
  this.numLegs = 4;
  this.name = name;
  this.color = color;
}

const terrier = new Dog();
```
