---
id: 587d7db0367417b2b2512b84
title: Urithi wa tabia kutoka kwa supertype
challengeType: 1
forumTopicId: 301319
dashedName: inherit-behaviors-from-a-supertype
---

# --description--

Katika changamoto iliyopita, uliunda `supertype` iitwayo `Animal` ambayo ilibainisha tabia zinazoshirikiwa na wanyama wote:

```js
function Animal() { }
Animal.prototype.eat = function() {
  console.log("nom nom nom");
};
```

Changamoto hii na ile inayofuata zitashughulikia jinsi ya kutumia tena njia za `Animal` ndani ya `Bird` na `Dog` bila kuzitengeneza tena. Inatumia mbinu inayoitwa urithi. Changamoto hii inahusu hatua ya kwanza: tengeneza mfano wa `supertype` (au mzazi). Tayari unajua njia moja ya kutengeneza mfano wa `Animal` kwa kutumia opereta `new`:

```js
let animal = new Animal();
```

Postoje određeni nedostatci prilikom korišćenja ove sintakse za nasleđivanje, što je veoma teško u opsegu ovog izazova. Umesto toga, ovde je alternativa metoda koja nema te gubitke:

```js
let animal = Object.create(Animal.prototype);
```

`Object.create(obj)` kreira novi objekat, i postavlja `obj` kao `prototype` tog novog objekta. Zapamtite da je `prototype` kao "recept" za kreiranje objekta. Postavljanjem `prototype` od `animal` u `prototype` od `Animal`, dajete primer za `animal` "recept" sličan bilo kom drugom primeru od `Animal`.

```js
animal.eat();
animal instanceof Animal;
```

Put do `instanceof` ovde će vratiti `true`.

# --instructions--

Tumia `Object.create` kutengeneza mifano miwili ya `Animal` iitwayo `duck` na `beagle`.

# --hints--

Kigezo cha `duck` kinapaswa kufafanuliwa.

```js
assert(typeof duck !== 'undefined');
```

Kigezo cha `beagle` kinapaswa kufafanuliwa.

```js
assert(typeof beagle !== 'undefined');
```

Parametar za `duck` mora biti inicijalizovan sa `Object.create`.

```js
assert(
  /(let|const|var)\s{1,}duck\s*=\s*Object\.create\s*\(\s*Animal\.prototype\s*\)\s*/.test(
    __helpers.removeJSComments(code)
  )
);
```

Kigezo cha `beagle` kinapaswa kuanzishwa na `Object.create`.

```js
assert(
  /(let|const|var)\s{1,}beagle\s*=\s*Object\.create\s*\(\s*Animal\.prototype\s*\)\s*/.test(
    __helpers.removeJSComments(code)
  )
);
```

`duck` treba da ima `prototype` od `Animal`.

```js
assert(duck instanceof Animal);
```

`beagle` treba da ima `prototype` sa `Animal`.

```js
assert(beagle instanceof Animal);
```

# --seed--

## --seed-contents--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

// Only change code below this line

let duck; // Change this line
let beagle; // Change this line
```

# --solutions--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};
let duck = Object.create(Animal.prototype);
let beagle = Object.create(Animal.prototype);

duck.eat();
beagle.eat();
```
