---
id: 66bba6fff611169359d9d36a
title: Napravite stranicu za privatnu e-poštu
challengeType: 25
dashedName: build-a-confidential-email-page
demoType: onClick
---

# --description--

**Cilj:** Popunite korisničke priče navedene ispod i prođite sve testove da biste završili laboratoriju.

**Priče korisnika:**

1. Mora da ima element sa `main` koji ima `id` od `email`.
1. Njegov element `#email` mora da ima `padding` od `50px`, gornji rub `50px`, `width` od `500px`, i `border` koji je `2px`.
1. Ukupna širina njegovog elementa `#email`, uključujući interne razmake i granice, mora biti `500px`.
1. Mora da ima dva elementa `div`, jedan sa ID-om `confidential` i drugi sa ID-om `top-secret`, unutar njegovog elementa `#email`.
1. Njegovi elementi za `#confidential` i `#top-secret` moraju imati `display` od `inline-block`.
1. Njegovi elementi za `#confidential` i `#top-secret` moraju imati `padding`, levi rub i `border`.
1. Element `#confidential` mora imati tekst `CONFIDENTIAL`.
1. Element `#top-secret` mora imati tekst `TOP SECRET`.
1. Njegovi elementi za `#confidential` i `#top-secret` moraju biti rotirani korišćenjem transformacija CSS.
1. Mora da ima najmanje tri paragraf elementa unutar njegovog elementa `#email`.
1. Mora da ima najmanje tri elementi `span` sa klasom objekta `blurred`, unutar svojih paragraf elemenata.
1. Mora da ima padajuću listu `blurred` koja čini element `3px` manje svetlim korišćenjem filtera CSS.

# --hints--

Trebalo bi da imaš element sa `main` koji ima `id` od `email`.

```js
assert.exists(document.querySelector('main#email'));
```

Trebalo bi da imaš selektor `#email` koji postavlja `padding` svih njegovih elemenata na `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('padding'), '50px');
```

Trebalo bi da imaš birač za `#email` koji postavlja `margin-top` njegovih svojstava na `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('margin-top'), '50px');
```

Trebalo bi da imaš birač za `#email` koji postavlja `width` njegovih elemenata na `500px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('width'), '500px');
```

Tvoj element za `#email` bi trebalo da ima granicu sa `2px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('border-width'), '2px');
```

Vaš element za `#email` treba da ima `box-sizing` od `border-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.boxSizing, 'border-box');
```

Trebalo bi da imate element od `div` sa `id` iz `confidential` unutar vašeg elementa `#email`.

```js
assert.exists(document.querySelector('#email > div#confidential'));
```

Unapaswa kuwa na kipengele cha `div` chenye `id` ya `top-secret` ndani ya kipengele chako cha `#email`.

```js
assert.exists(document.querySelector('#email > div#top-secret'));
```

Vaše polje `#confidential` mora imati `display` postavljeno na `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#confidential'))?.display, 'inline-block');
```

Vaše polje za `#top-secret` treba da ima `display` postavljeno na `inline-block`.

```js
assert.equal(window.getComputedStyle(document.querySelector('#top-secret'))?.display, 'inline-block');
```

Vaš element za `#confidential` mora imati `padding` sa svih strana.

```js
const style = window.getComputedStyle(document.querySelector('#confidential'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

Vaš element za `#top-secret` treba imati `padding` sa svih strana.

```js
const style = window.getComputedStyle(document.querySelector('#top-secret'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

Tvoj element za `#confidential` treba da ima levu ivicu.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.marginLeft, '0px');
```

Tvoj element za `#top-secret` treba da ima levu ivicu.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.marginLeft, '0px');
```

Vaš element od `#confidential` treba da ima `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.borderWidth, '0px');
```

Vaš deo `#top-secret` trebalo bi da ima `border`.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.borderWidth, '0px');
```

Tekst vašeg dela `#confidential` treba da bude `CONFIDENTIAL`. Izostavili ste tekst ili postoji greška u kucanju.

```js
assert.equal(document.querySelector('#confidential').innerText.toLowerCase(), 'confidential' );
```

Tekst vašeg elementa za `#top-secret` bi trebalo da bude `TOP SECRET`. Ostavili ste tekst ili postoje pravopisne greške.

```js
assert.equal(document.querySelector('#top-secret').innerText.toLowerCase(), 'top secret' );
```

Vaš element ``#confidential`` mora da koristi promene od `CSS` za rotiranje elemenata ``transform``.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.transform, 'none');
```

Tvoj element sa `#top-secret` treba da koristi izmjene od CSS za `transform` kako bi rotirao element.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.transform, 'none');
```

Trebalo bi da imaš najmanje tri elementa paragrafa unutar tvog elementa `#email`.

```js
assert.isAtLeast(document.querySelectorAll('#email > p').length, 3);
```

Trebalo bi da imaš najmanje tri elementa `span` sa klasom objekta `blurred` unutar svojih pasusa.

```js
assert.isAtLeast(document.querySelectorAll('p > span.blurred').length, 3);
```

Vaše komponente za `.blurred` ne smeju biti prazne.

```js
const els = document.querySelectorAll('.blurred');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.isAtLeast(el.innerText.length , 1))
```

Trebalo bi da imaš birač za `.blurred` koji postavlja `filter` njegovih elemenata na `blur(3px)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.blurred')?.filter, 'blur(3px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Confidential Email</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Secret Marshmallow Mission</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="email">
    <div id="confidential">CONFIDENTIAL</div>
    <p>Dear Agent <span class="blurred">S'more,</span></p>
    <p>We have an emergency. The secret formula for our <span class="blurred">Mega Marshmallow</span> has been
      compromised. This formula is what makes our marshmallows the fluffiest and most delicious.</p>
    <p>We suspect that <span class="blurred">Professor Puff</span> is behind this. He has taken the formula to his
      hidden laboratory. Your mission is to infiltrate the lab and secure the formula before it's too late.</p>
    <p>Be sure to keep the lab's location confidential. Any leak of this information could jeopardize the entire
      operation.</p>
    <div id="top-secret">TOP SECRET</div>
  </main>
</body>

</html>
```

```css
#email {
  background-color: SeaShell;
  margin: 50px auto;
  padding: 50px;
  width: 500px;
  border: 2px solid black;
  box-shadow: 5px 3px 3px gray;
  box-sizing: border-box;
}

#confidential {
  display: inline-block;
  margin-left: 100px;
  padding: 10px;
  text-align: center;
  font-size: 30px;
  transform: rotate(-20deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

#top-secret {
  margin-left: 200px;
  display: inline-block;
  padding: 10px;
  text-align: center;
  font-size: 20px;
  transform: rotate(15deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

.blurred {
  filter: blur(3px);
}
```
