---
id: 587d825a367417b2b2512c89
title: Implementiraj Quicksort Algoritam
challengeType: 26
dashedName: implement-the-quicksort-algorithm
---

# --description--

U ovom laboratorijumu ćete implementirati sortirajući algoritam: quicksort.

Quicksort je efikatan, rekurentan pristup tipa "podeli i vladaj" (divide-and-conquer) za sortiranje niza. U ovoj metodi, birja se pivot vrednost u originalnom nizu. Niz se zatim podeljuje na dva podniza sa vrednostima manjim i većim od pivot vrednosti. Zatim se kombinuju rezultati rekurentnog pozivanja quicksort algoritma na oba podniza. Ovo se nastavlja sve dok se ne dođe do osnovnog slučaja praznog ili niza sa jednim elementom, koji se vraća. Razmotavanje (unwinding) rekurentnih poziva vraća sortirani niz.

Iako je izbor pivot vrednosti važan, bilo koji pivot odgovara za naše svrhe ovde. Radi jednostavnosti, može se koristiti prvi ili poslednji element.

Quicksort je veoma efikatan metod sortiranja, pružajući performanse od *O(nlog(n))* u proseku. Takođe je relativno lako za implementaciju. Ove karakteristike čine ga popularnim i korisnim metodom sortiranja.

**Cilj:** Ispunite priče korisnika ispod i učinite da svi testovi prođu kako biste završili laboratoriju.

**Priče korisnika:**

1. Trebalo bi da imate funkciju `quicksort` koja prima niz celih brojeva kao ulaz.
1. Funkcija `quicksort` bi trebalo da vrati niz koji sadrži iste cele brojeve kao i ulaz, ali u redosledu od najmanjeg do najvećeg.
1. Funkcija `quicksort` bi trebalo rekurentno da poziva samu sebe kako bi sortirala niz.

# --hints--

You should have a `quicksort` function.

```js
assert.isFunction(quicksort);
```

`quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` should return an array that is unchanged except for order.

```js
assert.sameMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quicksort` should return a sorted array (least to greatest).

```js
assert.sameOrderedMembers(
  quicksort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`quicksort` should not use the built-in `.sort()` method.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quicksort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function quicksort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quicksort(lesser), ...equal, ...quicksort(greater)];
  }
}
```
