---
id: ac6993d51946422351508a41
title: Implementirajte algoritam za skraćivanje stringa
challengeType: 26
dashedName: truncate-a-string
---

# --description--

U ovom laboratorijumu ćete vežbati skraćivanje niza karaktera na određenu dužinu.

**Cilj:** Ispunite korisničke priče ispod i učinite da svi testovi prođu kako biste završili laboratorijum.

**Korisničke priče:**

1. Trebalo bi da imate funkciju `truncateString` koja prima dva argumenta: prvi je niz karaktera, a drugi je broj.
1. Ako je dužina niza karaktera veća od zadatog broja, niz bi trebalo skratiti da smanji dužinu tako da bude jednaka zadatom broju, a `...` treba dodati na kraj skraćenog niza.
1. Ako je dužina niza karaktera jednaka ili manja od zadatog broja, niz bi trebalo vratiti bez promena.

# --hints--

`truncateString("A-tisket a-tasket A green and yellow basket", 8)` should return the string `A-tisket...`.

```js
assert.strictEqual(
  truncateString('A-tisket a-tasket A green and yellow basket', 8),
  'A-tisket...'
);
```

`truncateString("Peter Piper picked a peck of pickled peppers", 11)` should return the string `Peter Piper...`.

```js
assert.strictEqual(
  truncateString('Peter Piper picked a peck of pickled peppers', 11),
  'Peter Piper...'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length)` should return the string `A-tisket a-tasket A green and yellow basket`.

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-tisket a-tasket A green and yellow basket", "A-tisket a-tasket A green and yellow basket".length + 2)` should return the string `A-tisket a-tasket A green and yellow basket`.

```js
assert.strictEqual(
  truncateString(
    'A-tisket a-tasket A green and yellow basket',
    'A-tisket a-tasket A green and yellow basket'.length + 2
  ),
  'A-tisket a-tasket A green and yellow basket'
);
```

`truncateString("A-", 1)` should return the string `A...`.

```js
assert.strictEqual(truncateString('A-', 1), 'A...');
```

`truncateString("Absolutely Longer", 2)` should return the string `Ab...`.

```js
assert.strictEqual(truncateString('Absolutely Longer', 2), 'Ab...');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function truncateString(str, num) {
  if (num >= str.length) {
    return str;
  }

  return str.slice(0, num) + '...';
}

```
