---
id: 686386e76197458b7dfd52ec
title: Korak 3
challengeType: 0
dashedName: step-3
---

# --description--

Ispod `h2`, kreiraj element `p` sa tekstom `$0`. Zatim, unutar tog elementa `p` i nakon teksta, kreiraj element `span` sa tekstom `/month`.

# --hints--

You should create a `p` element with the text `$0`.

```js
const p = document.querySelector("p");
const pTextBeforeSpan = p.childNodes[0].textContent.trim();

assert.equal(pTextBeforeSpan, "$0");
```

You should create a `span` element with the text `/month` inside the `p` element.

```js
const spanText = document.querySelector("p > span")?.textContent;
assert.equal(spanText, "/month");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            --fcc-editable-region--
            
            --fcc-editable-region--
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
