---
id: 686386e76197458b7dfd52ef
title: Korak 6
challengeType: 0
dashedName: step-6
---

# --description--

Sada ćete kreirati dugme za poziv na akciju (CTA). Poziv na akciju (Call to Action - CTA) je dugme ili link dizajniran da podstakne korisnika da preduzme određenu akciju, u ovom slučaju, da počne sa slušanjem.

Napolju od najdubljeg `div`, kreirajte ankorni element sa `href` vrednošću `#` i tekstom `Start listening`.

Time se završava HTML za prvu karticu sa cenama.

# --hints--

You should create an anchor element.

```js
assert.exists(document.querySelector("a"))
```

Your anchor element should have an `href` set to `#`.

```js
const link = document.querySelector("a")
assert.equal(link?.getAttribute("href"), "#")
```

Your anchor element should have the text `Start listening`.

```js
assert.equal(document.querySelector("a")?.textContent, "Start listening")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          --fcc-editable-region--
          
          --fcc-editable-region-- 
        </div>
      </div>
    </main>
  </body>
</html>
```
