---
id: 686383abac6c625a84458ea9
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U ovoj radionici, naučićete više o Tailwind CSS korišćenjem njegovih grid utilitetskih klasa za izgradnju komponente za cene.

Osnovni HTML šablon je dostavljen za vas. Počnite kreiranjem elementa `h1` sa tekstom `Choose your listening plan`.

# --hints--

You should create an `h1` element.

```js
assert.exists(document.querySelector("h1"));
```

Your `h1` element should have the text `Choose your listening plan`.

```js
assert.equal(document.querySelector("h1")?.textContent, "Choose your listening plan");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </main>
  </body>
</html>
```
