---
id: 686386e76197458b7dfd52eb
title: Korak 2
challengeType: 0
dashedName: step-2
---

# --description--

Kreiraj element `div` ispod `h1`. Unutar tog `div`, kreiraj još jedan element `div`, i u njega ugnjezdi još jedan element `div`.

I unutar najdubljeg `div`, kreiraj element `h2` sa tekstom `Listener`. Ovo će predstavljati prvi plan cena.

# --hints--

You should create a `div` element.

```js
assert.exists(document.querySelector("div"));
```

You should create another `div` element inside the outer `div`.

```js
assert.exists(document.querySelector("div > div"));
```

You should create another `div` element inside the inner `div`.

```js
assert.exists(document.querySelector("div > div > div"));
```

You should create an `h2` with the text `Listener` inside the innermost `div` element.

```js
assert.equal(document.querySelector("div > div > div > h2")?.textContent, "Listener");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
