---
id: 66c64210413532ee9d3bd342
title: Korak 3
challengeType: 1
dashedName: step-3
---

# --description--

Za ovaj sledeći deo radionice, veštaćete se dodavanju namirnica na listu za kupovinu.

Počnite korišćenjem `console.log` da logujete poruku `"It will be nice to have some fruit to eat."`.

# --hints--

You should log the message `"It will be nice to have some fruit to eat."` to the console.

```js
assert.match(code, /console\.log\s*\(\s*(['"])(It\s+will\s+be\s+nice\s+to\s+have\s+some\s+fruit\s+to\s+eat\.)\1\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

--fcc-editable-region--

--fcc-editable-region--
```
