---
id: 66c63bfa4685e4e3843fa0dc
title: Korak 1
challengeType: 1
dashedName: step-1
---

# --description--

U ovoj radionici, nastavićete da učite o nizovima kreiranjem liste za kupovinu namirnica.

Počnite dodavanjem `console.log` koji loguje string `"Grocery shopping list"` na konzolu.

# --hints--

You should have a `console.log` in your code.

```js
assert.match(code, /console\.log(.*)/);
```

Your `console.log` should output the message `"Grocery shopping list"`. Double check for spelling and capitalization errors.

```js
assert.match(code, /console\.log\((["'])(Grocery\s+shopping\s+list)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
