---
id: 66c645b345a39bf04864dc50
title: Korak 5
challengeType: 1
dashedName: step-5
---

# --description--

Da biste videli ažuriranu listu za kupovinu, moraćete da logujete trenutnu listu za kupovinu i kratku poruku u konzolu.

Pošto će se ova poruka koristiti više puta tokom radionice, najbolje je kreirati ponovo upotrebljivu funkciju.

Kreirajte funkciju nazvanu `getShoppingListMsg` koja prima niz kao parametar i vraća string `"Current Shopping List: "` praćen sadržajem priloženog niza.

Možete koristiti *template literals* ili spajanje stringova sa operaterom `+` da kombinujete string sa nizom prosleđenim funkciji.

# --hints--

You should have a function called `getShoppingListMsg` in your code.

```js
assert.isFunction(getShoppingListMsg);
```

Your `getShoppingListMsg` function should return a string.

```js
assert.isString(getShoppingListMsg(shoppingList));
```

Your `getShoppingListMsg` function should accept an array and return the message `"Current Shopping List: "` followed by the contents of that array. Double check spacing and spelling errors.

```js
assert.oneOf(getShoppingListMsg(["Anchovies", "Olives"]), ["Current Shopping List: Anchovies,Olives", "Current Shopping List: Anchovies, Olives"]);
assert.strictEqual(getShoppingListMsg(["Apples"]), "Current Shopping List: Apples");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

--fcc-editable-region--

--fcc-editable-region--
```
