---
id: 66c64095f61166ede6e39a84
title: Korak 2
challengeType: 1
dashedName: step-2
---

# --description--

Za ovu listu za kupovinu, koristićete niz da predstavite artikle koje treba da kupite.

U prethodnim lekcijama, naučili ste kako da kreirate nizove ovako:

```js
const fruits = ['apple', 'banana', 'orange'];
```

U ovom koraku, kreirajte varijablu nazvanu `shoppingList` i dodelite joj prazan niz.

# --hints--

You should have a variable called `shoppingList`.

```js
assert.isNotNull(shoppingList);
```

You should assign an array to the `shoppingList` variable.

```js
assert.isArray(shoppingList);
```

Your array should be empty.

```js
assert.isEmpty(shoppingList);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

--fcc-editable-region--

--fcc-editable-region--
```
