---
id: 66c6491fe4c8e0f16845425f
title: Korak 7
challengeType: 1
dashedName: step-7
---

# --description--

Sada je vreme da dodamo još jedno voće u listu.

Koristeći istu metodu niza kao ranije, dodajte string `"Grapes"` na kraj niza `shoppingList`.

Zatim, dodajte `console.log` i pozovite funkciju `getShoppingListMsg` sa nizom `shoppingList` kao argumentom unutar `console.log` kako biste videli ažuriranu listu logovanu u konzolu.

# --hints--

Your `shoppingList` array should now contain two items.

```js
assert.lengthOf(shoppingList, 2);
```

Your `shoppingList` array should contain the string `"Grapes"`.

```js
assert.equal(shoppingList[1], "Grapes");
```

You should call `getShoppingListMsg(shoppingList)` function inside of the `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
