---
id: 66c73a7798f6f62b2ae58f22
title: Korak 16
challengeType: 1
dashedName: step-16
---

# --description--

Sada je vreme da dodamo više stavki na početak liste namirnica.

Počnite dodavanjem izjave `console.log` koja loguje poruku `"It might be nice to get a dessert."`

Ispod te izjave `console`, koristite odgovarajuću metodu niza (array) da dodate string `"Chocolate Cake"` na početak niza `shoppingList`.

Konačno, dodajte `console.log` i pozovite funkciju `getShoppingListMsg` sa nizom `shoppingList` kao argumentom unutar `console.log` kako biste videli ažuriranu listu logovanu u konzolu.

# --hints--

You should have a `console.log` that logs the message `"It might be nice to get a dessert."`.

```js
assert.match(code, /console\.log\((['"])(It\s+might\s+be\s+nice\s+to\s+get\s+a\s+dessert\.)\1\)/);
```

Your `shoppingList` array should have a total of `6` items.

```js
assert.lengthOf(shoppingList, 6);
```

You should use the `unshift` method to add the string `"Chocolate Cake"` to the beginning of the `shoppingList` array.

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

You should call `getShoppingListMsg(shoppingList)` inside of the `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 6);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
