---
id: 66c726c34ecf1e238aa9d7d5
title: Korak 9
challengeType: 1
dashedName: step-9
---

# --description--

U prethodnim lekcijama, naučili ste kako da dodate elemente na početak niza koristeći metod `unshift()`.

Evo podsetnika o tome kako koristiti metod `unshift()`:

```js
array.unshift(item1, item2, ..., itemX);
```

Koristite metod `unshift()` da dodate string `"Vegetable Oil"` na početak niza `shoppingList`.

# --hints--

You should have three items in your `shoppingList` array.

```js
assert.lengthOf(shoppingList, 3);
```

You should use the `unshift()` method to add `"Vegetable Oil"` to the beginning of the `shoppingList` array.

```js
assert.equal(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

--fcc-editable-region--

--fcc-editable-region--
```
