---
id: 66c72b0ffbc5522525768558
title: Korak 11
challengeType: 1
dashedName: step-11
---

# --description--

U ranijim koracima, pregledali ste kako dodati element na kraj niza koristeći metod `push`.

Ali metoda `push` prihvata više argumenata, tako da možete dodati više elemenata na kraj niza ovako:

```js
array.push(item1, item2, item3);
```

U ovom koraku, koristite metodu `push` da dodate nizove `"Popcorn", "Beef Jerky", "Potato Chips"` u niz `shoppingList`.

Red je važan, pa se pobrinite da dodate elemente u redu kako su navedeni.

# --hints--

Your `shoppingList` array should have a total of `6` strings.

```js
assert.lengthOf(shoppingList, 6);
```

Your `shoppingList` should have the string `"Popcorn"`.

```js
assert.strictEqual(shoppingList[3], "Popcorn");
```

Your `shoppingList` should have the string `"Beef Jerky"`.

```js
assert.strictEqual(shoppingList[4], "Beef Jerky");
```

Your `shoppingList` should have the string `"Potato Chips"`.

```js
assert.strictEqual(shoppingList[5], "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
