---
id: 61fd67a656743144844941cb
title: Korak 4
challengeType: 0
dashedName: step-4
---

# --description--

Čitači ekrana najavljuju elemente HTML prema protoku dokumenta. Na kraju, želimo da izvještaj o stanju ima naslov sekcije "Balance Sheet" i podnaslov "AcmeWidgetCorp". Međutim, ovaj raspored se ne može dobro razumjeti ako ga čitač ekrana najavi.

Uzmi atribut/svojstvo `span` sa vrednošću `class` za `flex` koji imaš, a zatim dodaj dva elementa/komponente `span` unutar njega. Daj prvom elementu tekst `AcmeWidgetCorp`. Daj drugom elementu tekst `Balance Sheet`. Koristićeš CSS da promeniš redosled teksta na stranici, ali će redosled od HTML biti značajniji za čitač ekrana.

# --hints--

Vaše postojeće polje `span` mora imati svojstvo `class` sa vrednošću `flex`.

```js
assert(document.querySelector('h1')?.children?.[0]?.classList?.contains('flex'));
```

Kipengele chako cha `span` kilichopo kinapaswa kuwa na vipengele viwili vipya vya `span` ndani yake.

```js
assert(document.querySelector('.flex')?.children?.[0]?.localName === 'span');
assert(document.querySelector('.flex')?.children?.[1]?.localName === 'span');
```

Tvoji novi elementi `span` ne bi trebalo da imaju atribut `class`.

```js
assert(!document.querySelector('.flex')?.children?.[0]?.classList?.length);
assert(!document.querySelector('.flex')?.children?.[1]?.classList?.length);
```

Tvoj novi prvi element `span` treba da ima tekst `AcmeWidgetCorp`.

```js
assert(document.querySelector('.flex')?.children?.[0]?.textContent === 'AcmeWidgetCorp');
```

Vaš drugi novi element od `span` mora imati tekst `Balance Sheet`.

```js
assert(document.querySelector('.flex')?.children?.[1]?.textContent === 'Balance Sheet');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
