---
id: 686d749d6f32f708fdf2b7ad
title: Korak 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

U prethodnoj lekciji ste naučili o elementima `svg` i kako se često koriste kao ikone u projektima. U stvarnim bazama koda, obično se oslanjate na biblioteke ikona, tako da ne morate kreirati elemente `svg` od nule. Međutim, u ovom radionici, kreiraćete ikonu srca da biste naučili o osnovnim atributima korišćenim unutar elementa `svg`.

Počnite kreiranjem elementa `svg` na stranici.

# --hints--

You should have an `svg` element on the page.

```js
const svg = document.querySelector('svg');
assert.exists(svg);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region-- 
    
    --fcc-editable-region--
  </body>
</html>
```
