---
id: 68a47dcdbce12193caf77012
title: Implementujte Pronalazač sortirovanog indeksa
challengeType: 26
dashedName: lab-sorted-index-finder
---

# --description--

U ovom laboratorijumu ćete kreirati funkciju koja vraća najniži indeks na koji bi vrednost trebalo umetnuti u niz nakon što je sortiran po rastućem redosledu.

**Cilj**: Ispunite priče korisnika ispod i učinite da svi testovi prođu kako biste završili laboratorijum.

**Priče korisnika:**

1. Trebalo bi da imate funkciju `getIndexToIns` koja prima dva argumenta: niz i broj.  
2. Treba vam koristiti metod `sort` za sortiranje niza po rastućem redosledu.
3. Vaša funkcija `getIndexToIns` treba da vrati najniži indeks na koji bi broj trebalo umetnuti korišćenjem metoda `findIndex`.
4. Vaša funkcija `getIndexToIns` uvek treba da vrati broj.

**Savet:**  
Metod `findIndex` je ugrađeni niz metod u JavaScriptu. Prima povratnu funkciju (callback function) i vraća indeks prvog elementa koji zadovoljava uslov. I `findIndex` i `sort` su funkcije višeg reda (higher-order functions).

**Primeri:**

- `getIndexToIns([1, 2, 3, 4], 1.5)` bi trebalo da vrati `1` jer je `1.5` veći od `1` (indeks `0`) i manji od `2` (indeks `1`).
- `getIndexToIns([20, 3, 5], 19)` bi trebalo da vrati `2` jer, nakon sortiranja na `[3, 5, 20]`, `19` je manji od `20` (indeks `2`) i veći od `5` (indeks `1`).

# --hints--

You should have a `getIndexToIns` function.

```js
assert.isFunction(getIndexToIns);
```

`getIndexToIns` should always return a number.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

Your function should use the `sort` method.

```js
assert.match(getIndexToIns.toString(), /\.sort\(/);
```


Your function should make use of the `findIndex` method.

```js
assert.match(getIndexToIns.toString(), /\.findIndex\(/);
```

`getIndexToIns([10, 20, 30, 40, 50], 35)` should return `3`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` should return `2`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

`getIndexToIns([40, 60], 50)` should return `1`.

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

`getIndexToIns([3, 10, 5], 3)` should return `0`.

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

`getIndexToIns([5, 3, 20, 3], 5)` should return `2`.

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

`getIndexToIns([2, 20, 10], 19)` should return `2`.

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

`getIndexToIns([3, 10, 5], 11)` should return `3`

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 11), 3);
```

`getIndexToIns([], 5)` should return `0`

```js
assert.strictEqual(getIndexToIns([], 5), 0);
```

# --seed--

## --seed-contents--

```js
``` 

# --solutions--

```js
function getIndexToIns(arr, num) {
  const sorted = [...arr].sort((a, b) => a - b);
  const index = sorted.findIndex(el => el >= num);

  return index === -1 ? sorted.length : index;
}
``` 

