---
id: 587d7dbc367417b2b2512bb1
title: Kreiraj jednostavan JSX element
challengeType: 6
forumTopicId: 301390
dashedName: create-a-simple-jsx-element
---

# --description--

React je biblioteka za prikaze otvorenog izvora, kreirana i održavana od strane Facebooka. To je odličan alat za prikaz korisničkog interfejsa (UI) moderne veb aplikacije.

`React koristi sintaktičku ekstenziju JavaScript pod nazivom JSX, koja omogućava pisanje HTML direktno unutar JavaScript. Ovo ima nekoliko prednosti. Omogućava korišćenje pune moći programa JavaScript unutar HTML, i pomaže da vaš kod bude čitljiv. U velikoj meri, JSX je sličan HTML koji ste već naučili, međutim postoje neke ključne razlike koje ćemo pokriti u ovim izazovima.`

Na primer, jer je JSX sintaktička ekstenzija za JavaScript, možete napisati JavaScript direktno unutar JSX-a. Da biste to uradili, uključujete kod koji želite da se tretira kao JavaScript unutar zagrada: `{ 'this is treated as JavaScript code' }`. Zapamtite ovo, jer se koristi u nekoliko budućih izazova.

Hata hivyo, kwa sababu JSX si JavaScript halali, msimbo wa JSX lazima ukusanywe kuwa JavaScript. Transpiler Babel ni chombo maarufu kwa mchakato huu. Kwa urahisi wako, tayari imeongezwa nyuma ya pazia kwa changamoto hizi. Ikiwa utatengeneza JSX isiyo sahihi kisintaksia, utaona jaribio la kwanza katika changamoto hizi lishindikane.

Treba napomenuti da su ove izazove ispod haube nazvani `ReactDOM.render(JSX, document.getElementById('root'))`. Ovaj poziv funkcije je ono što stavlja vaš JSX u lagani prikaz/predstavljanje React od DOM. React koristi snimke (snapshot) svog DOM kako bi poboljšao ažuriranje samo za određene delove stvarnog DOM.

# --instructions--

Trenutni kod koristi JSX da postavi element `div` za konstantu `JSX`. Promeniте `div` u element `h1` i dodajte tekst `Hello JSX!` unutar njega.

# --hints--

Mora `JSX` da vrati funkciju/element `h1`.

```js
assert(JSX.type === 'h1');
```

Etiketa za `h1` treba da sadrži tekst `Hello JSX!`

```js
assert(Enzyme.shallow(JSX).contains('Hello JSX!'));
```

# --seed--

## --seed-contents--

```jsx
const JSX = <div></div>;
```

# --solutions--

```jsx
const JSX = <h1>Hello JSX!</h1>;
```
