---
id: 5a24c314108439a4d403616a
title: Prošaljite red podataka kao parametre
challengeType: 6
forumTopicId: 301401
dashedName: pass-an-array-as-props
---

# --description--

Changamoto ya mwisho ilionyesha jinsi ya kupitisha taarifa kutoka sehemu ya mzazi kwenda sehemu ya mtoto kama `props` au sifa. Changamoto hii inaangalia jinsi safu za data zinavyoweza kupitishwa kama `props`. Ili kupitisha safu ya data kwa kipengele cha JSX, lazima itendwe kama JavaScript na ifungwe ndani ya mabano ya mviringo.

```jsx
<ParentComponent>
  <ChildComponent colors={["green", "blue", "red"]} />
</ParentComponent>
```

Detilji deo zatim ima pristup svojstvu niza podataka `colors`. Metode nizova kao što su `join()` mogu se koristiti prilikom dobijanja tog svojstva.

```jsx
const ChildComponent = (props) => <p>{props.colors.join(', ')}</p>
```

Hii itachanganya vitu vyote vya safu ya `colors` kuwa mfuatano wa herufi ulio na koma kati na kutoa: `<p>green, blue, red</p>`. Baadaye, tutajifunza kuhusu njia nyingine za kawaida za kuonyesha safu za data katika React.

# --instructions--

Postoje sekcije `List` i `ToDo` u editoru koda. Kada prikažete svaki `List` iz sekcije `ToDo`, prođe svojstvo `tasks` koje je postavljeno na vrednost niza vežbi za rad, na primer `["walk dog", "workout"]`. Zatim pronađite ovaj niz `tasks` u sekciji `List` i prikažite njegovu vrednost unutar elementa `p`. Koristite `join(", ")` da prikažete niz `props.tasks` u elementu `p` kao listu od vrednosti razdvojenih zarezima. Današnja lista treba da ima najmanje 2 vežbe, a ona za sutra treba da ima najmanje 3 vežbe.

# --hints--

Dio `ToDo` mora da vrati `div` jednu instancu.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().type() === 'div';
  })()
);
```

Mtoto wa tatu wa sehemu ya `ToDo` inapaswa kuwa mfano wa sehemu ya `List`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().childAt(2).name() === 'List';
  })()
);
```

Peteti element sekcije `ToDo` treba da bude primer sekcije `List`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().childAt(4).name() === 'List';
  })()
);
```

Oba primera sekcije `List` moraju imati svojstvo nazvano `tasks`, a `tasks` mora biti tipa niza podataka.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return (
      Array.isArray(mockedComponent.find('List').get(0).props.tasks) &&
      Array.isArray(mockedComponent.find('List').get(1).props.tasks)
    );
  })()
);
```

Prvi deo `List` koji predstavlja današnju vežbu treba da ima najmanje 2 stavke.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.find('List').get(0).props.tasks.length >= 2;
  })()
);
```

Drugi deo `List` koji predstavlja vežbe za sutra mora imati 3 ili više stavki.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.find('List').get(1).props.tasks.length >= 3;
  })()
);
```

Deo `List` treba da prikaže vrednost iz atributa `tasks` u labeli `p`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return (
      mockedComponent
        .find('p')
        .get(0)
        .props.children.replace(/\s*,\s*/g, ',') ===
        mockedComponent
          .find('List')
          .get(0)
          .props.tasks.join(',')
          .replace(/\s*,\s*/g, ',') &&
      mockedComponent
        .find('p')
        .get(1)
        .props.children.replace(/\s*,\s*/g, ',') ===
        mockedComponent
          .find('List')
          .get(1)
          .props.tasks.join(',')
          .replace(/\s*,\s*/g, ',')
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const List = (props) => {
  { /* Change code below this line */ }
  return <p>{}</p>
  { /* Change code above this line */ }
};

class ToDo extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>To Do Lists</h1>
        <h2>Today</h2>
        { /* Change code below this line */ }
        <List/>
        <h2>Tomorrow</h2>
        <List/>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const List= (props) => {
  return <p>{props.tasks.join(', ')}</p>
};

class ToDo extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>To Do Lists</h1>
        <h2>Today</h2>
        <List tasks={['study', 'exercise']} />
        <h2>Tomorrow</h2>
        <List tasks={['call Sam', 'grocery shopping', 'order tickets']} />
      </div>
    );
  }
};
```
