---
id: 5a24c314108439a4d4036165
title: Koristi React da prikažeš organizovane delove unutra
challengeType: 6
forumTopicId: 301420
dashedName: use-react-to-render-nested-components
---

# --description--

Posledji izazov pokazao je jednostavan način za povezivanje dve komponente, ali postoji mnogo različitih načina na koje možete spojiti deo pomoću React.

Muundo wa sehemu ni mojawapo ya vipengele vyenye nguvu vya React. Unapofanya kazi na React, ni muhimu kuanza kufikiria kuhusu kiolesura chako cha mtumiaji kwa mujibu wa sehemu kama mfano wa App katika changamoto ya mwisho. Unavunja UI yako kuwa vipande vyake vya msingi vya ujenzi, na vipande hivyo vinakuwa sehemu. Hii husaidia kutenganisha msimbo unaohusika na UI kutoka kwa msimbo unaoshughulikia mantiki ya programu yako. Inaweza kurahisisha sana ukuzaji na matengenezo ya miradi tata.

# --instructions--

U uređivaču koda su deklarisane dve funkcionalne komponente, koje se zovu `TypesOfFruit` i `Fruits`. Uzmite deo `TypesOfFruit` i povežite ga, ili *ugnezdite*, unutar dela `Fruits`. Zatim uzmite deo `Fruits` i postavite ga unutar dela `TypesOfFood`. Rezultat bi trebalo da bude komponenta deteta, ugnježđena unutar komponente roditelja, koja je takođe ugnježđena unutar roditeljske komponente svoje vrste!

# --hints--

Dio od `TypesOfFood` treba da vrati jednu funkciju od `div`.

```js
assert(Enzyme.shallow(React.createElement(TypesOfFood)).type() === 'div');
```

Sekcija ``TypesOfFood`` treba da vrati sekciju ``Fruits``.

```js
assert(
  Enzyme.shallow(React.createElement(TypesOfFood)).props().children[1].type
    .name === 'Fruits'
);
```

Sehemu ya `Fruits` inapaswa kurudisha sehemu ya `TypesOfFruit`.

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('h2').html() ===
    '<h2>Fruits:</h2>'
);
```

Sehemu ya `TypesOfFruit` inapaswa kurudisha vipengele vya `h2` na `ul`.

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('ul').text() ===
    'ApplesBlueberriesStrawberriesBananas'
);
```

# --seed--

## --seed-contents--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }

      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }
        <TypesOfFruit />
      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }
        <Fruits />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
