---
id: 5a24c314108439a4d4036185
title: Koristi && za kraće uslovne izjave
challengeType: 6
forumTopicId: 301413
dashedName: use--for-a-more-concise-conditional
---

# --description--

Izrazi za `if/else` radili su u prethodnom izazovu, ali postoji kraći način da se postigne isti rezultat. Zamislite da pratite nekoliko uslova u komponenti i želite da se različiti elementi prikažu u zavisnosti od svakog od ovih uslova. Ako napišete mnogo izjava za `else if` koje vraćaju malo drugačiji UI, možete duplirati kod što ostavlja prostor za greške. Umesto toga, možete koristiti logički operator `&&` da izvršite uslovnu logiku na kraći način. Ovo je moguće jer želite da proverite da li je uslov `true`, i ako jeste, vratite niz znakova. Evo primera:

```jsx
{condition && <p>markup</p>}
```

Ako je `condition` jednako `true`, znaci za pisanje će biti vraćeni. Ako je uslov `false`, operacija će odmah vratiti `false` nakon procene `condition` i neće ništa vraćati. Možete uključiti ove izjave direktno u vaš JSX i kombinovati više uslova pisanjem `&&` nakon svakog. Ovo vam omogućava da rukujete sa složenijom logikom uslova u vašem `render()` bez ponavljanja mnogo koda.

# --instructions--

Tatua tena mfano uliopita, ili `h1` ionyeshwe tu ikiwa `display` ni `true`, lakini tumia kigezo cha mantiki cha `&&` badala ya kauli ya `if/else`.

# --hints--

`MyComponent` Trebalo bi da postoji i da se prikaže.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('MyComponent').length;
  })()
);
```

Kada je `display` postavljeno na `true`, `div`, `button`, na `h1` treba da se prikaže.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: true });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 2 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 1
  );
```

Wakati `display` imewekwa kuwa `false`, ni `div` na `button` tu zinapaswa kuonyeshwa.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: false });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 1 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 0
  );
```

Metoda prikaza treba da koristi logički parametar `&&` da proveri uslov `this.state.display`.

```js
assert(code.includes('&&'));
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
    this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         <h1>Displayed!</h1>
       </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
 this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         {this.state.display && <h1>Displayed!</h1>}
       </div>
    );
  }
};
```
