---
id: 5a24c314108439a4d403618a
title: Tumia Array.map() kuonyesha vipengele kwa njia ya mabadiliko
challengeType: 6
forumTopicId: 301417
dashedName: use-array-map-to-dynamically-render-elements
---

# --description--

Uslovno prikazivanje je važno, ali možda će vam trebati da vaši komponenti prikažu nepoznat broj elemenata. Često u aplikacijama za promene, programer nema način da zna stanje aplikacije do vremena izvršavanja, jer mnoge zavise od interakcije korisnika sa tom aplikacijom. Programeri moraju da napišu svoj kod kako bi rešili to nepoznato stanje unapred. Korišćenje `Array.map()` u React pokazuje ovaj koncept.

Kwa mfano, unaunda programu rahisi ya "Orodha ya Kufanya". Kama mpangaji programu, huna njia ya kujua ni vitu vingapi mtumizi anaweza kuwa navyo kwenye orodha yake. Unahitaji kuandaa sehemu yako kuonyesha idadi sahihi ya vipengele vya orodha kwa njia ya mabadiliko kabla mtu yeyote anaye tumia programu hiyo ataamua kuwa leo ni siku ya kufua nguo.

# --instructions--

Mhariri wa msimbo una sehemu kubwa ya sehemu ya `MyToDoList` tayari imeandaliwa. Baadhi ya msimbo huu unapaswa kuonekana wa kawaida ikiwa umekamilisha changamoto ya fomu iliyodhibitiwa. Utaona `textarea` na `button`, pamoja na njia kadhaa zinazofuatilia hali zao, lakini hakuna kinachoonyeshwa kwenye ukurasa bado.

Ndani ya `constructor`, tengeneza kitu cha `this.state` na fafanua hali mbili: `userInput` inapaswa kuanzishwa kama mfuatano wa herufi tupu, na `toDoList` inapaswa kuanzishwa kama safu ya data tupu. Kisha, futa thamani ya `null` katika njia ya `render()` karibu na mabadiliko ya `items`. Badala yake, tumia map kwenye safu ya `toDoList` iliyohifadhiwa katika hali ya ndani ya sehemu hiyo na onyesha kwa njia ya mabadiliko kipengele cha `li` kwa kila kipengee. Jaribu kuingiza mfuatano wa herufi `eat, code, sleep, repeat` kwenye `textarea`, kisha bofya kitufe na uone kinachotokea.

**Kumbuka:** Unaweza kujua kuwa vipengele vyote vya watoto wa daraja moja vinavyoundwa na operesheni ya map kama hii vinahitaji kupewa sifa ya `key` ya kipekee. Usijali, hili ni somo la changamoto inayofuata.

# --hints--

Sekcija MyToDoList mora da postoji i da se prikaže na stranici.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return mockedComponent.find('MyToDoList').length === 1;
  })()
);
```

Mtoto wa kwanza wa `MyToDoList` inapaswa kuwa kipengele cha `textarea`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(0).type() ===
      'textarea'
    );
  })()
);
```

Mtoto wa pili wa `MyToDoList` inapaswa kuwa kipengele cha `br`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(1).type() === 'br'
    );
  })()
);
```

Mtoto wa tatu wa `MyToDoList` inapaswa kuwa kipengele cha `button`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(2).type() ===
      'button'
    );
  })()
);
```

Stanje `MyToDoList` treba da se inicijalizuje sa `toDoList` kao prazan niz podataka.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    const initialState = mockedComponent.state();
    return (
      Array.isArray(initialState.toDoList) === true &&
      initialState.toDoList.length === 0
    );
  })()
);
```

Hali ya `MyToDoList` inapaswa kuanzishwa na `userInput` kama mfuatano wa herufi tupu.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    const initialState = mockedComponent.state();
    return (
      typeof initialState.userInput === 'string' &&
      initialState.userInput.length === 0
    );
  })()
);
```

Kada se dugme `Create List` klikne, deo `MyToDoList` bi trebalo da prikaže dinamički neuređeni spisak koji sadrži stavku za svaki element liste predmeta razdvojene zarezima unetih u polje `textarea`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
  const simulateChange = (el, value) =>
    el.simulate('change', { target: { value } });
  const state_1 = () => {
    return mockedComponent.find('ul').find('li');
  };
  const setInput = () => {
    return simulateChange(
      mockedComponent.find('textarea'),
      'testA, testB, testC'
    );
  };
  const click = () => {
    return mockedComponent.find('button').simulate('click');
  };
  const state_2 = () => {
    const nodes = mockedComponent.find('ul').find('li');
    return { nodes, text: nodes.reduce((t, n) => t + n.text().trim(), '') };
  };
  const setInput_2 = () => {
    return simulateChange(
      mockedComponent.find('textarea'),
      't1, t2, t3, t4, t5, t6'
    );
  };
  const click_1 = () => {
    return mockedComponent.find('button').simulate('click');
  };
  const state_3 = () => {
    const nodes = mockedComponent.find('ul').find('li');
    return { nodes, text: nodes.reduce((t, n) => t + n.text().trim(), '') };
  };
  const awaited_state_1 = state_1();
  const awaited_setInput = setInput();
  const awaited_click = click();
  const awaited_state_2 = state_2();
  const awaited_setInput_2 = setInput_2();
  const awaited_click_1 = click_1();
  const awaited_state_3 = state_3();
  assert(
    awaited_state_1.length === 0 &&
      awaited_state_2.nodes.length === 3 &&
      awaited_state_3.nodes.length === 6 &&
      awaited_state_2.text === 'testAtestBtestC' &&
      awaited_state_3.text === 't1t2t3t4t5t6'
  );
})();
```

# --seed--

## --seed-contents--

```jsx
const textAreaStyles = {
  width: 235,
  margin: 5
};

class MyToDoList extends React.Component {
  constructor(props) {
    super(props);
    // Change code below this line

    // Change code above this line
    this.handleSubmit = this.handleSubmit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleSubmit() {
    const itemsArray = this.state.userInput.split(',');
    this.setState({
      toDoList: itemsArray
    });
  }
  handleChange(e) {
    this.setState({
      userInput: e.target.value
    });
  }
  render() {
    const items = null; // Change this line
    return (
      <div>
        <textarea
          onChange={this.handleChange}
          value={this.state.userInput}
          style={textAreaStyles}
          placeholder='Separate Items With Commas'
        />
        <br />
        <button onClick={this.handleSubmit}>Create List</button>
        <h1>My "To Do" List:</h1>
        <ul>{items}</ul>
      </div>
    );
  }
}
```

# --solutions--

```jsx
const textAreaStyles = {
  width: 235,
  margin: 5
};

class MyToDoList extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      toDoList: [],
      userInput: ''
    };
    this.handleSubmit = this.handleSubmit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleSubmit() {
    const itemsArray = this.state.userInput.split(',');
    this.setState({
      toDoList: itemsArray
    });
  }
  handleChange(e) {
    this.setState({
      userInput: e.target.value
    });
  }
  render() {
    const items = this.state.toDoList.map((item, i) => {
      return <li key={i}>{item}</li>;
    });
    return (
      <div>
        <textarea
          onChange={this.handleChange}
          value={this.state.userInput}
          style={textAreaStyles}
          placeholder='Separate Items With Commas'
        />
        <br />
        <button onClick={this.handleSubmit}>Create List</button>
        <h1>My "To Do" List:</h1>
        <ul>{items}</ul>
      </div>
    );
  }
}
```
