---
id: 5a24c314108439a4d4036176
title: Koristi stanje da promeniš atribut
challengeType: 6
forumTopicId: 301421
dashedName: use-state-to-toggle-an-element
---

# --description--

Wakati mwingine unaweza kuhitaji kujua hali ya awali unaposasisha hali. Hata hivyo, masasisho ya hali yanaweza kuwa isiyo ya wakati mmoja - hii inamaanisha React inaweza kuunganisha wito nyingi za `setState()` kuwa sasisho moja. Hii inamaanisha huwezi kutegemea thamani ya awali ya `this.state` au `this.props` wakati unahesabu thamani inayofuata. Kwa hivyo, haipaswi kutumia msimbo kama huu:

```jsx
this.setState({
  counter: this.state.counter + this.props.increment
});
```

Umesto toga, trebalo bi da prosledite funkciju ``setState`` koja vam omogućava pristup trenutnom stanju i prosleđenim parametrima. Korišćenje funkcije sa ``setState`` osigurava da radite sa najnovijim vrednostima stanja i prosleđenih parametara. To znači da primer koji je prethodio treba prepisati na sledeći način:

```jsx
this.setState((state, props) => ({
  counter: state.counter + props.increment
}));
```

Takođe možete koristiti tip koji nema `props` ako vam je potrebno samo `state`:

```jsx
this.setState(state => ({
  counter: state.counter + 1
}));
```

Zapamti da moraš da obuhvatiš literalan sadržaj stvari unutar običnih zagrada, inače JavaScript će to smatrati deo koda.

# --instructions--

`MyComponent` ima svojstvo `visibility` koje je inicijalizovano sa `false`. Metoda za prikaz vraća jedan pogled ako je vrednost `visibility` tačna, i drugačiji pogled ako je netačan.

Kwa sasa, hakuna njia ya kusasisha sifa ya `visibility` katika `state` ya sehemu hiyo. Thamani inapaswa kubadilika mfululizo kati ya kweli na potofu. Kuna mshughulikiaji wa bofya kwenye kitufe unaochochea njia ya darasa iitwayo `toggleVisibility()`. Pitisha kitendakazi kwa `setState` ili kufafanua njia hii ili `state` ya `visibility` ibadilike kuwa thamani kinyume wakati njia hiyo inaitwa. Ikiwa `visibility` ni `false`, njia hiyo inaiweka kuwa `true`, na kinyume chake.

Konačno, kliknite na dugme da vidite prikaz te sekcije prema njenom `state`.

**Napomena:** Nemoj zaboraviti da zatvoriš glavnu frazu za `this` na taj način u `constructor`!

# --hints--

`MyComponent` bi trebalo da vrati funkciju iz `div` koja ima `button`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).find('div').find('button')
    .length,
  1
);
```

Hali ya `MyComponent` inapaswa kuanzishwa na sifa ya `visibility` iliyowekwa kuwa `false`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).state('visibility'),
  false
);
```

Klik na element dugmeta treba da promeni svojstvo `visibility` u stanju između `true` i `false` i prikaže element `h1` pod određenim uslovima.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ visibility: false });
    return mockedComponent.state('visibility');
  };
  const second = () => {
    mockedComponent.find('button').simulate('click');
    return mockedComponent.state('visibility') && mockedComponent.find('h1').exists();
  };
  const third = () => {
    mockedComponent.find('button').simulate('click');
    return mockedComponent.state('visibility') && mockedComponent.find('h1').exists();
  };
  const firstValue = first();
  const secondValue = second();
  const thirdValue = third();
  assert(!firstValue && secondValue && !thirdValue);
})();
```

Anonimna funkcija mora biti prosleđena na `setState`.

```js
const paramRegex = '[a-zA-Z$_]\\w*(,[a-zA-Z$_]\\w*)?';
assert(
  new RegExp(
    'this\\.setState\\((function\\(' +
      paramRegex +
      '\\){|([a-zA-Z$_]\\w*|\\(' +
      paramRegex +
      '\\))=>)'
  ).test(__helpers.removeWhiteSpace(code))
);
```

`this` ne bi trebalo koristiti unutar `setState`

```js
assert(!/this\.setState\([^}]*this/.test(code));
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      visibility: false
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    if (this.state.visibility) {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
          <h1>Now you see me!</h1>
        </div>
      );
    } else {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
        </div>
      );
    }
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      visibility: false
    };
    this.toggleVisibility = this.toggleVisibility.bind(this);
  }
  toggleVisibility() {
    this.setState(state => ({
      visibility: !state.visibility
    }));
  }
  render() {
    if (this.state.visibility) {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
          <h1>Now you see me!</h1>
        </div>
      );
    } else {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
        </div>
      );
    }
  }
}
```
