---
id: 5a24c314108439a4d4036161
title: Naučimo o samostalnim (self-closing) tagovima u JSX-u
challengeType: 6
forumTopicId: 301396
dashedName: learn-about-self-closing-jsx-tags
---

# --description--

Sada, vidite kako JSX odstupa od HTML na jedan važan način koristeći `className` naspram `class` za definisanje klase objekata HTML.

Drugi važan način na koji se JSX razlikuje od HTML je koncept samozatvarajućeg taga.

Katika HTML, karibu lebo zote zina lebo ya kufungua na lebo ya kufunga: `<div></div>`; lebo ya kufunga huwa na alama ya mstari wa mbele kabla ya jina la lebo unalofunga. Hata hivyo, kuna matukio maalum katika HTML yanayoitwa <dfn>vipengele vya void</dfn>, au vipengele ambavyo havihitaji lebo ya kufungua na lebo ya kufunga kabla ya kipengele kingine kuanza.

Na primer, oznaka za prekid linije može se napisati kao `<br>` ili kao `<br />`, ali ne bi trebalo da se piše kao `<br></br>`, jer nema nikakvog sadržaja.

U JSX-u, pravilo je malo drugačije. Bilo koji JSX element može se napisati sa samozatvarajućim tagom, i svaki element mora biti zatvoren. Tag za prekid linije, na primer, uvek mora se napisati kao `<br />` da bi bio validan i prevodljiv JSX. `<div>`, s druge strane, može se napisati kao `<div />` ili `<div></div>`. Razlika je u tome što u prvoj verziji sintakse nema načina da se nešto stavi unutar `<div />`. Videćete u budućim izazovima da je ova sintaksa važna prilikom prikazivanja delova React.

# --instructions--

Ispravite greške u kodnom editoru kako bi bio validan JSX i uspešno se kompajlirao/renderovao. Uverite se da niste promenili sadržaj – potrebno je samo zatvoriti tagove tamo gde su potrebni.

# --hints--

Konsanti `JSX` treba da vrati funkciju od `div`.

```js
assert.strictEqual(JSX.type, 'div');
```

`div` inapaswa kuwa na lebo ya `br`.

```js
assert(Enzyme.shallow(JSX).find('br').length === 1);
```

`div` trebalo bi da ima oznaku od `hr`.

```js
assert(Enzyme.shallow(JSX).find('hr').length === 1);
```

# --seed--

## --seed-contents--

```jsx
const JSX = (
  <div>
    <h2>Welcome to React!</h2> <br >
    <p>Be sure to close all tags!</p>
    <hr >
  </div>
);
```

# --solutions--

```jsx
const JSX = (
<div>
  <h2>Welcome to React!</h2> <br />
  <p>Be sure to close all tags!</p>
  <hr />
</div>
);
```
