---
id: 5a24c314108439a4d4036162
title: Kreirajte deo bez funkcionalnosti
challengeType: 6
forumTopicId: 301392
dashedName: create-a-stateless-functional-component
---

# --description--

Komponente su osnova za React. Sve u React je komponenta, i ovde ćete naučiti kako da napravite jednu.

Postoje dva načina za kreiranje komponente `React`. Prvi način je korišćenje funkcije `JavaScript`. Opisivanje komponente na ovaj način stvara *komponentu bez stanja funkcije*. Koncept stanja u aplikaciji će biti diskutovan u narednim izazovima. Trenutno, zamislite komponentu bez stanja kao nešto što može da prima i prikazuje podatke, ali ne kontroliše niti prati promene tih podataka. (Razgovaraćemo o drugom načinu kreiranja komponente `React` u sledećem izazovu.)

Da biste kreirali komponentu sa funkcijom, samo napišite funkciju JavaScript koja vraća ili JSX ili `null`. Važno je napomenuti da React zahteva da ime vaše funkcije počinje velikim slovom. Evo primera komponente bez funkcije koja dodaje klasu HTML u JSX:

```jsx
const DemoComponent = function() {
  return (
    <div className='customClass' />
  );
};
```

Baada ya kubadilishwa, `<div>` itakuwa na darasa la CSS la `customClass`.

Budući da JSX komponenta predstavlja HTML, možete spojiti nekoliko komponenti kako biste kreirali složeniju stranicu za HTML. Ovo je jedna od glavnih prednosti arhitekture komponenti koju pruža React. Omogućava vam da kreirate svoj UI iz mnogo različitih, razdvojenih komponenti. To olakšava izgradnju i održavanje složenih korisničkih interfejsa.

# --instructions--

Mhariri wa msimbo una kitendakazi kinachoitwa `MyComponent`. Kamilisha kitendakazi hiki ili kirejeshe kipengele kimoja cha `div` chenye mfuatano wa herufi fulani.

**Napomena:** Tekst se smatra potomkom elementa `div`, stoga neće moći da koristi samozatvarajuću oznaku.

# --hints--

`MyComponent` Trebalo bi da ponovo učita JSX.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.length === 1;
  })()
);
```

`MyComponent` inapaswa kurejesha kipengele cha `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Kipengele cha `div` kinapaswa kuwa na mfuatano wa herufi.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').text() !== '';
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const MyComponent = function() {
  // Change code below this line



  // Change code above this line
}
```

# --solutions--

```jsx
const MyComponent = function() {
  // Change code below this line
  return (
    <div>
      Demo Solution
    </div>
  );
  // Change code above this line
}
```
