---
id: 5a24c314108439a4d4036182
title: Dodajte inline stilove u React
challengeType: 6
forumTopicId: 301378
dashedName: add-inline-styles-in-react
---

# --description--

Možda ste primetili u prethodnom izazovu da je bilo nekoliko razlika u sintaksi iz inline stilova za HTML, kao i svojstvo `style` postavljeno kao nešto od JavaScript. Prvo, imena nekih svojstava stila za CSS koriste camel case. Na primer, poslednji izazov je postavio veličinu fonta za `fontSize` umesto `font-size`. Reči sa kistavim nastavkom kao što je `font-size` nisu validna sintaksa za svojstva elemenata JavaScript, pa zato React koristi camel case. Po pravilu, sva svojstva stila sa kistavim nastavkom pišu se koristeći camel case u JSX.

Sve jedinice dužine za vrednosti svojstava (kao što su `height`, `width`, i `fontSize`) smatraju se da su u `px` osim ako je drugačije navedeno. Ako želite da koristite `em`, na primer, morate obuhvatiti vrednost i jedinice navodnicima, kao što je `{fontSize: "4em"}`. Osim dužinskih vrednosti koje uzimaju `px` kao podrazumevano, sve ostale vrednosti svojstava moraju biti obuhvaćene navodnicima.

# --instructions--

Ako imate veliki set stilova, možete postaviti `object` stila u trajni parametar kako biste održali vaš kod organizovanim. Deklarišite svoj stilski parametar kao globalni parametar na fajlu. Inicijalizujte parametar `styles` i postavite `object` sa tri atributa stila i njihovim vrednostima. Definišite `div` boju od `purple`, veličinu fonta od `40` i ivicu od `2px solid purple`. Zatim postavite atribut `style` jednak parametru `styles`.

# --hints--

Kriterijum za `styles` treba da bude `object` sa tri osobine.

```js
assert(Object.keys(styles).length === 3);
```

Parametar za `styles` mora imati svojstvo `color` postavljeno na vrednost `purple`.

```js
assert(styles.color === 'purple');
```

Kriterijum za `styles` mora imati svojstvo `fontSize` koje je postavljeno na vrednost `40`.

```js
assert(styles.fontSize == 40);
```

Parametar za `styles` treba da ima svojstvo `border` koje je postavljeno na vrednost `2px solid purple`.

```js
assert(styles.border === '2px solid purple');
```

Sekcija treba da prikaže element `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.shallow(React.createElement(Colorful));
    return mockedComponent.type() === 'div';
  })()
);
```

Element za `div` treba da ima stilove definisane od objekta `styles`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.shallow(React.createElement(Colorful));
    return (
      mockedComponent.props().style.color === 'purple' &&
      mockedComponent.props().style.fontSize == 40 &&
      mockedComponent.props().style.border === '2px solid purple'
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Change code above this line
class Colorful extends React.Component {
  render() {
    // Change code below this line
    return (
      <div style={{color: "yellow", fontSize: 24}}>Style Me!</div>
    );
    // Change code above this line
  }
};
```

# --solutions--

```jsx
const styles = {
  color: "purple",
  fontSize: 40,
  border: "2px solid purple"
};
// Change code above this line
class Colorful extends React.Component {
  render() {
    // Change code below this line
    return (
      <div style={styles}>Style Me!</div>
    );
    // Change code above this line
  }
};
```
