---
id: 5a24c314108439a4d4036172
title: Prikaži stanje u korisničkom interfejsu na drugi način
challengeType: 6
forumTopicId: 301408
dashedName: render-state-in-the-user-interface-another-way
---

# --description--

Postoji drugi način da se pristupi `state` u sekciji. U načinu `render()`, pre deklaracije `return`, možete direktno napisati JavaScript. Na primer, možete deklarisati funkcije, dohvatiti podatke iz `state` ili `props`, izvršiti proračune na ovim podacima i tako dalje. Zatim, možete postaviti bilo koju vrednost za parametre, koji su dostupni u deklaraciji `return`.

# --instructions--

U toku rendera `MyComponent`, definišite `const` nazvano `name` i postavite ga jednak vrednosti imena u `state` sekcije. Budući da možete direktno napisati JavaScript u ovu kod sekciju, ne morate obuhvatati ovu referencu uglastim zagradama.

Sledeće, u izveštaju return, prikažite ovu vrednost u labeli `h1` koristeći parametar `name`. Zapamtite, morate koristiti JSX sintaksu (kvadesne zagrade za JavaScript) u izveštaju return.

# --hints--

`MyComponent` treba imati `name` sa vrednošću `freeCodeCamp` sačuvanu u svom stanju.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'freeCodeCamp'
);
```

`MyComponent` treba da prikaže naslov sekcije za `h1`, okružen jednim `div`.

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.mount(React.createElement(MyComponent)).html()
  )
);
```

Etiketa za `h1` koja je nacrtana mora da uključuje referencu na `{name}`.

```js
assert(/<h1>\n*\s*\{\s*name\s*\}\s*\n*<\/h1>/.test(code));
```

Zaglavlje dela `h1` koje je prikazano treba da ima tekst koji se prikazuje iz stanja dela.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  assert(firstValue === '<div><h1>TestName</h1></div>');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    // Change code below this line

    // Change code above this line
    return (
      <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    // Change code below this line
    const name = this.state.name;
    // Change code above this line
    return (
      <div>
        { /* Change code below this line */ }
        <h1>{name}</h1>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
