---
id: 5a24c314108439a4d403616e
title: Pristup parametrima korišćenjem this.props
challengeType: 6
forumTopicId: 301375
dashedName: access-props-using-this-props
---

# --description--

Nedavni izazovi opisali su osnovne načine prosleđivanja parametara podklasama objekata. Ali šta ako je podklasa objekta kojoj prosleđujete parametre deo klase ES6, umesto dela koji nije funkcionalan? Deo klase objekta ES6 koristi malo drugačiji ugovor da bi dobio pristup prosleđenim parametrima.

Kila unaporejelea sehemu ya darasa la vitu ndani yake yenyewe, unatumia neno kuu `this`. Ili kupata upatikanaji wa vigezo vinavyopitishwa ndani ya sehemu ya darasa la vitu, unaweka msimbo unaotumia kupata upatikanaji huo mbele ya `this`. Kwa mfano, ikiwa sehemu ya darasa la vitu la ES6 ina kigezo kinachoitwa `data`, unaandika `{this.props.data}` katika JSX.

# --instructions--

Pokaži primer dela `Welcome` u glavnom delu `App`. Ovde, sa `Welcome` proveri parametar `name` i postavi mu vrednost niza karaktera. Unutar poddela, `Welcome`, dođi do pristupa parametru `name` unutar oznaka `strong`.

# --hints--

Deo `App` treba da vrati jedan element od `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Deca od `App` treba da budu deo `Welcome`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return (
      mockedComponent.children().childAt(0).name() === 'Welcome'
    );
  })()
);
```

Deo `Welcome` treba da ima parametar nazvan `name`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return mockedComponent.find('Welcome').props().name;
  })()
);
```

Dio `Welcome` treba da prikaže niz karaktera koji je prebačen kao parametar za `name` unutar oznaka od `strong`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return (
      mockedComponent.find('strong').text() ===
      mockedComponent.find('Welcome').props().name
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class App extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
            { /* Change code below this line */ }
            <Welcome />
            { /* Change code above this line */ }
        </div>
    );
  }
};

class Welcome extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
          { /* Change code below this line */ }
          <p>Hello, <strong></strong>!</p>
          { /* Change code above this line */ }
        </div>
    );
  }
};
```

# --solutions--

```jsx
class Welcome extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
          { /* Change code below this line */ }
          <p>Hello, <strong>{this.props.name}</strong>!</p>
          { /* Change code above this line */ }
        </div>
    );
  }
};

class App extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
            { /* Change code below this line */ }
            <Welcome name="Quincy"/>
            { /* Change code above this line */ }
        </div>
    );
  }
};
```
