---
id: 5a24c314108439a4d403618d
title: Prikaži React na serveru koristeći renderToString
challengeType: 6
forumTopicId: 301407
dashedName: render-react-on-the-server-with-rendertostring
---

# --description--

Do sada si prikazivao delove React klijentu. Obično je ovo ono što ćeš raditi svaki put. Međutim, postoje određene upotrebe gde je odgovarajuće prikazati deo React na serveru. Pošto je React biblioteka pogleda za JavaScript i možeš pokrenuti JavaScript na serveru koristeći Node, ovo je moguće. Zapravo, React pruža način za `renderToString()` koji možeš koristiti za ovu svrhu.

Postoji dva glavna razloga zašto se prikazivanje na serveru može koristiti u stvarnoj aplikaciji. Prvo, bez toga, vaše aplikacije za React bi imale fajl od HTML sa nedostatkom i veliku grupu JavaScript kada se prvi put preuzmu u pregledaču. Ovo možda neće biti dobro za pretražne motore koji pokušavaju da indeksiraju sadržaj na vašim stranicama kako bi ljudi mogli da ga vide. Ako prikažete početne oznake za HTML na serveru i pošaljete ih klijentu, inicijalni preuzimanje stranice će imati sve oznake stranica koje pretražni motori mogu da analiziraju. Drugo, ovo stvara brže iskustvo učitavanja početne stranice jer je prikazani HTML manji od celokupnog koda programa JavaScript. React će i dalje biti u stanju da prepozna vaš program i upravlja njime nakon inicijalnog preuzimanja.

# --instructions--

Putanja `renderToString()` dostupna je na `ReactDOMServer`, koja postoji ovde kao globalni objekat. Ova putanja prima jedan argument, koji je svojstvo React. Koristite ovo da pokažete `App` kao niz znakova.

# --hints--

Dio `App` treba da se prikaže kao niz slova koristeći `ReactDOMServer.renderToString`.

```js
  assert(
    code.replace(/ /g, '')
      .includes('ReactDOMServer.renderToString(<App/>)') &&
      Enzyme.mount(React.createElement(App)).children().name() === 'div'
  );
```

# --before-all--

```js
var ReactDOMServer = { renderToString(x) { return null; } };
```

# --seed--

## --seed-contents--

```jsx
class App extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <div/>
  }
};

// Change code below this line
```

# --solutions--

```jsx
class App extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <div/>
  }
};

// Change code below this line
ReactDOMServer.renderToString(<App/>);
```
