---
id: 56533eb9ac21ba0edf2244cf
title: Izgradite kolekciju zapisa
challengeType: 26
dashedName: build-a-record-collection
---

# --description--

Kreirate funkciju koja pomaže u održavanju kolekcije muzičkih albuma. Kolekcija je organizovana kao objekat koji sadrži više albuma, a svaki od njih je takođe objekat. Svaki album je predstavljen u kolekciji jedinstvenim ``id`` kao ime svojstva (property name). Unutar svakog objekta albuma postoje razna svojstva koja opisuju informacije o albumu. Nisu svi albumi kompletno informisani.

Funkcija ``updateRecords`` prima 4 argumenta predstavljena sledećim parametrima funkcije:

-   ``records`` - objekat koji sadrži više pojedinačnih albuma
-   ``id`` - broj koji predstavlja određeni album u objektu ``records``
-   ``prop`` - niz znakova (string) koji predstavlja ime svojstva albuma koje treba ažurirati
-   ``value`` - niz znakova (string) koji sadrži informacije korišćene za ažuriranje svojstva albuma

**Cilj:** Ispunite korisničke priče ispod i osigurajte da se svi testovi prođu kako biste završili laboratoriju.

**Korisničke priče:**

1. Vaša funkcija mora uvek vratiti ceo objekat ``records``.
2. Ako je ``value`` prazan niz znakova (string), obrišite dato svojstvo ``prop`` sa albuma.
3. Ako ``prop`` nije ``tracks`` i ``value`` nije prazan niz znakova (string), dodelite ``value`` na to albumovo ``prop``.
4. Ako je ``prop`` jednako ``tracks`` i ``value`` nije prazan niz znakova (string), ali album nema svojstvo ``tracks``, kreirajte prazan niz i dodajte ``value`` u njega.
5. Ako je `prop` jednako ``tracks`` i ``value`` nije prazan niz znakova (string), dodajte ``value`` na kraj postojećeg niza ``tracks`` albuma.

**Napomena:** Za testove koristi se kopija objekta ``recordCollection``. Vaša funkcija ne bi trebalo direktno da se odnosi na objekat ``recordCollection``, već samo na parametar funkcije.

# --before-each--

```js
const _recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};
```

# --hints--

You should have a `updateRecords` function.

```js
assert.isFunction(updateRecords);
```

After `updateRecords(recordCollection, 5439, "artist", "ABBA")`, `artist` should be the string `ABBA`

```js
assert.equal(
  updateRecords(_recordCollection, 5439, 'artist', 'ABBA')[5439]['artist'],
    'ABBA'
);
```

After `updateRecords(recordCollection, 5439, "tracks", "Take a Chance on Me")`, `tracks` should have the string `Take a Chance on Me` as the last and only element.

```js
updateRecords(_recordCollection, 5439, 'tracks', 'Take a Chance on Me');
assert.lengthOf(_recordCollection[5439]['tracks'], 1);
assert.equal(_recordCollection[5439]['tracks'].pop(), 'Take a Chance on Me');
```

After `updateRecords(recordCollection, 2548, "artist", "")`, `artist` should not be set

```js
updateRecords(_recordCollection, 2548, 'artist', '');
assert.isFalse(_recordCollection[2548].hasOwnProperty('artist'));
```

After `updateRecords(recordCollection, 1245, "tracks", "Addicted to Love")`, `tracks` should have the string `Addicted to Love` as the last element.

```js
assert.equal(
  updateRecords(_recordCollection, 1245, 'tracks', 'Addicted to Love')[1245][
    'tracks'
  ].pop(), 'Addicted to Love'
);
```

After `updateRecords(recordCollection, 2468, "tracks", "Free")`, `tracks` should have the string `1999` as the first element.

```js
assert.equal(
  updateRecords(_recordCollection, 2468, 'tracks', 'Free')[2468][
    'tracks'
  ][0], '1999'
);
```

After `updateRecords(recordCollection, 2548, "tracks", "")`, `tracks` should not be set

```js
updateRecords(_recordCollection, 2548, 'tracks', '');
assert.isFalse(_recordCollection[2548].hasOwnProperty('tracks'));
```

After `updateRecords(recordCollection, 1245, "albumTitle", "Riptide")`, `albumTitle` should be the string `Riptide`

```js
assert.equal(
  updateRecords(_recordCollection, 1245, 'albumTitle', 'Riptide')[1245][
    'albumTitle'
  ], 'Riptide'
);
```

# --seed--

## --seed-contents--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};



```

# --solutions--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};

function updateRecords(records, id, prop, value) {
  if (value === '') delete records[id][prop];
  else if (prop === 'tracks') {
    records[id][prop] = records[id][prop] || [];
    records[id][prop].push(value);
  } else {
    records[id][prop] = value;
  }

  return records;
}
```
