id: 681cb1b0dab50c87ddb2e51a
title: "Izazov 9: Zbir kvadrata"
challengeType: 29
dashedName: challenge-9
---

# --description--

Ako vam je dat pozitivan ceo broj do 1.000, vratite zbir svih celih brojeva podignutih na kvadrat od 1 do tog broja.

# --hints--

`sum_of_squares(5)` bi trebalo da vrati `55`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(5), 55)`)
}})
```

`sum_of_squares(10)` bi trebalo da vrati `385`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(10), 385)`)
}})
```

`sum_of_squares(25)` bi trebalo da vrati `5525`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(25), 5525)`)
}})
```

`sum_of_squares(500)` bi trebalo da vrati `41791750`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(500), 41791750)`)
}})
```

`sum_of_squares(1000)` bi trebalo da vrati `333833500`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_of_squares(1000), 333833500)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_of_squares(n):

    return n
```

# --solutions--

```py
def sum_of_squares(n):
    sum = ((n)*(n+1)*(2*n+1))//6
    return sum
```