id: 681cb1b0dab50c87ddb2e518
title: "Izazov 7: Ciljani zbir"
challengeType: 29
dashedName: challenge-7
---

# --description--

Ako imate niz brojeva i specifičnu ciljnu vrednost, pronađite dva različita broja u nizu čija je suma jednaka toj ciljnoj vrednosti. Vratite niz koji sadrži indekse (pozicije) tih dva broja, ili `"Target not found"` ako se nijedna par brojeva ne sabira do ciljne vrednosti.

- Niz koji se vraća mora sadržati indekse u rastućem redosledu.

# --hints--

`find_target([2, 7, 11, 15], 9)` bi trebalo da vrati `[0, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([2, 7, 11, 15], 9), [0, 1])`)
}})
```

`find_target([3, 2, 4, 5], 6)` bi trebalo da vrati `[1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([3, 2, 4, 5], 6), [1, 2])`)
}})
```

`find_target([1, 3, 5, 6, 7, 8], 15)` bi trebalo da vrati `[4, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([1, 3, 5, 6, 7, 8], 15), [4, 5])`)
}})
```

`find_target([1, 3, 5, 7], 14)` bi trebalo da vrati `'Target not found'`.

```js

({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_target([1, 3, 5, 7], 14), "Target not found")`)
}})
```

# --seed--

## --seed-contents--

```py
def find_target(arr, target):

    return arr
```

# --solutions--

```py
def find_target(arr, target):
    for i in range(len(arr)):
        for j in range(i + 1, len(arr)):
            if arr[i] + arr[j] == target:
                return [i, j]
    return 'Target not found'
```