id: 681cb1b0dab50c87ddb2e519
title: "Challenge 8: Factorializer"
challengeType: 29
dashedName: challenge-8
---

# --description--

Ako vam je dat ceo broj od nule do dvadeset, vratite faktorijalu tog broja. Faktorijal broja je proizvod svih brojeva između 1 i datog broja.

- Faktorijal nule je 1.

# --hints--

`factorial(0)` bi trebalo da vrati `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(0), 1)`)
}})
```

`factorial(5)` bi trebalo da vrati `120`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(5), 120)`)
}})
```

`factorial(20)` bi trebalo da vrati `2432902008176640000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(factorial(20), 2432902008176640000)`)
}})
```

# --seed--

## --seed-contents--

```py
def factorial(n):

    return n
```

# --solutions--

```py
def factorial(n):
    result = 1
    for i in range(1, n + 1):
        result *= i
    return result
```