id: 6821ebda237de8297eaee792
title: "Challenge 16: Reverse Parenthesis"
challengeType: 29
dashedName: challenge-16
---

# --description--

Kada dobijete niz slova koji sadrži pravilno raspoređene zagrade, vratite dešifrovisanu verziju tog niza koristeći sledeća pravila:

- Sva slova unutar svake pare zagrada moraju biti obrnuta.
- Zagrade se moraju ukloniti iz konačnog rezultata.
- Ako postoje ugnježdene zagrade, najunutrašnja para mora prvo biti obrnuta, a zatim njen rezultat uključena u obrnivanje spoljne pare.
- Pretpostavlja se da su sve zagrade pravilno i tačno raspoređene.

# --hints--

`decode("(f(b(dc)e)a)")` should return `"abcdef"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("(f(b(dc)e)a)"), "abcdef")`)
}})
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` should return `"Can you read this?"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?")`)
}})
```

`decode("f(Ce(re))o((e(aC)m)d)p")` should return `"freeCodeCamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(s):

    return s
```

# --solutions--

```py
def decode(s):
    while ')' in s:
        close_index = s.index(')')
        open_index = s.rindex('(', 0, close_index)
        inner = s[open_index + 1:close_index][::-1]
        s = s[:open_index] + inner + s[close_index + 1:]
    return s
```