id: 6821ebea237de8297eaee795
title: "Izazov 19: Svetlost sveće"
challengeType: 29
dashedName: challenge-19
---

# --description--

Dato vam je potpuno brojanje koje predstavlja početni broj sveća i potpuno brojanje koliko sveća je potrebno za napraviti novu sveću, vratite ukupan broj sveća koje ćete koristiti nakon što napravite i sagorete sve što možete.

Na primer, ako dobijete 7 sveća, a za pravljenje nove sveće potrebne su 2 sagorene sveće:

1. Sagorete 7 sveća i ostaje vam 7 ostataka,
2. Vratite 6 ostataka u 3 nove sveće (ostaje 1 sveća),
3. Sagorete 3 sveće i ostaje vam još 3 ostatka (ukupno 4),
4. Vratite 4 ostatka u 2 nove sveće,
5. Sagorete 2 sveće i ostaje vam 2 ostatka,
6. Vratite 2 ostataka u 1 novu sveću,
7. Sagorete 1 sveću.

Ukupno je sagorelo 13 sveća u ovom primeru.

# --hints--

`burn_candles(7, 2)` should return `13`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(7, 2), 13)`)
}})
```

`burn_candles(10, 5)` should return `12`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(10, 5), 12)`)
}})
```

`burn_candles(20, 3)` should return `29`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(20, 3), 29)`)
}})
```

`burn_candles(17, 4)` should return `22`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(17, 4), 22)`)
}})
```

`burn_candles(2345, 3)` should return `3517`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(burn_candles(2345, 3), 3517)`)
}})
```

# --seed--

## --seed-contents--

```py
def burn_candles(candles, leftovers_needed):

    return candles
```

# --solutions--

```py
def burn_candles(candles, leftovers_needed):
    total_burned = 0
    unused_leftovers = 0

    while candles > 0:
        total_burned += candles
        leftovers = candles + unused_leftovers
        candles = leftovers // leftovers_needed
        unused_leftovers = leftovers % leftovers_needed

    return total_burned
```