---
id: 6821ebce237de8297eaee790
title: "Challenge 14: Character Battle"
challengeType: 29
dashedName: challenge-14
---

# --description--

Dati je niz od dve slova koji predstavljaju vašu vojsku i protivničku vojsku. Svako slovo iz vaše vojske se bori sa slovom koje se nalazi na istoj poziciji u protivničkoj vojskoj koristeći sledeće pravilo:

- Slova `a-z` imaju snagu od `1-26`, respektivno.
- Slova `A-Z` imaju snagu od `27-52`, respektivno.
- Brojevi `0-9` imaju snagu jednaku njihovoj nominalnoj vrednosti.
- Sva druga slova imaju vrednost nula.
- Svako slovo može učestvovati samo u jednom бою.

U svakoj bitci pobeđuje slovo sa većom snagom. Vojska koja osvoji više bitaka pobediće. Vratite sledeće vrednosti:

- `"Opponent retreated"` ako vaša vojska ima više slova nego protivnička vojska.
- `"We retreated"` ako protivnička vojska ima više slova nego vaša.
- `"We won"` ako je vaša vojska osvojila više bitaka.
- `"We lost"` ako je protivnička vojska osvojila više bitaka.
- `"It was a tie"` ako su obe vojske osvojile isti broj bitaka.

# --hints--

`battle("Hello", "World")` should return `"We lost"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Hello", "World"), "We lost")`)
}})
```

`battle("pizza", "salad")` should return `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("pizza", "salad"), "We won")`)
}})
```

`battle("C@T5", "D0G$")` should return `"We won"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("C@T5", "D0G$"), "We won")`)
}})
```

`battle("kn!ght", "orc")` should return `"Opponent retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("kn!ght", "orc"), "Opponent retreated")`)
}})
```

`battle("PC", "Mac")` should return `"We retreated"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("PC", "Mac"), "We retreated")`)
}})
```

`battle("Wizards", "Dragons")` should return `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Wizards", "Dragons"), "It was a tie")`)
}})
```

`battle("Mr. Smith", "Dr. Jones")` should return `"It was a tie"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(battle("Mr. Smith", "Dr. Jones"), "It was a tie")`)
}})
```

# --seed--

## --seed-contents--

```py
def battle(my_army, opposing_army):

    return my_army
```

# --solutions--

```py
def get_strength(soldier):
    letters = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ'

    if soldier.isdigit():
        return int(soldier)
    elif soldier in letters:
        return letters.index(soldier) + 1
    else:
        return 0

def battle(my_army, opposing_army):
    if len(my_army) > len(opposing_army):
        return 'Opponent retreated'
    if len(opposing_army) > len(my_army):
        return 'We retreated'

    my_wins = 0
    their_wins = 0

    for my_soldier, their_soldier in zip(my_army, opposing_army):
        my_strength = get_strength(my_soldier)
        their_strength = get_strength(their_soldier)

        if my_strength > their_strength:
            my_wins += 1
        elif their_strength > my_strength:
            their_wins += 1

    if my_wins > their_wins:
        return 'We won'
    elif their_wins > my_wins:
        return 'We lost'
    else:
        return 'It was a tie'
```