id: 68216ef80f957572e7c340c5
title: "Challenge 12: Message Decoder"
challengeType: 29
dashedName: challenge-12
---

# --description--

Kada dobijete sekvencu slova tajne poruke i ceo broj koji predstavlja koliko je slova korišćeno za pomeranje/šifrovanje poruke radi sakrivanja, vratite dešifrovovan sekvencu slova.

- Pozitivan broj znači da je poruka pomerena napred u alfabeti.
- Negativan broj znači da je poruka pomerena unazad u alfabeti.
- Razlika između velikih i malih slova je važna; dešifrovovana slova moraju sačuvati razliku između velikog ili malog slova kao što je bilo u originalnoj poruci.
- Nealfabetički znakovi ne bi trebalo dešifrovovati.

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` should return `"This is a secret message."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.")`)
}})
```

`decode("Byffi Qilfx!", 20)` should return `"Hello World!"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Byffi Qilfx!", 20), "Hello World!")`)
}})
```

`decode("Zqd xnt njzx?", -1)` should return `"Are you okay?"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Zqd xnt njzx?", -1), "Are you okay?")`)
}})
```

`decode("oannLxmnLjvy", 9)` should return `"freeCodeCamp"`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("oannLxmnLjvy", 9), "freeCodeCamp")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message, shift):

    return message
```

# --solutions--

```py
def decode(message, shift):
    decoded_message = []
    for char in message:
        if char.isalpha():
            base = ord('a') if char.islower() else ord('A')
            new_char = chr((ord(char) - base - shift) % 26 + base)
            decoded_message.append(new_char)
        else:
            decoded_message.append(char)
    return ''.join(decoded_message)
```