id: 6821ebc9237de8297eaee78f
title: "Challenge 13: Unnatural Prime"
challengeType: 29
dashedName: challenge-13
---

# --description--

Za dato celo число, utvrdite da li je ono primbroj ili negativan primbroj.

- Primbroj je celo pozitivno число veće od 1 koje se deli samo sa 1 i samim sobom.
- Negativan primbroj je negativna vrednost pozitivnog primbroja.
- `1` i `0` se ne smatraju primbrojima.

# --hints--

`is_unnatural_prime(1)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(1), False)`)
}})
```

`is_unnatural_prime(-1)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(-1), False)`)
}})
```

`is_unnatural_prime(19)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(19), True)`)
}})
```

`is_unnatural_prime(-23)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(-23), True)`)
}})
```

`is_unnatural_prime(0)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(0), False)`)
}})
```

`is_unnatural_prime(97)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(97), True)`)
}})
```

`is_unnatural_prime(-61)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(-61), True)`)
}})
```

`is_unnatural_prime(99)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(99), False)`)
}})
```

`is_unnatural_prime(-44)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_unnatural_prime(-44), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_unnatural_prime(n):

    return n
```

# --solutions--

```py
def is_unnatural_prime(n):
    abs_n = abs(n)

    if abs_n <= 1:
        return False

    for i in range(2, int(abs_n ** 0.5) + 1):
        if abs_n % i == 0:
            return False

    return True
```