id: 6821ebd4237de8297eaee791
title: "Izazov 15: camelCase"
challengeType: 29
dashedName: challenge-15
---

# --description--

Dato vam je niz slova, vratite mu verziju u stilu camelCase koristeći sledeća pravila:

- Reči u argumentu niza slova su razdvojene jednom ili više karaktera iz sledećeg skupa: razmak (` `), crtica (`-`) ili donji vek (`_`). Svaka sekvenca ovoga se tretira kao granica reči.
- Prva reč mora biti u malim slovima (svih mala).
- Svaka sledeća reč mora početi velikim slovom, a ostatak mora biti u malim slovima.
- Svi razmaci i separatori moraju biti uklonjeni.

# --hints--

`to_camel_case("hello world")` should return `"helloWorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("hello world"), "helloWorld")`)
}})
```

`to_camel_case("HELLO WORLD")` should return `"helloWorld"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("HELLO WORLD"), "helloWorld")`)
}})
```

`to_camel_case("secret agent-X")` should return `"secretAgentX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("secret agent-X"), "secretAgentX")`)
}})
```

`to_camel_case("FREE cODE cAMP")` should return `"freeCodeCamp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("FREE cODE cAMP"), "freeCodeCamp")`)
}})
```

`to_camel_case("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")` should return `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_camel_case("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_camel_case(s):

    return s
```

# --solutions--

```py
import re
def to_camel_case(s):
    words = re.split(r'[_\- ]+', s)

    camel = [
        words[0].lower() if words else ''
    ] + [
        word.capitalize() for word in words[1:]
    ]

    return ''.join(camel)
```