---
id: 6821ebf3237de8297eaee797
title: "Vežba 21: Generator Hex"
challengeType: 29
dashedName: challenge-21
---

# --description--

Kada mu započete sa imenom CSS boje, napišite nasumični heksadecimalni (hex) kod boje koji ima dominantnu boju u datoj boji.

- Funkcija treba da obrađuje `"red"`, `"green"` ili `"blue"` kao ulazni argument.
- Ako ulaz nije jedna od navedenih, funkcija bi trebalo da vrati `"Invalid color"`.
- Funkcija bi trebalo da vrati heksadecimalni kod boje sa šest nasumičnih slova gde je vrednost boje iz ulaza veća od bilo koje druge.
- Primer validnih rezultata za dati ulaz:

| Ulaz   | Rezultat   |
|---------|----------|
| `"red"`   | `"FF0000"` |
| `"green"` | `"00FF00"` |
| `"blue"`  | `"0000FF"` |

# --hints--

`generate_hex("yellow")` bi trebalo da vrati `"Invalid color"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_hex("yellow"), "Invalid color")`);
}})
```

`generate_hex("red")` bi trebalo da vratiti niz od šest slova.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(len(generate_hex("red")), 6)`);
}})
```

`generate_hex("red")` bi trebalo da vrati validan heksadecimalni kod sa šest slova.

```js
({test: () => { runPython(`
from unittest import TestCase

hex = generate_hex("red").upper()
is_valid_hex = len(hex) == 6 and all(c in "0123456789ABCDEF" for c in hex)
TestCase().assertTrue(is_valid_hex)`)
}})
```

`generate_hex("red")` bi trebalo da vrati validan heksadecimalni kod sa šest slova gde je vrednost crvene boje veća od bilo koje druge.

```js
({test: () => { runPython(`
from unittest import TestCase

hex = generate_hex("red").upper()
is_valid_hex = len(hex) == 6 and all(c in "0123456789ABCDEF" for c in hex)
TestCase().assertTrue(is_valid_hex)

r = int(hex[:2], 16)
g = int(hex[2:4], 16)
b = int(hex[4:], 16)

TestCase().assertGreater(r, g)
TestCase().assertGreater(r, b)`)
}})
```

Pozivanje `generate_hex("red")` dva puta bi trebalo da vrati dve različite heksadecimalne boje gde je crvena dominantna.

```js
({test: () => { runPython(`
from unittest import TestCase

hex1 = generate_hex("red").upper()
is_valid_hex1 = len(hex1) == 6 and all(c in "0123456789ABCDEF" for c in hex1)
TestCase().assertTrue(is_valid_hex1)

r1 = int(hex1[:2], 16)
g1 = int(hex1[2:4], 16)
b1 = int(hex1[4:], 16)

TestCase().assertGreater(r1, g1)
TestCase().assertGreater(r1, b1)

hex2 = generate_hex("red").upper()
is_valid_hex2 = len(hex2) == 6 and all(c in "0123456789ABCDEF" for c in hex2)
TestCase().assertTrue(is_valid_hex2)

r2 = int(hex2[:2], 16)
g2 = int(hex2[2:4], 16)
b2 = int(hex2[4:], 16)

TestCase().assertGreater(r2, g2)
TestCase().assertGreater(r2, b2)
TestCase().assertNotEqual(hex1, hex2)`);
}})
```

Pozivanje `generate_hex("green")` dva puta bi trebalo da vrati dve različite heksadecimalne boje gde je zelena dominantna.

```js
({test: () => { runPython(`
from unittest import TestCase

hex1 = generate_hex("green").upper()
is_valid_hex1 = len(hex1) == 6 and all(c in "0123456789ABCDEF" for c in hex1)
TestCase().assertTrue(is_valid_hex1)

r1 = int(hex1[:2], 16)
g1 = int(hex1[2:4], 16)
b1 = int(hex1[4:], 16)

TestCase().assertGreater(g1, r1)
TestCase().assertGreater(g1, b1)

hex2 = generate_hex("green").upper()
is_valid_hex2 = len(hex2) == 6 and all(c in "0123456789ABCDEF" for c in hex2)
TestCase().assertTrue(is_valid_hex2)

r2 = int(hex2[:2], 16)
g2 = int(hex2[2:4], 16)
b2 = int(hex2[4:], 16)

TestCase().assertGreater(g2, r2)
TestCase().assertGreater(g2, b2)
TestCase().assertNotEqual(hex1, hex2)`);
}})
```

Pozivanje `generate_hex("blue")` dva puta bi trebalo da vrati dve različite heksadecimalne boje gde je plava dominantna.

```js
({test: () => { runPython(`
from unittest import TestCase

hex1 = generate_hex("blue").upper()
is_valid_hex1 = len(hex1) == 6 and all(c in "0123456789ABCDEF" for c in hex1)
TestCase().assertTrue(is_valid_hex1)

r1 = int(hex1[:2], 16)
g1 = int(hex1[2:4], 16)
b1 = int(hex1[4:], 16)

TestCase().assertGreater(b1, r1)
TestCase().assertGreater(b1, g1)

hex2 = generate_hex("blue").upper()
is_valid_hex2 = len(hex2) == 6 and all(c in "0123456789ABCDEF" for c in hex2)
TestCase().assertTrue(is_valid_hex2)

r2 = int(hex2[:2], 16)
g2 = int(hex2[2:4], 16)
b2 = int(hex2[4:], 16)

TestCase().assertGreater(b2, r2)
TestCase().assertGreater(b2, g2)
TestCase().assertNotEqual(hex1, hex2)`);
}})
```

# --seed--

## --seed-contents--

```py
def generate_hex(color):

    return color
```

# --solutions--

```py
import random
def generate_hex(color):
    def to_hex(n):
        return hex(n)[2:].upper().zfill(2)

    dominant = random.randint(170, 255)
    weak1 = random.randint(0, 169)
    weak2 = random.randint(0, 169)

    if color.lower() == "red":
        r = dominant
        g = weak1
        b = weak2
    elif color.lower() == "green":
        r = weak1
        g = dominant
        b = weak2
    elif color.lower() == "blue":
        r = weak1
        g = weak2
        b = dominant
    else:
        return "Invalid color"

    return f'{to_hex(r)}{to_hex(g)}{to_hex(b)}'
```