id: 6821ebee237de8297eaee796
title: "Izazov 20: Duplikati u nizu"
challengeType: 29
dashedName: challenge-20
---

# --description--

Dato niz neličnih brojeva, vratite niz neličnih brojeva koji se više od jednom ponavljaju u originalnom nizu, sortiran po rastućem redosledu. Ako nema vrednosti koja se ponavlja više od jednom, vratite prazan niz.

- Uključite samo jedan primer svake vrednosti u vraćeni niz.

# --hints--

`find_duplicates([1, 2, 3, 4, 5])` bi trebalo da vrati `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_duplicates([1, 2, 3, 4, 5]), [])`)
}})
```

`find_duplicates([1, 2, 3, 4, 1, 2])` bi trebalo da vrati `[1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_duplicates([1, 2, 3, 4, 1, 2]), [1, 2])`)
}})
```

`find_duplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])` bi trebalo da vrati `[-6, 0, 2, 4, 5, 23]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_duplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_duplicates(arr):

    return arr
```

# --solutions--

```py
def find_duplicates(arr):
    seen = set()
    duplicates = set()

    for num in arr:
        if num in seen:
            duplicates.add(num)
        else:
            seen.add(num)

    return sorted(duplicates)
```