id: 68af0687ef34c76c28ffa54b
title: "Challenge 32: Obrnuti rečenicu"
challengeType: 29
dashedName: challenge-32
---

# --description--

Dati niz reči, vratite novi niz reči u obrnutom redosledu. Na primer, prva reč bi trebalo da bude na kraju vraćenog niza, a poslednja reč bi trebalo da bude na početku vraćenog niza.

- U datom nizu, reči mogu biti razdvojene jednim ili više razmaka.
- Vraćeni niz bi trebalo da sadrži samo jedan razmak između reči.

# --hints--

`reverse_sentence("world hello")` should return `"hello world"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("world hello"), "hello world")`)
}})
```

`reverse_sentence("push commit git")` should return `"git commit push"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("push commit git"), "git commit push")`)
}})
```

`reverse_sentence("npm  install   apt    sudo")` should return `"sudo apt install npm"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("npm  install   apt    sudo"), "sudo apt install npm")`)
}})
```

`reverse_sentence("import    default   function  export")` should return `"export function default import"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(reverse_sentence("import    default   function  export"), "export function default import")`)
}})
```

# --seed--

## --seed-contents--

```py
def reverse_sentence(sentence):

    return sentence
```

# --solutions--

```py
def reverse_sentence(sentence):

   return " ".join(sentence.split()[::-1])
```