id: 68af0687ef34c76c28ffa549
title: "Challenge 31: Array Diff"
challengeType: 29
dashedName: challenge-31
---

# --description--

Kada dobijete dva niza podataka sa string vrednostima, vratite novi niz koji sadrži sve vrednosti koje se pojavljuju u samo jednom od tih nizova.

- Vraćeni niz podataka mora biti sortiran po abecednom redosledu.

# --hints--

`array_diff(["apple", "banana"], ["apple", "banana", "cherry"])` should return `["cherry"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_diff(["apple", "banana"], ["apple", "banana", "cherry"]), ["cherry"])`)
}})
```

`array_diff(["apple", "banana", "cherry"], ["apple", "banana"])` should return `["cherry"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_diff(["apple", "banana", "cherry"], ["apple", "banana"]), ["cherry"])`)
}})
```

`array_diff(["one", "two", "three", "four", "six"], ["one", "three", "eight"])` should return `["eight", "four", "six", "two"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_diff(["one", "two", "three", "four", "six"], ["one", "three", "eight"]), ["eight", "four", "six", "two"])`)
}})
```

`array_diff(["two", "four", "five", "eight"], ["one", "two", "three", "four", "seven", "eight"])` should return `["five", "one", "seven", "three"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_diff(["two", "four", "five", "eight"], ["one", "two", "three", "four", "seven", "eight"]), ["five", "one", "seven", "three"])`)
}})
```

`array_diff(["I", "like", "freeCodeCamp"], ["I", "like", "rocks"])` should return `["freeCodeCamp", "rocks"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(array_diff(["I", "like", "freeCodeCamp"], ["I", "like", "rocks"]), ["freeCodeCamp", "rocks"])`)
}})
```