id: 68b06e589bf227324381476f
title: "Izazov 35: Frekvencija reči"
challengeType: 29
dashedName: challenge-35
---

# --description--

Kada dobijete odlomak, vratite niz podataka sa tri najčešće reči.

- Reči u odlomku su razdvojene razmakom.
- Ignorišite razliku između malih i velikih slova u datom odlomku. Na primer, smatrate `Hello` i `hello` kao jednu reč.
- Ignorišite interpunkciju u datom odlomku. Interpunkcija su zarezi (`,`), tačke (`.`) i uzvičnici (`!`).
- Niz podataka koji se vraća mora sadržati sve reči u malim slovima.
- Niz podataka koji se vraća mora biti u silaznom redosledu po frekvenciji pojavljivanja, sa najčešće pojavljenom reči na prvom mestu.

# --hints--

`get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` should return `["coding", "python", "in"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"])`)
}})
```

`get_words("I like coding. I like testing. I love debugging!")` should return `["i", "like", "coding"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"])`)
}})
```

`get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` should return `["debug", "test", "deploy"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_words("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_words(paragraph):

    return paragraph
```

# --solutions--

```py
import re
def get_words(paragraph):
    cleaned = re.sub(r'[.,?]', '', paragraph.lower())

    words = cleaned.split()

    freq = {}
    for word in words:
        freq[word] = freq.get(word, 0) + 1

    sorted_words = sorted(freq.keys(), key=lambda w: freq[w], reverse=True)

    return sorted_words[:3]
```