id: 68b06e589bf2273243814771
title: "Izazov 36: Regulator temperature"
challengeType: 29
dashedName: challenge-36
---

# --description--

Dajući trenutnu temperaturu sobe i željenu temperaturu, vratite niz znakova koji pokazuje kako podesiti temperaturu sobe u skladu sa ovim uslovima:

- Vratite `"heat"` ako je trenutna temperatura niža od željene temperature.
- Vratite `"cool"` ako je trenutna temperatura viša od željene temperature.
- Vratite `"hold"` ako je trenutna temperatura jednaka željenoj temperaturi.

# --hints--

`adjust_thermostat(68, 72)` treba da vrati `"heat"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(68, 72), "heat")`)
}})
```

`adjust_thermostat(75, 72)` treba da vrati `"cool"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(75, 72), "cool")`)
}})
```

`adjust_thermostat(72, 72)` treba da vrati `"hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 72), "hold")`)
}})
```

`adjust_thermostat(-20.5, -10.1)` treba da vrati `"heat"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(-20.5, -10.1), "heat")`)
}})
```

`adjust_thermostat(100, 99.9)` treba da vrati `"cool"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(100, 99.9), "cool")`)
}})
```

`adjust_thermostat(0.0, 0.0)` treba da vrati `"hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(0.0, 0.0), "hold")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(temp, target):

    return temp
```

# --solutions--

```py
def adjust_thermostat(temp, target):
    if temp < target:
        return "heat"
    elif temp > target:
        return "cool"
    else:
        return "hold"
```