id: 68b1f72371a5ac895ac70a06
title: "Izazov 42: Čuvanje videa"
challengeType: 29
dashedName: challenge-42
---

# --description--

Uzimajući u obzir veličinu videa, jedinicu za veličinu videa, kapacitet hard diska i jedinicu za hard disk, vratite broj videa koje hard disk može sačuvati koristeći sledeće uslove:

- Jedinica za veličinu videa može biti bajti (`"B"`), kilobajti (`"KB"`), megabajti (`"MB"`) ili gigabajti (`"GB"`).
- Ako nije navedena nijedna od gore spomenutih video jedinica, vratite `"Invalid video unit"`.
- Jedinica kapaciteta hard diska može biti gigabajti (`"GB"`) ili terabajti (`"TB"`).
- Ako nije navedena nijedna od gore spomenutih jedinica za disk, vratite `"Invalid drive unit"`.
- Vratite ceo broj videa koje disk može sačuvati.
- Koristite sledeće konverzije:

| Jedinica | Ekvivalentno |
|:-------:|:--------:|
| 1 B     |   1 B    |
| 1 KB    |  1000 B  |
| 1 MB    |  1000 KB |
| 1 GB    |  1000 MB |
| 1 TB    |  1000 GB |

Na primer, ako primite `500`, `"MB"`, `100` i `"GB"` kao argumente, odredite koliko videa od 500 MB se može sačuvati na hard disku od 100 GB.

# --hints--

`number_of_videos(500, "MB", 100, "GB")` should return `200`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500, "MB", 100, "GB"), 200)`)
}})
```

`number_of_videos(1, "TB", 10, "TB")` should return `"Invalid video unit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1, "TB", 10, "TB"), "Invalid video unit")`)
}})
```

`number_of_videos(2000, "MB", 100000, "MB")` should return `"Invalid drive unit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(2000, "MB", 100000, "MB"), "Invalid drive unit")`)
}})
```

`number_of_videos(500000, "KB", 2, "TB")` should return `4000`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(500000, "KB", 2, "TB"), 4000)`)
}})
```

`number_of_videos(1.5, "GB", 2.2, "TB")` should return `1466`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(number_of_videos(1.5, "GB", 2.2, "TB"), 1466)`)
}})
```

# --seed--

## --seed-contents--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):

    return video_size
```

# --solutions--

```py
def number_of_videos(video_size, video_unit, drive_size, drive_unit):
    video_units = {"KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000}
    drive_units = {"GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000}

    if video_unit not in video_units:
        return "Invalid video unit"
    if drive_unit not in drive_units:
        return "Invalid drive unit"

    video_bytes = video_size * video_units[video_unit]
    drive_bytes = drive_size * drive_units[drive_unit]

    return int(drive_bytes // video_bytes)
```