id: 68b7cadffed0e75a517da673
title: "Izazov 52: Binarno na decimalno"
challengeType: 29
dashedName: challenge-52
---

# --description--

Ako vam je dat niz karaktera koji predstavlja binarni broj, vratite njegov ekvivalent u decimalnom sistemu kao broj.

Binarni brojevi koriste samo cifre `0` i `1` za predstavljanje bilo kog broja. Da biste pretvorili binarno u decimalno, pomnožite svaku cifru sa potencijom `2`, a zatim ih sabrnite. Počnite množenjem najdesnije cifre sa `2^0`, sledeće cifre sa levog sa `2^1` i tako dalje. Kada se sve cifre pomnože sa potencijama `2`, saberi rezultate.

Na primer, binarni broj `101` je jednak `5` u decimalnom sistemu jer:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`to_decimal("101")` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("101"), 5)`)
}})
```

`to_decimal("1010")` should return `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010"), 10)`)
}})
```

`to_decimal("10010")` should return `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("10010"), 18)`)
}})
```

`to_decimal("1010101")` should return `85`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_decimal("1010101"), 85)`)
}})
```

# --seed--

## --seed-contents--

```py
def to_decimal(binary):

    return binary
```

# --solutions--

```py
def to_decimal(binary):
    decimal = 0
    for i, bit_char in enumerate(reversed(binary)):
        bit = int(bit_char)
        decimal += bit * (2 ** i)
    return decimal
```