---
id: 68b7cadffed0e75a517da675
title: "Challenge 53: Decimal to Binary"
challengeType: 29
dashedName: challenge-53
---

# --description--

Ako vam je dat celobrojni, pozitivan broj, vratite njegov binarni prikaz kao niz karaktera.

Binarni brojevi koriste samo cifre `0` i `1` da predstavljaju bilo koji broj. Da biste pretvorili decimalni broj u binarni, stalno delite taj broj sa `2` i beležite ostatak. Ponovite dok broj ne postane nula. Pročitajte beležene ostatke od poslednjeg do prvog. Na primer, konverzija broja `12` u binarni sistem:

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12` u binarnom sistemu je `1100`.

# --hints--

`to_binary(5)` bi trebalo da vrati `"101"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(5), "101")`)
}})
```

`to_binary(12)` bi trebalo da vrati `"1100"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(12), "1100")`)
}})
```

`to_binary(50)` bi trebalo da vrati `"110010"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(50), "110010")`)
}})
```

`to_binary(99)` bi trebalo da vrati `"1100011"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_binary(99), "1100011")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_binary(decimal):

    return decimal
```

# --solutions--

```py
def to_binary(decimal):
    if decimal == 0:
        return "0"
    binary = ""
    while decimal > 0:
        binary = str(decimal % 2) + binary
        decimal //= 2
    return binary
```