id: 68c1a929005bf54d342aa8d5
title: "Challenge 58: Space Week Day 4: Landing Spot"
challengeType: 29
dashedName: challenge-58
---

# --description--

Na četvrtom danu Tjedna u svemiru, dobijate matricu brojeva (dvodimenzionalni niz podataka), koji predstavljaju zone gdje vaš rover može sletjeti. Pronađite najsigurnije mjesto za slijetanje prema sljedećim pravilima:

- Svaka ćelija u matrici će imati broj od `0-9`, uključujući i to.
- Svaki `0` predstavlja potencijalno mjesto za slijetanje.
- Svaki broj osim `0` predstavlja opasno mjesto za slijetanje. Što je veći broj, to je veća opasnost.
- Najsigurnije mjesto definisano je kao ćelija `0` čije susjedne ćelije (do 4 susjeda, ne računajući uglove) imaju najmanji ukupan indeks opasnosti.
- Ignorišite susjede koji su van granica (uglasi i rubovi imaju samo malo susjeda).
- Vratite indekse najsigurnijeg mjesta za slijetanje. Uvijek će postojati samo jedno najsigurnije mjesto.

Na primjer, ako dobijete:

```js
[
  [1, 0],
  [2, 0]
]
```

Vratite `[0, 1]`, indekse `0` u prvom redu.

# --hints--

`find_landing_spot([[1, 0], [2, 0]])` treba vratiti `[0, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[1, 0], [2, 0]]), [0, 1])`)
}})
```

`find_landing_spot([[9, 0, 3], [7, 0, 4], [8, 0, 5]])` treba vratiti `[1, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[9, 0, 3], [7, 0, 4], [8, 0, 5]]), [1, 1])`)
}})
```

`find_landing_spot([[1, 2, 1], [0, 0, 2], [3, 0, 0]])` treba vratiti `[2, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[1, 2, 1], [0, 0, 2], [3, 0, 0]]), [2, 2])`)
}})
```

`find_landing_spot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]])` treba vratiti `[2, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_landing_spot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]]), [2, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def find_landing_spot(matrix):

    return matrix
```

# --solutions--

```py
def find_landing_spot(matrix):
    best_spot = None
    lowest_neighbor_sum = float('inf')

    for i in range(len(matrix)):
        for j in range(len(matrix[i])):
            if matrix[i][j] == 0:
                current_neighbor_sum = 0

                if i > 0:
                    current_neighbor_sum += matrix[i - 1][j]
                if j < len(matrix[i]) - 1:
                    current_neighbor_sum += matrix[i][j + 1]
                if i < len(matrix) - 1:
                    current_neighbor_sum += matrix[i + 1][j]
                if j > 0:
                    current_neighbor_sum += matrix[i][j - 1]

                if current_neighbor_sum < lowest_neighbor_sum:
                    lowest_neighbor_sum = current_neighbor_sum
                    best_spot = [i, j]

    return best_spot
```