id: 68c1a929005bf54d342aa8d2
title: "Izazov 55: Nedelja u svemiru Dan 1: Klasifikacija zvezda"
challengeType: 29
dashedName: challenge-55
---

# --description--

4. oktobar označava početak Nedelje u svemiru na Zemlji. Sledećih sedam dana donosiće vam izazove programiranja sa tematikom astronomije.

Za današnji izazov, dato je tempertura površine zvezde u Kelvinima (K) i potrebno je utvrditi njenu klasifikaciju zvezda prema sledećim standardima:

- `"O"`: 30.000 K ili više
- `"B"`: 10.000 K - 29.999 K
- `"A"`: 7.500 K - 9.999 K
- `"F"`: 6.000 K - 7.499 K
- `"G"`: 5.200 K - 5.999 K
- `"K"`: 3.700 K - 5.199 K
- `"M"`: 0 K - 3.699 K

- Vratite klasifikaciju zvezde koja je data.

# --hints--

`classification(5778)` bi trebalo da vrati `"G"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(classification(5778), "G")`)
}})
```

`classification(2400)` bi trebalo da vrati `"M"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(classification(2400), "M")`)
}})
```

`classification(9999)` bi trebalo da vrati `"A"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(classification(9999), "A")`)
}})
```

`classification(3700)` bi trebalo da vrati `"K"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(classification(3700), "K")`)
}})
```

`classification(3699)` bi trebalo da vrati `"M"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(classification(3699), "M")`)
}})
```

`classification(210000)` bi trebalo da vrati `"O"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(classification(210000), "O")`)
}})
```

`classification(6000)` bi trebalo da vrati `"F"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(classification(6000), "F")`)
}})
```

`classification(11432)` bi trebalo da vrati `"B"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(classification(11432), "B")`)
}})
```

# --seed--

## --seed-contents--

```py
def classification(temp):

    return temp
```

# --solutions--

```py
def classification(temp):
    if temp >= 30000:
        return "O"
    elif temp >= 10000:
        return "B"
    elif temp >= 7500:
        return "A"
    elif temp >= 6000:
        return "F"
    elif temp >= 5200:
        return "G"
    elif temp >= 3700:
        return "K"
    else:
        return "M"
```