id: 68c1a929005bf54d342aa8d6
title: "Izazov 59: Peti dan u Nedelji svemira: Zona Goldiloks"
challengeType: 29
dashedName: challenge-59
---

# --description--

Za peti dan u Nedelji svemira, izračunaćeš "Zonu Goldiloks" zvezde – oblast oko zvezde gde su uslovi "savršeno pravi" za postojanje tečnosti vode.

Uzimajući u obzir masu zvezde, vrati niz podataka sa početnom i krajnjom udaljenošću njene Zone Goldiloks za Jedinice svemira.

Da bi se izračunala Zona Goldiloks:

1. Pronađi svetlo zvezde podizanjem njene mase na stepen 3.5.
2. Početak zone je 0.95 puta kvadratni koren njenog svetla.
3. Kraj zone je 1.37 puta kvadratni koren njenog svetla.

- Vrati udaljenost zaokruženu na dve decimalne začišlja.

Na primer, ako dobiješ `1` kao masu, vrati `[0.95, 1.37]`.

# --hints--

`goldilocks_zone(1)` bi trebalo da vrati `[0.95, 1.37]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(1), [0.95, 1.37])`)
}})
```

`goldilocks_zone(0.5)` bi trebalo da vrati `[0.28, 0.41]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(0.5), [0.28, 0.41])`)
}})
```

`goldilocks_zone(6)` bi trebalo da vrati `[21.85, 31.51]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(6), [21.85, 31.51])`)
}})
```

`goldilocks_zone(3.7)` bi trebalo da vrati `[9.38, 13.52]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(3.7), [9.38, 13.52])`)
}})
```

`goldilocks_zone(20)` bi trebalo da vrati `[179.69, 259.13]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(goldilocks_zone(20), [179.69, 259.13])`)
}})
```

# --seed--

## --seed-contents--

```py
def goldilocks_zone(mass):

    return mass
```

# --solutions--

```py
import math
def goldilocks_zone(mass):
    luminosity = mass ** 3.5
    start = 0.95 * math.sqrt(luminosity)
    end = 1.37 * math.sqrt(luminosity)
    return [round(start, 2), round(end, 2)]
```