id: 68c1a929005bf54d342aa8d3
title: "Challenge 56: Space Week Day 2: Exoplanet Search"
challengeType: 29
dashedName: challenge-56
---

# --description--

Za drugog dana u Nedelji svemira, dat vam je niz slova gde svako slovo predstavlja meru svetlosne jačine zvezde. Utvrdite da li su mere otkrile ekoplanet korišćenjem tranzitne metode (transit method). Tranzitna metoda se dešava kada planeta prođe ispred zvezde, smanjujući vidljivu svetlost.

- Mere svetlosne jačine sadrže samo slova `0-9` i `A-Z`, gde svaka mera odgovara sledećim numeričkim vrednostima:
- Slova `0-9` predstavljaju nivoe svetlosne jačine `0-9`.
- Slova `A-Z` predstavljaju nivoe svetlosne jačine `10-35`.

Zvezda se smatra da ima ekoplanet ako je bilo koja mera manja ili jednaka 80% proseka svih mera. Na primer, ako je prosek svetlosne jačine zvezde 10, smatraće se da ima ekoplanet ako je bilo koja mera 8 ili niža.

# --hints--

`has_exoplanet("665544554")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("665544554"), False)`)
}})
```

`has_exoplanet("FGFFCFFGG")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FGFFCFFGG"), True)`)
}})
```

`has_exoplanet("MONOPLONOMONPLNOMPNOMP")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("MONOPLONOMONPLNOMPNOMP"), False)`)
}})
```

`has_exoplanet("FREECODECAMP")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("FREECODECAMP"), True)`)
}})
```

`has_exoplanet("9AB98AB9BC98A")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("9AB98AB9BC98A"), False)`)
}})
```

`has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(has_exoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def has_exoplanet(readings):

    return readings
```

# --solutions--

```py
def has_exoplanet(readings):
    values = [int(c, 36) for c in readings]
    average = sum(values) / len(values)
    threshold = average * 0.8
    for v in values:
        if v <= threshold:
            return True
    return False
```