id: 68c497f3aaefc9fd9f1b0e25
title: "Izazov Challenge 61: Nedelja u svemiru Dan 7: Gorivo za lansiranje"
challengeType: 29
dashedName: challenge-61
---

# --description--

Za poslednji dan Nedelje u svemiru, biće vam dat težina (kg) tereta koji želite poslati u svemir. Odredite količinu goriva potrebnu za lansiranje vašeg tereta u svemir koristeći sledeća pravila:

- Raketni motori zahtevaju 1 kg goriva na svaki 5 kg težine za lansiranje.
- Samo gorivo ima težinu. Dakle, kada dodate gorivo, ukupna težina raste, što zahteva više goriva, što povećava težinu, i tako dalje.
- Da biste izračunali ukuruhan potreban gorivo: počnite sa težinom tereta, izračunajte gorivo potrebno za tu težinu, dodajte to gorivo na ukupnu težinu, pa ponovo izračunajte. Ponovite ovaj proces sve dok ne bude potrebno manje od 1 kg dodatnog goriva, a zatim se zaustavite.
- Ignorišite težinu samog raketnog motora. Izračunajte samo gorivo potrebne za lansiranje tereta i njegovog sopstvenog goriva.

Na primer, ako imate teret od 50 kg, biće vam potrebno 10 kg goriva za lansiranje tog tereta (težina / 5), što povećava ukupnu težinu na 60 kg, što zahteva 12 kg za lansiranje (2 kg više), što povećava ukupnu težinu na 62 kg, što zahteva 12.4 kg za lansiranje - 0.4 kg više - što je manje od 1 kg dodatnog goriva, pa se ovde zaustavljamo. Ukupna težina za lansiranje je 62.4 kg, gde je 50 početni teret, a 12.4 je gorivo.

- Vratite količinu potrebnog goriva zaokruženu na jednu decimalnu tačku.

# --hints--

`launch_fuel(50)` should return `12.4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(50), 12.4)`)
}})
```

`launch_fuel(500)` should return `124.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(500), 124.8)`)
}})
```

`launch_fuel(243)` should return `60.7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(243), 60.7)`)
}})
```

`launch_fuel(11000)` should return `2749.8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(11000), 2749.8)`)
}})
```

`launch_fuel(6214)` should return `1553.4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(launch_fuel(6214), 1553.4)`)
}})
```

# --seed--

## --seed-contents--

```py
def launch_fuel(payload):

    return payload
```

# --solutions--

```py
def launch_fuel(payload):
    total_mass = payload
    additional_fuel = total_mass / 5
    total_fuel = additional_fuel

    while additional_fuel >= 1:
        total_mass += additional_fuel
        additional_fuel = (total_mass / 5) - total_fuel
        total_fuel += additional_fuel

    return round(total_fuel, 1)
```