id: 68c497f3aaefc9fd9f1b0e24
title: "Izazov 60: Nedelja u svemiru Dan 6: Faza Meseca"
challengeType: 29
dashedName: challenge-60
---

# --description--

Za šestog dana Nedelje u svemiru, biće vam dat datum u formatu `"YYYY-MM-DD"` i morate odrediti fazu Meseca za taj dan koristeći sledeća pravila:

Koristite jednostavan ciklusa Meseca od 28 dana, podeljen na četiri jednake faze:

- `"New"`: dani 1 - 7
- `"Waxing"`: dani 8 - 14
- `"Full"`: dani 15 - 21
- `"Waning"`: dani 22 - 28

Nakon 28 dana, ciklus se ponavlja sa danom 1, novim mesecem.

- Koristite `"2000-01-06"` kao referencu za novi mesec (dan 1 ciklusa) da biste odredili fazu datuma.
- Neće vam biti dat nijedan datum pre referentnog datuma.
- Vratite tačnu fazu kao niz karaktera.

**Napomena:** Dan 1 predstavlja dan novog meseca, što znači da je prošlo 0 dana od prethodnog novog meseca.

# --hints--

`moon_phase("2000-01-12")` should return `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-12"), "New")`)
}})
```

`moon_phase("2000-01-13")` should return `"Waxing"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2000-01-13"), "Waxing")`)
}})
```

`moon_phase("2014-10-15")` should return `"Full"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2014-10-15"), "Full")`)
}})
```

`moon_phase("2012-10-21")` should return `"Waning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2012-10-21"), "Waning")`)
}})
```

`moon_phase("2022-12-14")` should return `"New"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(moon_phase("2022-12-14"), "New")`)
}})
```

# --seed--

## --seed-contents--

```py
def moon_phase(date_string):

    return date_string
```

# --solutions--

```py
from datetime import datetime
def moon_phase(date_string):
    ref_date = datetime.strptime("2000-01-06", "%Y-%m-%d")
    target_date = datetime.strptime(date_string, "%Y-%m-%d")

    diff_days = (target_date - ref_date).days
    cycle_day = (diff_days % 28) + 1

    if cycle_day <= 7:
        return "New"
    elif cycle_day <= 14:
        return "Waxing"
    elif cycle_day <= 21:
        return "Full"
    else:
        return "Waning"
```