id: 68cae5b538ff798bbd4da006
title: "Izazov 68: Maskiranje kreditne kartice"
challengeType: 29
dashedName: challenge-68
---

# --description--

Dato je niz brojeva kreditne kartice, vratite maskiranu verziju koristeći sljedeće uvjete:

- Niz će imati četiri grupe od po četiri broja (`0-9`), a sve grupe su razdvojene razmakom ili crticom (`-`).
- Promijenite sve brojeve osim zadnje četiri u zvjezdice (`*`).
- Ostale znakove ostavite nepromijenjene.

Na primjer, ako je dato `"4012-8888-8888-1881"`, vratit će se `"****-****-****-1881"`.

# --hints--

`mask("4012-8888-8888-1881")` should return `"****-****-****-1881"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("4012-8888-8888-1881"), "****-****-****-1881")`)
}})
```

`mask("5105 1051 0510 5100")` should return `"**** **** **** 5100"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("5105 1051 0510 5100"), "**** **** **** 5100")`)
}})
```

`mask("6011 1111 1111 1117")` should return `"**** **** **** 1117"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("6011 1111 1111 1117"), "**** **** **** 1117")`)
}})
```

`mask("2223-0000-4845-0010")` should return `"****-****-****-0010"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mask("2223-0000-4845-0010"), "****-****-****-0010")`)
}})
```

# --seed--

## --seed-contents--

```py
def mask(card):

    return card
```

# --solutions--

```py
def mask(card):
    if '-' in card:
        split = card.split('-')
        return f"****-****-****-{split[3]}"
    else:
        split = card.split(' ')
        return f"**** **** **** {split[3]}"
```