id: 68cae5b538ff798bbd4da004
title: "Vežba 66: Uklanjanje HTML oznaka"
challengeType: 29
dashedName: challenge-66
---

# --description--

Izvučeno je niz znakova HTML koda; uklonite oznake i vratite običan tekstualni sadržaj.

- Ulazni niz će sadržati samo validan HTML.
- HTML oznake mogu biti ugnježđene.
- Uklonite sve oznake i atribute elemenata.

Na primer, `'<a href="#">Click here</a>'` bi trebalo da vrati `"Click here"`.

# --hints--

`strip_tags('<a href="#">Click here</a>')` bi trebalo da vrati `"Click here"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<a href="#">Click here</a>'), "Click here")`)
}})
```

`strip_tags('<p class="center">Hello <b>World</b>!</p>')` bi trebalo da vrati `"Hello World!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!")`)
}})
```

`strip_tags('<img src="cat.jpg" alt="Cat">')` bi trebalo da vrati prazan niz znakova (`""`).

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<img src="cat.jpg" alt="Cat">'), "")`)
}})
```

`strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` bi trebalo da vrati `sectionsection`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strip_tags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection")`)
}})
```

# --seed--

## --seed-contents--

```py
def strip_tags(html):

    return html
```

# --solutions--

```py
import re
def strip_tags(html):
    return re.sub(r'<[^>]*>', '', html)
```