id: 68cae5b538ff798bbd4da007
title: "Challenge 69: Potpajani čarape"
challengeType: 29
dashedName: challenge-69
---

# --description--

Uzimajući u obzir celonočnu brojku koja predstavlja početni broj parova čarapa i drugu brojku koja predstavlja koliko je ciklusa pranja urađeno, vratite trenutni broj kompletnih parova čarapa koristeći sledeće uslove:

- Svakih 2 ciklusa pranja, gubiš jednu čarappu.
- Svakih 3 ciklusa pranja, dobijaš jednu izgubljenu čarapu.
- Svakih 5 ciklusa pranja, jedna čarapa se uništava i mora biti bacena.
- Svakih 10 ciklusa pranja, kupuješ novi par čarapa.
- Ne možeš imati manje od nule čarapa.
- Pravila mogu da se preklapaju. Na primer, u 10. ciklusu pranja, gubiš jednu čarappu, bacaš jednu i kupuješ novi par čarapa.
- Vrati broj kompletnih parova čarapa.

# --hints--

`sock_pairs(2, 5)` bi trebalo da vrati `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(2, 5), 1)`)
}})
```

`sock_pairs(1, 2)` bi trebalo da vrati `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 2), 0)`)
}})
```

`sock_pairs(5, 11)` bi trebalo da vrati `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(5, 11), 4)`)
}})
```

`sock_pairs(6, 25)` bi trebalo da vrati `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(6, 25), 3)`)
}})
```

`sock_pairs(1, 8)` bi trebalo da vrati `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sock_pairs(1, 8), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sock_pairs(pairs, cycles):

    return pairs
```

# --solutions--

```py
def sock_pairs(pairs, cycles):
    socks = pairs * 2

    for i in range(1, cycles + 1):
        if i % 2 == 0:
            socks -= 1
        if i % 3 == 0:
            socks += 1
        if i % 5 == 0:
            socks -= 1
        if i % 10 == 0:
            socks += 2

        if socks < 0:
            socks = 0

    return socks // 2
```