id: 68cae5b538ff798bbd4da00a
title: "Challenge 72: Thermostat Adjuster 2"
challengeType: 29
dashedName: challenge-72
---

# --description--

Uzimajući u obzir trenutnu temperaturu sobe u Farenhajtu i željenu temperaturu u Celzijusima, vratite niz znakova koji opisuje kako podesiti temperaturu sobe prema ovim uslovima:

- Vratite `"Heat: X degrees Fahrenheit"` ako je trenutna temperatura niža od ciljne. Gde je `X` broj Farenhajt stepeni potrebni za podizanje temperature sobe do cilja, zaokruženo na jednu decimalnu tačku.
- Vratite `"Cool: X degrees Fahrenheit"` ako je trenutna temperatura viša od ciljne. Gde je `X` broj Farenhajt stepeni potrebni za hlađenje sobe do cilja, zaokruženo na jednu decimalnu tačku.
- Vratite `"Hold"` ako je trenutna temperatura jednaka ciljnoj.

Da biste pretvorili Celzijus u Farenhajt, pomnožite temperaturu u Celzijusima sa 1.8, a zatim dodajte 32 rezultatu (`F = (C * 1.8) + 32`).

# --hints--

`adjust_thermostat(32, 0)` should return `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(32, 0), "Hold")`)
}})
```

`adjust_thermostat(70, 25)` should return `"Heat: 7.0 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(70, 25), "Heat: 7.0 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(72, 18)` should return `"Cool: 7.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(72, 18), "Cool: 7.6 degrees Fahrenheit")`)
}})
```

`adjust_thermostat(212, 100)` should return `"Hold"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(212, 100), "Hold")`)
}})
```

`adjust_thermostat(59, 22)` should return `"Heat: 12.6 degrees Fahrenheit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(adjust_thermostat(59, 22), "Heat: 12.6 degrees Fahrenheit")`)
}})
```

# --seed--

## --seed-contents--

```py
def adjust_thermostat(current_f, target_c):

    return current_f
```

# --solutions--

```py
def adjust_thermostat(current_f, target_c):
    target_f = (target_c * 1.8) + 32
    diff = round(abs(target_f - current_f), 1)

    if current_f < target_f:
        return f"Heat: {diff} degrees Fahrenheit"
    elif current_f > target_f:
        return f"Cool: {diff} degrees Fahrenheit"
    else:
        return "Hold"
```