---
id: 68d2ba1468508398389487d0
title: "Challenge 75: Hidden Treasure"
challengeType: 29
dashedName: challenge-75
---

# --description--

Pruženi 2D nizova podataka koji predstavlja mapu dna mora koja sadrži skriveno blago, i niz podataka sa koordinatama (`[redak, kolona]`) za ronjenje u potrazi za blagom. Vratite `"Empty"`, `"Found"` ili `"Recovered"` koristeći sledeće pravilo:

- Pruženi 2D niz podataka će sadržati samo jedno neotkriveno blago koje zauzima više ćelija.
- Svaka ćelija u 2D nizu podataka će imati jednu od sledećih vrednosti:
  - `"-"`: Nema blaga.
  - `"O"`: Deo blaga koji nije pronađen.
  - `"X"`: Deo blaga koji je već pronađen.

- Ako područje za ronjenje nema blago, vratite `"Empty"`.
- Ako područje za ronjenje pronađe blago, ali barem jedan deo blaga još uvek nije pronađen, vratite `"Found"`.
- Ako područje za ronjenje pronađe poslednji deo blaga koji nije pronađen, vratite `"Recovered"`.

Na primer, ako dobijete:

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

I sa `[2, 1]` kao koordinatama područja za ronjenje, vratite `"Recovered"` jer je ronjenje pronašlo poslednji deo blaga koji nije pronađen.

# --hints--

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` should return `"Recovered"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Recovered")`)
}})
```

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` should return `"Empty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Empty")`)
}})
```

`dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` should return `"Found"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Found")`)
}})
```

`dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` should return `"Found"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Found")`)
}})
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` should return `"Recovered"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Recovered")`)
}})
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` should return `"Empty"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Empty")`)
}})
```

# --seed--

## --seed-contents--

```py
def dive(map, coordinates):

    return map
```

# --solutions--

```py
def dive(map, coordinates):
    row, col = coordinates
    target = map[row][col]

    if target == "-":
        return "Empty"

    if target == "O":
        for r in range(len(map)):
            for c in range(len(map[r])):
                if (r != row or c != col) and map[r][c] == "O":
                    return "Found"
        return "Recovered"

    return "Found"
```