id: 68d30845cc08266018fc46bc
title: "Challenge 76: Komplementarna DNK"
challengeType: 29
dashedName: challenge-76
---

# --description--

Dati je niz slova koji predstavlja DNK lanac; vratite njegov komplementarni niz koristeći sledeća pravila:

- DNK sadrži slova "A", "C", "G" i "T".
- Slovo "A" i "T" su komplementarni jedno drugom.
- Slovo "C" i "G" su komplementarni jedno drugom.

Na primer, ako dobijete "ACGT", vratite "TGCA".

# --hints--

`complementary_dna("ACGT")` bi trebalo da vrati `"TGCA"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ACGT"), "TGCA")`)
}})
```

`complementary_dna("ATGCGTACGTTAGC")` bi trebalo da vrati `"TACGCATGCAATCG"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("ATGCGTACGTTAGC"), "TACGCATGCAATCG")`)
}})
```

`complementary_dna("GGCTTACGATCGAAG")` bi trebalo da vrati `"CCGAATGCTAGCTTC"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC")`)
}})
```

`complementary_dna("GATCTAGCTAGGCTAGCTAG")` bi trebalo da vrati `"CTAGATCGATCCGATCGATC"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(complementary_dna("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC")`)
}})
```

# --seed--

## --seed-contents--

```py
def complementary_dna(strand):

    return strand
```

# --solutions--

```py
def complementary_dna(strand):
    complements = { "A": "T", "T": "A", "C": "G", "G": "C" }
    return "".join(complements[n] for n in strand)
```