id: 68d30845cc08266018fc46be
title: "Izazov 78: Niz celih brojeva"
challengeType: 29
dashedName: challenge-78
---

# --description--

Ako dobijete pozitivan ceo broj, vratite niz karaktera koji prikazuje sve cele brojeve od 1 do, i uključujući, taj broj, u numeričkom redosledu.

Na primer, ako dobijete `5`, vratite `"12345"`.

# --hints--

`sequence(5)` should return `"12345"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(5), "12345")`)
}})
```

`sequence(10)` should return `"12345678910"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(10), "12345678910")`)
}})
```

`sequence(1)` should return `"1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(1), "1")`)
}})
```

`sequence(27)` should return `"123456789101112131415161718192021222324252627"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sequence(27), "123456789101112131415161718192021222324252627")`)
}})
```

# --seed--

## --seed-contents--

```py
def sequence(n):

    return n
```

# --solutions--

```py
def sequence(n):
    result = ""
    for i in range(1, n + 1):
        result += str(i)
    return result
```