id: 68d30fc57588d97fd3027b30
title: "Izazov 79: Navigator"
challengeType: 29
dashedName: challenge-79
---

# --description--

28. oktobar 1994. godine, Netscape Navigator je lansiran, pomažući milionima ljudi da istražuju ranu mrežu.

Na osnovu niza podataka o komandama pretraživača koje ste izvršili u Netscape Navigator-u, vratite trenutnu stranicu koju imate nakon izvršenja svih komandi koristeći sledeća pravila:

- Uvek počinjete na stranici `"Home"`, koja nije uključena u niz komandi.
- Validne komande su:
  - `"Visit Page"`: Gde je `"Page"` naziv stranice koju posetite. Na primer, `"Visit About"` vodi vas na stranicu `"About"`. Kada posetite novu stranicu, uverite se da brišete bilo koju istoriju napred koju imate.
  - `"Back"`: Vodi vas nazad na prethodno posetačenu stranicu u vašoj istoriji ili ostaje na trenutnoj stranici ako nema takve.
  - `"Forward"`: Voči vas napred u istoriji do stranice od koje ste otišli ili ostaje na trenutnoj stranici ako nema takve.

Na primer, ako dobijete `["Visit About Us", "Back", "Forward"]`, vratite `"About Us"`.

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` should return `"About Us"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Back", "Forward"]), "About Us")`)
}})
```

`navigate(["Forward"])` should return `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Forward"]), "Home")`)
}})
```

`navigate(["Back"])` should return `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Back"]), "Home")`)
}})
```

`navigate(["Visit About Us", "Visit Gallery"])` should return `"Gallery"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Gallery"]), "Gallery")`)
}})
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` should return `"Home"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home")`)
}})
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` should return `"Contact"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact")`)
}})
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` should return `"Visit Us"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us")`)
}})
```

# --seed--

## --seed-contents--

```py
def navigate(commands):

    return commands
```

# --solutions--

```py
def navigate(commands):
    history = ["Home"]
    current_index = 0

    for command in commands:
        if command.startswith("Visit "):
            history = history[:current_index + 1]
            history.append(command[6:])
            current_index += 1
        elif command == "Back" and current_index > 0:
            current_index -= 1
        elif command == "Forward" and current_index < len(history) - 1:
            current_index += 1

    return history[current_index]
```