id: 68e39ed6106dac2f0a98fd66
title: "Izazov 84: Zarazjen"
challengeType: 29
dashedName: challenge-84
---

# --description--

Danas, 2. novembra 1988. godine, pustjen je prvi veliki internet vir, zarazivši približno 10% povezanih računara na internet samo za jedan dan.

U ovom izazovu, dobijate broj dana od kada je virus pušten u internet i morate utvrditi koliko je računara zaraženo koristeći sledeća pravila:

- Dan 0, zaraziće se prvi računar.
- Svaki naredni dan, broj zaraženih računara udvostručava se.
- Na svakom trećem danu, primenjuje se zaštita nakon što virus proširi i smanji broj zaraženih računara za 20%. Zaokružite broj računara koji su dobili zaštitu na sledeći ceo broj.

Na primer, dana:

- Dan 0: Zarazeno je ukupno 1 računar.
- Dan 1: Zarazeno je ukupno 2 računara.
- Dan 2: Zarazeno je ukupno 4 računara.
- Dan 3: Zarazeno je ukupno 8 računara. Zatim, primenite zaštitu: 8 zaraženih računara * 20% = 1.6 računara dobilo zaštitu. Zaokružite 1.6 na 2. 8 zaraženih računara - 2 dobila zaštitu = ukupno 6 računara zarazeno nakon dana 3.

Vratite ukupan broj zarazenih računara nakon datoг broja dana.

# --hints--

`infected(1)` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(1), 2)`)
}})
```

`infected(3)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(3), 6)`)
}})
```

`infected(8)` should return `152`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(8), 152)`)
}})
```

`infected(17)` should return `39808`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(17), 39808)`)
}})
```

`infected(25)` should return `5217638`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(infected(25), 5217638)`)
}})
```

# --seed--

## --seed-contents--

```py
def infected(days):

    return days
```

# --solutions--

```py
import math
def infected(days):
    infected = 1
    for day in range(1, days + 1):
        infected *= 2
        if day % 3 == 0:
            patched = math.ceil(infected * 0.2)
            infected -= patched

    return infected
```