id: 68ee9e3066cfd4eb2328e8a8
title: "Challenge 89: Counting Cards"
challengeType: 29
dashedName: challenge-89
---

# --description--

Standardni pak karti za igranja imaju po 13 jedinstvenih karata u svakoj šardi. Ako dobijete pun broj koji predstavlja broj karata koje treba izabrati iz pakovanja, vratite broj jedinstvenih kombinacija karata koje možete izabrati.

- Redosled nije bitan. Izbor karte A pa karte B jednak je izboru karte B pa karte A.

Na primer, ako dobijete `52`, vratite `1`. Postoji samo jedna kombinacija od 52 karte za izbor iz pakovanja od 52 karte. I ako dobijete `2`, vratite `1326`. Postoji 1326 kombinacija karata koje možete dobiti kada izaberete 2 karte iz pakovanja.

# --hints--

`combinations(52)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(52), 1)`)
}})
```

`combinations(1)` should return `52`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(1), 52)`)
}})
```

`combinations(2)` should return `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(2), 1326)`)
}})
```

`combinations(5)` should return `2598960`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(5), 2598960)`)
}})
```

`combinations(10)` should return `15820024220`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(10), 15820024220)`)
}})
```

`combinations(50)` should return `1326`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(combinations(50), 1326)`)
}})
```

# --seed--

## --seed-contents--

```py
def combinations(cards):

    return cards
```

# --solutions--

```py
from math import factorial
def combinations(cards):
    n = 52
    
    return factorial(n) // (factorial(cards) * factorial(n - cards))
```