---
id: 68ee9e3066cfd4eb2328e8a6
title: "Izazov 87: Izgradnja matrice"
challengeType: 29
dashedName: challenge-87
---

# --description--

Kada dobijete dva celobrojna broja (broj redova i broj kolona), vratite matricu (list podataka redova podataka) ispunjenu nulama (`0`) dimenzija koje su date.

Na primer, ako dobijete `2` i `3`, vratite:

```json
[
  [0, 0, 0],
  [0, 0, 0]
]
```

# --hints--

`build_matrix(2, 3)` bi trebalo da vrati `[[0, 0, 0], [0, 0, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(2, 3), [[0, 0, 0], [0, 0, 0]])`)
}})
```

`build_matrix(3, 2)` bi trebalo da vrati `[[0, 0], [0, 0], [0, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(3, 2), [[0, 0], [0, 0], [0, 0]])`)
}})
```

`build_matrix(4, 3)` bi trebalo da vrati `[[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(4, 3), [[0, 0, 0], [0, 0, 0], [0, 0, 0], [0, 0, 0]])`)
}})
```

`build_matrix(9, 1)` bi trebalo da vrati `[[0], [0], [0], [0], [0], [0], [0], [0], [0]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(build_matrix(9, 1), [[0], [0], [0], [0], [0], [0], [0], [0], [0]])`)
}})
```

# --seed--

## --seed-contents--

```py
def build_matrix(rows, cols):

    return rows
```

# --solutions--

```py
def build_matrix(rows, cols):

    return [[0 for _ in range(cols)] for _ in range(rows)]
```