id: 68f6587287ad1f4ad39b0c80
title: "Izazov 94: Generator e-mail potpisa"
challengeType: 29
dashedName: challenge-94
---

# --description--

Uz dato niz slova imena osobe, titula i kompanije, vratite e-mail potpis kao jedan niz znakova koristeći sledeće pravilo:

- Ime mora biti navedeno prvo, praćeno prefiksom koji zavisi od prvog slova imena. Za imena koja počinju sa (bez obzira na veliko ili malo slovo):
  - `A-I`: Koristite `>>` kao prefiks.
  - `J-R`: Koristite `--` kao prefiks.
  - `S-Z`: Koristite `::` kao prefiks.
- Zarez i razmak (`, `) mora slediti ime.
- Titula i kompanija moraju slediti zarez i razmak, razdvojeni sa `" at "` (sa razmacima na obe strane).

Na primer, ako vam je dato `"Quinn Waverly"`, `"Founder and CEO"` i `"TechCo"`, vratite `"--Quinn Waverly, Founder and CEO at TechCo"`.

# --hints--

`generate_signature("Quinn Waverly", "Founder and CEO", "TechCo")` should return `"--Quinn Waverly, Founder and CEO at TechCo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo")`)
}})
```

`generate_signature("Alice Reed", "Engineer", "TechCo")` should return `">>Alice Reed, Engineer at TechCo"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo")`)
}})
```

`generate_signature("Tina Vaughn", "Developer", "example.com")` should return `"::Tina Vaughn, Developer at example.com"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com")`)
}})
```

`generate_signature("B. B.", "Product Tester", "AcmeCorp")` should return `">>B. B., Product Tester at AcmeCorp"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp")`)
}})
```

`generate_signature("windstorm", "Cloud Architect", "Atmospheronics")` should return `"::windstorm, Cloud Architect at Atmospheronics"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_signature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_signature(name, title, company):

    return name
```

# --solutions--

```py
def generate_signature(name, title, company):
    first_letter = name[0].upper()
    if first_letter in "ABCDEFGHI":
        prefix = ">>"
    elif first_letter in "JKLMNOPQR":
        prefix = "--"
    else:
        prefix = "::"

    return f"{prefix}{name}, {title} at {company}"
```