id: 68f6587287ad1f4ad39b0c7f
title: "Challenge 93: Vowels and Consonants"
challengeType: 29
dashedName: challenge-93
---

# --description--

Ako dobijete niz slova, vratite niz podataka sa brojem samoglasnika i brojem suglasnika u tom nizu.

- Samoglasnici su `a`, `e`, `i`, `o`, `u` u bilo kom slučaju.
- Suglasnici su sva ostala slova u bilo kom slučaju.
- Ignorišite nealfabetičke karaktere.

Na primer, ako dobijete `"Hello World"`, vratite `[3, 7]`.

# --hints--

`count("Hello World")` bi trebalo da vrati `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

`count("JavaScript")` bi trebalo da vrati `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("JavaScript"), [3, 7])`)
}})
```

`count("Python")` bi trebalo da vrati `[1, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Python"), [1, 5])`)
}})
```

`count("freeCodeCamp")` bi trebalo da vrati `[5, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("freeCodeCamp"), [5, 7])`)
}})
```

`count("Hello, World!")` bi trebalo da vrati `[3, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

`count("The quick brown fox jumps over the lazy dog.")` bi trebalo da vrati `[11, 24]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(count("Hello World"), [3, 7])`)
}})
```

# --seed--

## --seed-contents--

```py
def count(s):

    return s
```

# --solutions--

```py
def count(s):
    vowels = "aeiou"
    consonants = "bcdfghjklmnpqrstvwxyz"
    v = 0
    c = 0

    for char in s:
        ch = char.lower()
        if ch in vowels:
            v += 1
        if ch in consonants:
            c += 1

    return [v, c]
```