id: 68ffb91507a5b645769328c4
title: "Izazov 101: Konverter zaglavlja u Markdownu"
challengeType: 29
dashedName: challenge-101
---

# --description--

Ako vam je dat niz znakova koji predstavlja zaglavlje sekcije u Markdownu, vratite odgovarajuće HTML zaglavlje sekcije.

Važeće zaglavlje sekcije u Markdownu mora:

- Počinjati sa nula ili više razmaka, praćeno sa
- 1 do 6 uzastopnih znakova hash (`#`), zatim
- Bar jednim razmakom. I na kraju,
- Tekstom zaglavlja sekcije.

Broj znakova hash određuje nivo zaglavlja sekcije. Na primer, jedan znak hash odgovara etiketi `h1`, a šest znakova hash odgovara etiketi `h6`.

Ako dati niz znakova ne odgovara gore navedenom formatu, vratite `"Invalid format"`.

Na primer, ako primite `"# My level 1 heading"`, vratite `"<h1>My level 1 heading</h1>"`.

Napomena: Konzola možda neće prikazati HTML etikete unutar niza tokom snimanja poruke. Proverite konzolu pregledača da vidite zapis sa ugrađenim etiketama.

# --hints--

`convert("# My level 1 heading")` should return `"<h1>My level 1 heading</h1>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("# My level 1 heading"), "<h1>My level 1 heading</h1>")`)
}})
```

`convert("My heading")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("My heading"), "Invalid format")`)
}})
```

`convert("##### My level 5 heading")` should return `"<h5>My level 5 heading</h5>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("##### My level 5 heading"), "<h5>My level 5 heading</h5>")`)
}})
```

`convert("#My heading")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("#My heading"), "Invalid format")`)
}})
```

`convert("  ###  My level 3 heading")` should return `"<h3>My level 3 heading</h3>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("  ###  My level 3 heading"), "<h3>My level 3 heading</h3>")`)
}})
```

`convert("####### My level 7 heading")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("####### My level 7 heading"), "Invalid format")`)
}})
```

`convert("## My #2 heading")` should return `"<h2>My #2 heading</h2>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert("## My #2 heading"), "<h2>My #2 heading</h2>")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert(heading):

    return heading
```

# --solutions--

```py
import re
def convert(heading):
    trimmed = heading.strip()

    if not re.match(r'^#{1,6}\s+\S+', trimmed):
        return "Invalid format"

    hash_count = trimmed.find(' ')
    text = trimmed[hash_count:].lstrip()

    return f"<h{hash_count}>{text}</h{hash_count}>"
```