id: 69162d64f96574d9bb629efb
title: "Challenge 110: Guesser reči"
challengeType: 29
dashedName: challenge-110
---

# --description--

Uzimajući niz od dva karaktera iste dužine, tajno reč i pogađanje, uporedite pogađanje sa tajnom rečju koristeći sledeća pravila:

- Tajna reč i pogađanje će sadržati samo velika slova (`"A"` do `"Z"`);
- Za svaki karakter u pogadanju, zamenite ga brojem na osnovu kako odgovara tajnoj reči:
  - `"2"` ako je karakter prisutan u tajnoj reči i na ispravnom mestu.
  - `"1"` ako je karakter prisutan u tajnoj reči, ali na pogrešnom mestu.
  - `"0"` ako karakter nije prisutan u tajnoj reči.
- Svaki karakter u tajnoj reči može se koristiti samo jednom.
- Prvo se dodeljuje potpuno poklapanje (`"2"`), zatim delimično poklapanje (`"1"`) se dodeljuje od leve ka desno za preostale karaktere.
- Ako se određeni karakter često pojavljuje u pogadanju, on može odgovarati samo onoliko puta koliko se pojavljuje u tajnoj reči.

Na primer, uzimajući tajnu reč `"APPLE"` i pogađanje `"POPPA"`, vratite `"10201"`:

Prvi karakter `"P"` nije na pravom mestu (`"1"`), `"O"` nije prisutan u tajnoj reči (`"0"`), drugi karakter `"P"` je na pravom mestu (`"2"`), treći karakter `"P"` je nula (`"0"`) jer su dva karaktera `"P"` u tajnoj reči korišćena, i `"A"` nije na pravom mestu (`"1"`).

# --hints--

`compare("APPLE", "POPPA")` should return `"10201"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("APPLE", "POPPA"), "10201")`)
}})
```

`compare("REACT", "TRACE")` should return `"11221"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("REACT", "TRACE"), "11221")`)
}})
```

`compare("DEBUGS", "PYTHON")` should return `"000000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("DEBUGS", "PYTHON"), "000000")`)
}})
```

`compare("JAVASCRIPT", "TYPESCRIPT")` should return `"0000222222"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222")`)
}})
```

`compare("ORANGE", "ROUNDS")` should return `"110200"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("ORANGE", "ROUNDS"), "110200")`)
}})
```

`compare("WIRELESS", "ETHERNET")` should return `"10021000"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare("WIRELESS", "ETHERNET"), "10021000")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare(word, guess):

    return word
```

# --solutions--

```py
def compare(word, guess):
    n = len(word)
    result = ["0"] * n
    used_secret = [False] * n

    for i in range(n):
        if guess[i] == word[i]:
            result[i] = "2"
            used_secret[i] = True

    for i in range(n):
        if result[i] == "0":
            for j in range(n):
                if not used_secret[j] and guess[i] == word[j]:
                    result[i] = "1"
                    used_secret[j] = True
                    break

    return "".join(result)
```