---
id: 69162d64f96574d9bb629f00
title: "Challenge 115: Markdown Ordered List Item Converter"
challengeType: 29
dashedName: challenge-115
---

# --description--

Ako vam je dat niz znakova koji predstavlja stavku u numerisanom spisku u Markdownu, vratite odgovarajući HTML niz znakova.

Validna stavka numerisanog spiska u Markdownu mora:

- Počinjati sa nula ili više razmaka, praćeno sa
- Brojem (jedan ili više) i tačkom (`.`), praćeno sa
- Bar jednim razmakom, a zatim
- Tekstom stavke spiska.

Ako niz znakova ne odgovara u potpunosti gore navedenom formatu, vratite `"Invalid format"`. U suprotnom, obavijte tekst stavke spiska unutar `li` oznaka i vratite niz znakova.

Na primer, ako primite `"1. My item"`, vratite `"<li>My item</li>"`.

Napomena: Konsola možda neće prikazivati HTML oznake unutar niza tokom snimanja poruke. Proverite konzolu pretraživača da vidite snimke sa ugrađenim oznakama.

# --hints--

`convert_list_item("1. My item")` should return `"<li>My item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1. My item"), "<li>My item</li>")`)
}})
```

`convert_list_item(" 1.  Another item")` should return `"<li>Another item</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(" 1.  Another item"), "<li>Another item</li>")`)
}})
```

`convert_list_item("1 . invalid item")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("1 . invalid item"), "Invalid format")`)
}})
```

`convert_list_item("2. list item text")` should return `"<li>list item text</li>"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("2. list item text"), "<li>list item text</li>")`)
}})
```

`convert_list_item(". invalid again")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item(". invalid again"), "Invalid format")`)
}})
```

`convert_list_item("A. last invalid")` should return `"Invalid format"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_list_item("A. last invalid"), "Invalid format")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_list_item(markdown):

    return markdown
```

# --solutions--

```py
import re
def convert_list_item(markdown):
    match = re.match(r'^\s*(\d+)\.\s+(.+)$', markdown)
    if not match:
        return "Invalid format"
    text = match.group(2)
    return f"<li>{text}</li>"
```