id: 69162d64f96574d9bb629efd
title: "Izazov 112: Detektor AI"
challengeType: 29
dashedName: challenge-112
---

# --description--

Ovaj izazov potekao je sa lansiranja ChatGPT-a 30. novembra 2022. godine.

Razmatrajući sekvencu karaktera u jednoj ili više rečenica, utvrdite da li je verovatno generisano od strane AI-ja koristeći sledeća pravila:

- Ima dve ili više crtica (`-`).
- Ima dve ili više parazisa (`()`). Tekst može biti unutar zagrada.
- Ima tri ili više reči sa 7 ili više slova.

- Reči su razdvojene samo jednim razmakom i sastoje se samo od slova (`A-Z`). Ne uključujte interpunkcijske znakove ili ne-slovočne karaktere kao deo reči.

Ako zadata rečenica ispunjava bilo koje od ovih pravila, vratite `"AI"`, inače, vratite `"Human"`.

# --hints--

`detect_ai("The quick brown fox jumped over the lazy dog.")` should return `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The quick brown fox jumped over the lazy dog."), "Human")`)
}})
```

`detect_ai("The hypersonic brown fox - jumped (over) the lazy dog.")` should return `"Human"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The hypersonic brown fox - jumped (over) the lazy dog."), "Human")`)
}})
```

`detect_ai("Yes - you're right! I made a mistake there - let me try again.")` should return `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("Yes - you're right! I made a mistake there - let me try again."), "AI")`)
}})
```

`detect_ai("The extraordinary students were studying vivaciously.")` should return `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The extraordinary students were studying vivaciously."), "AI")`)
}})
```

`detect_ai("The (excited) student was (coding) in the library.")` should return `"AI"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_ai("The (excited) student was (coding) in the library."), "AI")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_ai(text):

    return text
```

# --solutions--

```py
import re
def detect_ai(text):
    if len(re.findall(r'-', text)) >= 2:
        return "AI"

    if len(re.findall(r'\([^)]*\)', text)) >= 2:
        return "AI"

    words = text.split()
    long_word_count = sum(1 for w in words if len(re.sub(r'[^A-Za-z]', '', w)) >= 7)
    if long_word_count >= 3:
        return "AI"

    return "Human"
```