id: 69162d64f96574d9bb629efc
title: "Izazov 111: Trajektorija lopte"
challengeType: 29
dashedName: challenge-111
---

# --description--

Danasnji izazov potekao je iz video igre Pong, lansirane 29. novembra 1972. godine.

Uzimajući u obzir matricu (dati niz redova podataka) koja sadrži poziciju lopte (`2`), i prethodnu poziciju lopte (`1`), vratite matrica koordinate sledeće pozicije lopte.

- Lopta se uvek kreće po pravoj liniji.
- Smer kretanja zavisi od toga kako je lopta prešla iz `1` u `2`.
- Granice matrice smatraju se zidovima. Ako loptica udari:
  - o gornji ili donji zid, ona skokom menja svoj vertikalni smer.
  - o levi ili desni zid, ona skokom menja svoj horizontalni smer.
  - u ćošak, ona skokom menja oba smera.

# --hints--

`get_next_location([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]])` bi trebalo da vrati `[2, 3]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_location([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]]), [2, 3])`)
}})
```

`get_next_location([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]])` bi trebalo da vrati `[3, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_location([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]]), [3, 0])`)
}})
```

`get_next_location([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]])` bi trebalo da vrati `[1, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_location([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]]), [1, 2])`)
}})
```

`get_next_location([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]])` bi trebalo da vrati `[1, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_location([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]]), [1, 1])`)
}})
```

`get_next_location([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]])` bi trebalo da vrati `[2, 2]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_next_location([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]]), [2, 2])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_next_location(matrix):

    return matrix
```

# --solutions--

```py
def get_next_location(matrix):
    prev = None
    curr = None

    for r in range(len(matrix)):
        for c in range(len(matrix[0])):
            if matrix[r][c] == 1:
                prev = (r, c)
            if matrix[r][c] == 2:
                curr = (r, c)

    prev_row, prev_col = prev
    curr_row, curr_col = curr

    dir_x = curr_col - prev_col
    dir_y = curr_row - prev_row

    next_row = curr_row + dir_y
    next_col = curr_col + dir_x

    max_row = len(matrix) - 1
    max_col = len(matrix[0]) - 1

    if next_col < 0 or next_col > max_col:
        dir_x *= -1
    if next_row < 0 or next_row > max_row:
        dir_y *= -1

    next_row = curr_row + dir_y
    next_col = curr_col + dir_x

    return [next_row, next_col]
```