id: 691b559495c5cb5a37b9b483
title: "Izazov 123: Izgradnja Rimskih brojeva"
challengeType: 29
dashedName: challenge-123
---

# --description--

Ako dobijete kompletan broj, vratite njegov ekvivalent u rimskim brojevima.

Rimski brojevi koriste ove simbole:

| Symbol | Vrednost |
| ------ | ------- |
| I      | 1       |
| V      | 5       |
| X      | 10      |
| L      | 50      |
| C      | 100     |
| D      | 500     |
| M      | 1000    |

Rimski brojevi se pišu od najvećeg ka najmanjem, od levi ka desni koristeći sledeća pravila:

- Dodavanje se koristi kada simbol sledi drugi sa jednakom ili manjom vrednošću. Na primer, `18` se piše kao `XVIII` (10 + 5 + 1 + 1 + 1 = 18).
- Oduzimanje se koristi kada je manji simbol vidljiv pre većeg, predstavljajući 4 ili 9 u bilo kojoj poziciji. Na primer, 19 se piše kao `XIX` (10 + (10 - 1)).
- Nijedan simbol ne sme biti ponovljen više od tri puta uzastopno. Oduzimanje se koristi kada je potrebno napisati više od tri simbola uzastopno. Stoga, najveći broj koji možete napisati je 3999.

Evo još jednog primera: ako dobijete `1464`, vratite `"MCDLXIV"` (1000 + (500 - 100) + 50 + 10 + (5 - 1)).

# --hints--

`to_roman(18)` treba da vrati `"XVIII"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(18), "XVIII")`)
}})
```

`to_roman(19)` treba da vrati `"XIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(19), "XIX")`)
}})
```

`to_roman(1464)` treba da vrati `"MCDLXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(1464), "MCDLXIV")`)
}})
```

`to_roman(2025)` treba da vrati `"MMXXV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(2025), "MMXXV")`)
}})
```

`to_roman(3999)` treba da vrati `"MMMCMXCIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(to_roman(3999), "MMMCMXCIX")`)
}})
```

# --seed--

## --seed-contents--

```py
def to_roman(num):

    return num
```

# --solutions--

```py
def to_roman(num):
    symbols = [
        ("M", 1000),
        ("CM", 900),
        ("D", 500),
        ("CD", 400),
        ("C", 100),
        ("XC", 90),
        ("L", 50),
        ("XL", 40),
        ("X", 10),
        ("IX", 9),
        ("V", 5),
        ("IV", 4),
        ("I", 1)
    ]

    result = ""
    for sym, val in symbols:
        while num >= val:
            result += sym
            num -= val
    return result
```