id: 691b559495c5cb5a37b9b487
title: "Izazov 127: Provera brzine"
challengeType: 29
dashedName: challenge-127
---

# --description--

Na osnovu brzine koju prelazite u miljama na sat (MPH) i ograničenja brzine u kilometrima na sat (KPH), utvrdite da li prelazite brzinu i da ćete dobiti upozorenje ili kaznu.

- 1 mila je jednako 1,60934 kilometra.
- Ako vozite brzinom koja ne prelazi ili je jednaka ograničenju brzine, vratite `"Not Speeding"`.
- Ako vozite brzinom koja prelazi ograničenje brzine za maksimalno KPH 5, vratite `"Warning"`.
- Ako vozite brzinom koja prelazi ograničenje brzine za više od KPH 5, vratite `"Ticket"`.

# --hints--

`speed_check(30, 70)` should return `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(30, 70), "Not Speeding")`)
}})
```

`speed_check(40, 60)` should return `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 60), "Warning")`)
}})
```

`speed_check(40, 65)` should return `"Not Speeding"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(40, 65), "Not Speeding")`)
}})
```

`speed_check(60, 90)` should return `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(60, 90), "Ticket")`)
}})
```

`speed_check(65, 100)` should return `"Warning"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(65, 100), "Warning")`)
}})
```

`speed_check(88, 40)` should return `"Ticket"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(speed_check(88, 40), "Ticket")`)
}})
```

# --seed--

## --seed-contents--

```py
def speed_check(speed_mph, speed_limit_kph):

    return speed_mph
```

# --solutions--

```py
def speed_check(speed_mph, speed_limit_kph):
    speed_kph = speed_mph * 1.60934

    if speed_kph <= speed_limit_kph:
        return "Not Speeding"

    over = speed_kph - speed_limit_kph

    if over <= 5:
        return "Warning"

    return "Ticket"
```