id: 691b559495c5cb5a37b9b489
title: "Izazov 129: Markdown Parser za citate (Blockquote)"
challengeType: 29
dashedName: challenge-129
---

# --description--

Ako vam je dat niz znakova koji sadrži blokcitiranje u Markdownu, vratite odgovarajući HTML niz znakova.

Blokcitiranje u Markdownu je bilo koji red koji:

- Počinje sa nula ili više razmaka
- Praćen znakiem većegovića (>)
- Zatim, jedan ili više razmaka
- I na kraju, tekst cita.

Vratite tekst cita obavijen HTML oznakama za otvaranje i zatvaranje `blockquote`.

Na primer, ako primite `"> Ovo je citat"`, vratite `<blockquote >Ovo je citat</blockquote >`.

Napomena: Konzola možda neće prikazati HTML oznake unutar niza tokom snimanja poruke. Proverite pregledačku konzolu da vidite snimke sa ugrađenim oznakama.

# --hints--

`parse_blockquote("> This is a quote")` bi trebalo da vrati `"<blockquote >This is a quote</blockquote >"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("> This is a quote"), "<blockquote >This is a quote</blockquote >")`)
}})
```

`parse_blockquote(" > This is also a quote")` bi trebalo da vrati `"<blockquote >This is also a quote</blockquote >"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote(" > This is also a quote"), "<blockquote >This is also a quote</blockquote >")`)
}})
```

`parse_blockquote("  >    So  Is  This")` bi trebalo da vrati `"<blockquote >So  Is  This</blockquote >"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_blockquote("  >    So  Is  This"), "<blockquote >So  Is  This</blockquote >")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_blockquote(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_blockquote(markdown):
    pattern = r"^\s*>\s+(.+)$"
    match = re.match(pattern, markdown)

    text = match.group(1)
    return f"<blockquote>{text}</blockquote>"
```