id: 6925e2068081f40f549ced1a
title: "Challenge 136: Markdown Image Parser"
challengeType: 29
dashedName: challenge-136
---

# --description--

Kada pronađete niz znakova slike u Markdownu, vratite odgovarajući niz znaka HTML.

Markdown slika ima sledeći format: `"![alt text](image_url)"`. Gde:

- `alt text` je opis slike (vrednost atributa `alt`).
- `image_url` je izvorni URL slike (vrednost atributa `src`).

Vratite niz znakova HTML taga `img` sa `src`-om postavljeno na URL slike i `alt`-om postavljeno na tekst alt.

Na primer, ako pronađete `"![Cute cat](cat.png)"`, vratite `'<img src="cat.png" alt="Cute cat">'`;

- Uverite se da su tag, redosled atributa, razmak i upotreba navodnika isti kao gore.

Napomena: Konola možda neće prikazivati HTML tagove u nizu znakova prilikom snimanja poruke — proverite konzolu pretraživača da vidite snimke sa priloženim tagovima.

# --hints--

`parse_image("![Cute cat](cat.png)")` should return `'<img src="cat.png" alt="Cute cat">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">')`)
}})
```

`parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` should return `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">')`)
}})
```

`parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` should return `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_image("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">')`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_image(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_image(markdown):
    match = re.search(r'!\[(.*?)\]\((.*?)\)', markdown)
    if not match:
        return ""
    alt, src = match.groups()
    return f'<img src="{src}" alt="{alt}">'
```