id: 6925e2068081f40f549ced1b
title: "Izazov 137: Generator snježadine"
challengeType: 29
dashedName: challenge-137
---

# --description--

Ako vam je dat niz znakova sa više redova koji koriste karakter nove linije (`\n`) za predstavljanje prekida redova, vratite novi niz gde je svaki red savršeno ogledalan i spojen na kraj originalnog reda.

- Ogledajte red tako što ćete obrnuti sve sekvence znakova, uključujući razmake.

Na primer, ako primite `"* \n *\n* "`, što se prikazuje u konzoli kao:

```sh
* 
 *
* 
```

Vratite `"*  *\n ** \n*  *"`, što se prikazuje u konzoli kao:

```sh
*  *
 ** 
*  *
```

Obratite posebnu pažnju na razmake u datom i vraćenom nizu. Uverite se da ne izbacujete nijedan razmak.

# --hints--

`generate_snowflake("* \n *\n* ")` should return `"*  *\n ** \n*  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("* \\n *\\n* "), "*  *\\n ** \\n*  *")`)
}})
```

`generate_snowflake("X=~")` should return `"X=~~=X"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("X=~"), "X=~~=X")`)
}})
```

`generate_snowflake(" X  \n  v \nX--=\n  ^ \n X  ")` should return `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake(" X  \\n  v \\nX--=\\n  ^ \\n X  "), " X    X \\n  v  v  \\nX--==--X\\n  ^  ^  \\n X    X ")`)
}})
```

`generate_snowflake("*   *\n * * \n* * *\n * * \n*   *")` should return `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*   *\\n * * \\n* * *\\n * * \\n*   *"), "*   **   *\\n * *  * * \\n* * ** * *\\n * *  * * \\n*   **   *")`)
}})
```

`generate_snowflake("*  -\n * -\n*  -")` should return `"*  --  *\n * -- * \n*  --  *"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(generate_snowflake("*  -\\n * -\\n*  -"), "*  --  *\\n * -- * \\n*  --  *")`)
}})
```

# --seed--

## --seed-contents--

```py
def generate_snowflake(crystals):

    return crystals
```

# --solutions--

```py
def generate_snowflake(crystals):
    lines = crystals.split("\n")
    mirrored = []

    for line in lines:
        reversed_line = line[::-1]
        mirrored.append(line + reversed_line)

    return "\n".join(mirrored)
```