id: 6925e2068081f40f549ced1c
title: "Challenge 138: Zbir delitelja"
challengeType: 29
dashedName: challenge-138
---

# --description--

Ako vam je dat pozitivan ceo broj, vratite zbir svih njegovih delitelja.

- Delitelj je bilo koji ceo broj koji deli taj broj bez ostatka (ostatak je `0`).
- Brojte svaki delitelj samo jednom.

Na primer, ako dobijete `6`, vratite `12` jer su delitelji `6` brojevi `1`, `2`, `3` i `6`, a njihov zbir je `12`.

# --hints--

`sum_divisors(6)` should return `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(6), 12)`)
}})
```

`sum_divisors(13)` should return `14`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(13), 14)`)
}})
```

`sum_divisors(28)` should return `56`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(28), 56)`)
}})
```

`sum_divisors(84)` should return `224`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(84), 224)`)
}})
```

`sum_divisors(549)` should return `806`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(549), 806)`)
}})
```

`sum_divisors(9348)` should return `23520`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sum_divisors(9348), 23520)`)
}})
```

# --seed--

## --seed-contents--

```py
def sum_divisors(n):

    return n
```

# --solutions--

```py
def sum_divisors(n):
    total = 0
    for i in range(1, n + 1):
        if n % i == 0:
            total += i
    return total
```