id: 69272dcf1c24b44fd79137c5
title: "Izazov 142: Sabiranje stringa"
challengeType: 29
dashedName: challenge-142
---

# --description--

Ako vam je dat niz karaktera koji uključuje brojeve i druge slova, vratite zbir svih brojeva prisutnih u tom nizu.

- Razmatrajte brojeve koji slede jedan za drugim kao jednu celobrojnu vrednost. Na primer, `"13"` se računa kao 13, a ne 1 + 3.
- Ignorišite sva slova koja nisu brojevi.

# --hints--

`string_sum("3apples2bananas")` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("3apples2bananas"), 5)`)
}})
```

`string_sum("10cats5dogs2birds")` should return `17`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("10cats5dogs2birds"), 17)`)
}})
```

`string_sum("125344")` should return `125344`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("125344"), 125344)`)
}})
```

`string_sum("a1b20c300")` should return `321`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("a1b20c300"), 321)`)
}})
```

`string_sum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` should return `1653`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(string_sum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653)`)
}})
```

# --seed--

## --seed-contents--

```py
def string_sum(s):

    return s
```

# --solutions--

```py
import re
def string_sum(s):
    numbers = re.findall(r'\d+', s)
    return sum(int(num) for num in numbers)
```