id: 69306364df283fcaff2e1ad7
title: "Izazov 146: Sedište za levokakarnika na stolu"
challengeType: 29
dashedName: challenge-146
---

# --description--

Dato je sto dimenzija 4x2 (redova redova) koji predstavlja raspored sedišta za obrok, odredite koliko stolica može da zauzme osoba koja piše levom rukom.

- Osoba koja piše levom rukom ne može da sedi na mestu gde bi osoba koja piše desnom rukom sedela na stolicu sa leve strane te osobe.

U datoj matrici:

- `"R"` je stolica zauzeta od strane osobe koja piše desnom rukom.
- `"L"` je stolica zauzeta od strane osobe koja piše levom rukom.
- `"U"` je prazna stolica.
- Samo prazne stolice su dostupne za sedenje.
- Stolice u gornjem redu okrenute su "dole", a stolice u donjem redu okrenute su "gore" (kao sto), pa zavise da li je levo ili desno od orijentacije stolice.
- Uglove stolice imaju samo jednu stolicu pored sebe.

Na primer, na datoj tabeli:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

Gornja leva stolica ne može biti zauzeta jer sa leve strane stoji osoba koja piše desnom rukom. Druga dve prazne stolice mogu biti zauzete jer sa leve strane nema osobe koja piše desnom rukom.

# --hints--

`find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2)`)
}})
```

`find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` should return `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0)`)
}})
```

`find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1)`)
}})
```

`find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(find_left_handed_seats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5)`)
}})
```

# --seed--

## --seed-contents--

```py
def find_left_handed_seats(table):

    return table
```

# --solutions--

```py
def find_left_handed_seats(table):
    available_seats = 0
    top_row, bottom_row = table

    for i in range(len(top_row)):
        if top_row[i] == "U":
            if i == len(top_row) - 1:
                available_seats += 1
            elif top_row[i + 1] != "R":
                available_seats += 1

    for i in range(len(bottom_row)):
        if bottom_row[i] == "U":
            if i == 0:
                available_seats += 1
            elif bottom_row[i - 1] != "R":
                available_seats += 1

    return available_seats
```