id: 69306364df283fcaff2e1ad6
title: "Izazov 145: N-ti Fibonacci Broj"
challengeType: 29
dashedName: challenge-145
---

# --description--

Ako vam je dat ceo broj `n`, vratite n-ti broj u Fibonacci sekvenci.

Fibonacci sekvenca je niz brojeva gde je svaki broj zbir prethodnih dva broja. Prvih deset brojeva u sekvenci su `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

# --hints--

`nth_fibonacci(4)` bi trebalo da vrati `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(4), 2)`)
}})
```

`nth_fibonacci(10)` bi trebalo da vrati `34`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(10), 34)`)
}})
```

`nth_fibonacci(15)` bi trebalo da vrati `377`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(15), 377)`)
}})
```

`nth_fibonacci(40)` bi trebalo da vrati `63245986`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(40), 63245986)`)
}})
```

`nth_fibonacci(75)` bi trebalo da vrati `1304969544928657`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(nth_fibonacci(75), 1304969544928657)`)
}})
```

# --seed--

## --seed-contents--

```py
def nth_fibonacci(n):

    return n
```

# --solutions--

```py
def nth_fibonacci(n):
    if n == 1:
        return 0
    if n == 2:
        return 1

    a, b = 0, 1
    for _ in range(3, n + 1):
        a, b = b, a + b
    return b
```