id: 69306364df283fcaff2e1ad9
title: "Izazov 148: Pritisak u gumama"
challengeType: 29
dashedName: challenge-148
---

# --description--

Dato je niz podataka sa četiri broja koji predstavljaju pritisak u gumama u psi za četiri gume vašeg vozila, i drugi niz od dva broja koji predstavljaju minimalni i maksimalni nivo pritiska za vaše gume u barima. Vratite niz od četiri karaktera koji određuje status svake gume.

- 1 bar je jednak 14.5038 psi.

Vratite niz podataka sa sledećim vrednostima za svaku gumu:

- `"Low"` ako je pritisak u gumama niži od dozvoljenog minimalnog nivoa.
- `"Good"` ako se nalazi između dozvoljenog minimalnog i maksimalnog nivoa.
- `"High"` ako je veći od dozvoljenog maksimalnog nivoa.

# --hints--

`tire_status([32, 28, 35, 29], [2, 3])` should return `["Good", "Low", "Good", "Low"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([32, 28, 35, 29], [2, 3]), ["Good", "Low", "Good", "Low"])`)
}})
```

`tire_status([32, 28, 35, 30], [2, 2.3])` should return `["Good", "Low", "High", "Good"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([32, 28, 35, 30], [2, 2.3]), ["Good", "Low", "High", "Good"])`)
}})
```

`tire_status([29, 26, 31, 28], [2.1, 2.5])` should return `["Low", "Low", "Good", "Low"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([29, 26, 31, 28], [2.1, 2.5]), ["Low", "Low", "Good", "Low"])`)
}})
```

`tire_status([31, 31, 30, 29], [1.5, 2])` should return `["High", "High", "High", "Good"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([31, 31, 30, 29], [1.5, 2]), ["High", "High", "High", "Good"])`)
}})
```

`tire_status([30, 28, 30, 29], [1.9, 2.1])` should return `["Good", "Good", "Good", "Good"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(tire_status([30, 28, 30, 29], [1.9, 2.1]), ["Good", "Good", "Good", "Good"])`)
}})
```

# --seed--

## --seed-contents--

```py
def tire_status(pressures_psi, range_bar):

    return pressures_psi
```

# --solutions--

```py
def tire_status(pressures_psi, range_bar):
    psi_to_bar = 1 / 14.5038
    min_bar, max_bar = range_bar
    result = []

    for psi in pressures_psi:
        bar = psi * psi_to_bar
        if bar < min_bar:
            result.append("Low")
        elif bar > max_bar:
            result.append("High")
        else:
            result.append("Good")
    return result
```