id: 69373793f5a867f769cde135
title: "Challenge 150: Markdown Unordered List Parser"
challengeType: 29
dashedName: challenge-150
---

# --description--

Izvući validan neuređeni spisak iz Markdowna, vratiti odgovarajući HTML niz.

Neuređeni spisak ima jedan ili više stavki spiska. Validna stavka spiska izgleda u svom redu i:

- Počinje sa crticom (`"-"`), praćeno sa
- Bareman razmakom, zatim
- Tekstom stavke spiska.

Spisak se vraća kao jedan niz znakova sa novim redovima odvojenima karakterom za novi red (`"\n"`). Ne uključujte karaktere novog reda u tekst stavke.

Omotajte svaku stavku spiska unutar HTML oznaka `li`, a ceo spisak stavki unutar oznaka `ul`.

Na primer, ako primite `"- Item A\n- Item B"`, vratiti `<ul><li class="list-item">Item A</li class="list-item"><li class="list-item">Item B</li class="list-item></ul class="list-item">`.

Napomena: Konzola možda neće prikazati HTML oznake unutar niza tokom snimanja poruke. Proverite pregledački konzolu da vidite snimke sa ugrađenim oznakama.

# --hints--

`parse_unordered_list("- Item A\n- Item B")` bi trebalo da vrati `"<ul class="list-item"><li class="list-item">Item A</li class="list-item"><li class="list-item">Item B</li class="list-item></ul class="list-item">"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("- Item A\\n- Item B"), "<ul class="list-item"><li class="list-item">Item A</li class="list-item"><li class="list-item">Item B</li class="list-item></ul class="list-item">")`)
}})
```

`parse_unordered_list("-  JavaScript\n-  Python")` bi trebalo da vrati `"<ul class="list-item"><li class="list-item">JavaScript</li class="list-item"><li class="list-item">Python</li class="list-item></ul class="list-item">"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("-  JavaScript\\n-  Python"), "<ul class="list-item"><li class="list-item">JavaScript</li class="list-item"><li class="list-item">Python</li class="list-item></ul class="list-item">")`)
}})
```

`parse_unordered_list("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla")` bi trebalo da vrati `"<ul class="list-item"><li class="list-item">2 C Flour</li class="list-item"><li class="list-item">1/2 C Sugar</li class="list-item"><li class="list-item">1 Tsp Vanilla</li class="list-item></ul class="list-item">"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("- 2 C Flour\\n- 1/2 C Sugar\\n- 1 Tsp Vanilla"), "<ul class="list-item"><li class="list-item">2 C Flour</li class="list-item"><li class="list-item">1/2 C Sugar</li class="list-item"><li class="list-item">1 Tsp Vanilla</li class="list-item></ul class="list-item">")`)
}})
```

`parse_unordered_list("- A-1\n- A-2\n- B-1")` bi trebalo da vrati `"<ul class="list-item"><li class="list-item">A-1</li class="list-item"><li class="list-item">A-2</li class="list-item"><li class="list-item">B-1</li class="list-item></ul class="list-item">"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_unordered_list("- A-1\\n- A-2\\n- B-1"), "<ul class="list-item"><li class="list-item">A-1</li class="list-item"><li class="list-item">A-2</li class="list-item"><li class="list-item">B-1</li class="list-item></ul class="list-item">")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_unordered_list(markdown):

    return markdown
```

# --solutions--

```py
def parse_unordered_list(markdown):
    lines = markdown.split("\n")
    list_items = []

    for line in lines:
        list_items.append(f"<li class="list-item">{line[1:].strip()}</li>")

    return f"<ul class="list-item>{''.join(list_items)}</ul>"
```