---
id: 6939b873185d8e00d453563f
title: "Izazov 160: Pokreti konja"
challengeType: 29
dashedName: challenge-160
---

# --description--

Iz pozicije konja na šahovskoj tabli, vratite broj validnih kvadrata na koje konj može da se pomakne.

Standardna šahovska tabla je dimenzija 8x8, a redovi su označeni slovima `A` do `H` (od leve ka desno) i kolone brojevima `1` do `8` (od dole ka gore). Izgleda ovako:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

Konj se kreće u "L" pravcu: dva kvadrata u jednom pravcu (vertikalno ili horizontalno), i jedan kvadrat u pravcu okomito na njega.

To znači da konj može da se pomakne do osam mogućih pozicija, ali manje ako je blizu granica table. Na primer, ako je konj na `A1`, može samo da se pomeri na `B3` ili `C2`.

# --hints--

`knight_moves("A1")` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("A1"), 2)`)
}})
```

`knight_moves("D4")` should return `8`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("D4"), 8)`)
}})
```

`knight_moves("G6")` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("G6"), 6)`)
}})
```

`knight_moves("B8")` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("B8"), 3)`)
}})
```

`knight_moves("H3")` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(knight_moves("H3"), 4)`)
}})
```

# --seed--

## --seed-contents--

```py
def knight_moves(position):

    return position
```

# --solutions--

```py
def knight_moves(position):
    col = ord(position[0]) - ord('A')
    row = int(position[1]) - 1

    moves = [
        (2, 1), (2, -1),
        (-2, 1), (-2, -1),
        (1, 2), (1, -2),
        (-1, 2), (-1, -2)
    ]

    valid = 0

    for dx, dy in moves:
        new_col = col + dx
        new_row = row + dy

        if 0 <= new_col < 8 and 0 <= new_row < 8:
            valid += 1

    return valid
```