id: 694596b0585c11170ac7c7fa
title: "Challenge 162: Energy Consumption"
challengeType: 29
dashedName: challenge-162
---

# --description--

Na osnovu broja kalorija potrošenih tokom vežbanja i broja vati-sati korišćenih od strane vašeg elektronskog uređaja tokom tih vežbanja, utvrdite koji je potrošio više energije.

Da biste ih uporedili, pretvorite obe vrednosti u džoule koristeći sledeće konverzije:

- 1 Kalorija je jednako 4184 džoula.
- 1 vati-sat je jednako 3600 džoula.

Vraćanje (Return):

- `"Workout"` ako su vežbanja potrošila više energije.
- `"Devices"` ako je uređaj potrošio više energije.
- `"Equal"` ako su sve potrošile jednak količinu energije.

# --hints--

`compare_energy(250, 50)` should return `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(250, 50), "Workout")`)
}})
```

`compare_energy(100, 200)` should return `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(100, 200), "Devices")`)
}})
```

`compare_energy(450, 523)` should return `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(450, 523), "Equal")`)
}})
```

`compare_energy(300, 75)` should return `"Workout"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(300, 75), "Workout")`)
}})
```

`compare_energy(200, 250)` should return `"Devices"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(200, 250), "Devices")`)
}})
```

`compare_energy(900, 1046)` should return `"Equal"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(compare_energy(900, 1046), "Equal")`)
}})
```

# --seed--

## --seed-contents--

```py
def compare_energy(calories_burned, watt_hours_used):

    return calories_burned
```

# --solutions--

```py
def compare_energy(calories_burned, watt_hours_used):
    workout_energy = calories_burned * 4184
    device_energy = watt_hours_used * 3600

    if workout_energy > device_energy:
        return "Workout"
    if device_energy > workout_energy:
        return "Devices"
    return "Equal"
```