id: 694596b0585c11170ac7c7fc
title: "Challenge 164: Markdown Inline Code Parser"
challengeType: 29
dashedName: challenge-164
---

# --description--

Ako vam je dat niz znakova Markdown koji sadrži jednu ili više sekvenci koda unutar linije, vratite odgovarajući niz HTML znakova.

Sekvencijalni kod u liniji u Markdownu koristi jedan backtick karakter (``` ` ```) na početku i kraju teksta sekvence koda.

Vratite zadati niz tako da su sve sekvencije koda zamijenjene HTML tagovima `<code/>`.

Na primjer, ako primite niz ```"Use `let` to declare the variable."```, vratite `"Koristite <code>let</code> za deklarisanje varijable."`.

Napomena: Konsola možda neće prikazati HTML tagove unutar niza tokom snimanja poruke. Proverite konzolu pretraživača da vidite snimke sa ugrađenim tagovima.

# --hints--

```parse_inline_code("Use `let` to declare the variable.")``` bi trebalo da vrati `"Use <code>let</code> to declare the variable."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Use \`let\` to declare the variable."), "Use <code>let</code> to declare the variable.")`)
}})
```

```parse_inline_code("Use `let` or `const` to declare a variable.")``` bi trebalo da vrati `"Use <code>let</code> or <code>const</code> to declare a variable."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Use \`let\` or \`const\` to declare a variable."), "Use <code>let</code> or <code>const</code> to declare a variable.")`)
}})
```

```parse_inline_code("Run `npm install` then `npm start`.")``` bi trebalo da vrati `"Run <code>npm install</code> then <code>npm start</code>."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_inline_code("Run \`npm install\` then \`npm start\`."), "Run <code>npm install</code> then <code>npm start</code>.")`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_inline_code(markdown):

    return markdown
```

# --solutions--

```py
import re
def parse_inline_code(markdown):
    return re.sub(r"`([^`]+)`", r"<code>\1</code>", markdown)
```