id: 696655d24b614176d4c9b78d
title: "Challenge 170: Odd or Even Day"
challengeType: 29
dashedName: challenge-170
---

# --description--

Pored date-a (broj milisekundi od početka Unix epoha), vratite:

- `"odd"` ako je broj dana meseca za taj timestamp neparizan.
- `"even"` ako je broj dana meseca za taj timestamp paran.

Na primer, ako date `1769472000000`, timestamp za 27. januar 2026. godine, vratite `"odd"` jer je broj dana (`27`) neparizan broj.

Napomena: Timestamp je u milisekundama i morate koristiti datum po UTC vremenskoj zoni, a ne lokalnu vremensku zonu.

# --hints--

`odd_or_even_day(1769472000000)` should return `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769472000000), "odd")`)
}})
```

`odd_or_even_day(1769444440000)` should return `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1769444440000), "even")`)
}})
```

`odd_or_even_day(6739456780000)` should return `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(6739456780000), "odd")`)
}})
```

`odd_or_even_day(1)` should return `"odd"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(1), "odd")`)
}})
```

`odd_or_even_day(86400000)` should return `"even"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(odd_or_even_day(86400000), "even")`)
}})
```

# --seed--

## --seed-contents--

```py
def odd_or_even_day(timestamp):

    return timestamp
```

# --solutions--

```py
from datetime import datetime

def odd_or_even_day(timestamp):
    dt = datetime.utcfromtimestamp(timestamp / 1000)
    day = dt.day
    return "even" if day % 2 == 0 else "odd"
```