id: 69738771fb5a7b8b24cca2a4
title: "Izazov 178: Skraćivanje teksta"
challengeType: 29
dashedName: challenge-178
---

# --description--

Ako vam je dat niz znakova, vratite ga takav kakav jeste ako ima 20 ili manje znakova. Ako ima više od 20 znakova, smanjite ga na prvih 17 znakova, a zatim dodajte `"..."` na kraj (da bi ukupan broj bio 20 znaka) i vratite rezultat.

# --hints--

`truncate_text("Hello, world!")` bi trebalo da vrati `"Hello, world!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Hello, world!"), "Hello, world!")`)
}})
```

`truncate_text("This string should get truncated.")` bi trebalo da vrati `"This string shoul..."`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("This string should get truncated."), "This string shoul...")`)
}})
```

`truncate_text("Exactly twenty chars")` bi trebalo da vrati `"Exactly twenty chars"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("Exactly twenty chars"), "Exactly twenty chars")`)
}})
```

`truncate_text(".....................")` bi trebalo da vrati `"...................."`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(truncate_text("....................."), "....................")`)
}})
```

# --seed--

## --seed-contents--

```py
def truncate_text(text):

    return text
```

# --solutions--

```py
def truncate_text(text):
    if len(text) <= 20:
        return text
    return text[:17] + "..."
```