id: 697a49e9860d24853adef67d
title: "Izazov 192: Zimske Igre 2026, Dan 13: Nordički kombinovani skok"
challengeType: 29
dashedName: challenge-192
---

# --description--

Dato setu sa skokovskih rezultata sportista, izračunajte njihovo kašnjenje u početku trke Nordic Combined.

Sportista sa najvišim skorom na skoku počinje prvi (kašnjenje od 0 sekundi). Svi ostali sportisti počinju kasnije, u zavisnosti od razlike njihovog skora sa skoka od najboljeg skora.

Da biste izračunali kašnjenje za svakog sportistu, oduzmite skor na skoku sportiste od apsolutno najboljeg skora skoka, a zatim rezultat pomnožite sa 1.5. Zaokružite kašnjenje do sledećeg celog broja.

# --hints--

`calculate_start_delays([120, 110, 125])` should return `[8, 23, 0]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([120, 110, 125]), [8, 23, 0])`)
}})
```

`calculate_start_delays([118, 125, 122, 120])` should return `[11, 0, 5, 8]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([118, 125, 122, 120]), [11, 0, 5, 8])`)
}})
```

`calculate_start_delays([100, 105, 95, 110, 120, 115, 108])` should return `[30, 23, 38, 15, 0, 8, 18]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([100, 105, 95, 110, 120, 115, 108]), [30, 23, 38, 15, 0, 8, 18])`)
}})
```

`calculate_start_delays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124])` should return `[3, 11, 6, 18, 21, 15, 8, 26, 0, 12]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(calculate_start_delays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124]), [3, 11, 6, 18, 21, 15, 8, 26, 0, 12])`)
}})
```

# --seed--

## --seed-contents--

```py
def calculate_start_delays(jump_scores):

    return jump_scores
```

# --solutions--

```py
import math
def calculate_start_delays(jump_scores):
    best_jump = max(jump_scores)
    return [math.ceil((best_jump - score) * 1.5) for score in jump_scores]
```