id: 698a1a73ade5ac0e19180fa0
title: "Izazov 197: Kompatibilnost krvnih grupa"
challengeType: 29
dashedName: challenge-197
---

# --description--

Uzimajući u obzir krvnu grupu donatora i krvnu grupu primaoca, utvrdite da li donator može dati krv primaocu.

Svaka grupa krvi ima:

- Slovo: `"A"`, `"B"`, `"AB"`, ili `"O"`
- I Rh faktor: `"+"` ili `"-"`

Grupe krvi će biti jedno važeće slovo praćeno Rh faktorom. Na primer, `"AB+"` i `"O-"` su validne grupe krvi.

Pravila za Slova (Letter Rules):

- `"O"` može dati krv bilo kojoj drugoj grupi slova.
- `"A"` može dati krv `"A"` i `"AB"`.
- `"B"` može dati krv `"B"` i `"AB"`.
- `"AB"` može dati krv samo `"AB"`.

Pravila za Rh (Rh Rules):

- Negativan (`"-"`) može dati krv `"-"` i `"+"`.
- Pozitivan (`"+"`) može dati krv samo `"+"`.

I pravila slova i Rh moraju biti ispunjena kako bi donator mogao dati krv primaocu.

# --hints--

`can_donate("B+", "B+")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B+", "B+"), True)`)
}})
```

`can_donate("O-", "AB-")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O-", "AB-"), True)`)
}})
```

`can_donate("O+", "A-")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O+", "A-"), False)`)
}})
```

`can_donate("A+", "AB+")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("A+", "AB+"), True)`)
}})
```

`can_donate("A-", "B-")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("A-", "B-"), False)`)
}})
```

`can_donate("B-", "AB+")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B-", "AB+"), True)`)
}})
```

`can_donate("B-", "A+")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("B-", "A+"), False)`)
}})
```

`can_donate("O-", "O+")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O-", "O+"), True)`)
}})
```

`can_donate("O+", "O-")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("O+", "O-"), False)`)
}})
```

`can_donate("AB+", "AB-")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(can_donate("AB+", "AB-"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def can_donate(donor, recipient):

    return donor
```

# --solutions--

```py
def can_donate(donor, recipient):
    donor_type = donor[:-1]
    donor_rh = donor[-1]

    recipient_type = recipient[:-1]
    recipient_rh = recipient[-1]

    abo_compatibility = {
        "O": ["A", "B", "AB", "O"],
        "A": ["A", "AB"],
        "B": ["B", "AB"],
        "AB": ["AB"]
    }

    abo_match = recipient_type in abo_compatibility[donor_type]
    rh_match = (
        donor_rh == "-" or
        (donor_rh == "+" and recipient_rh == "+")
    )

    return abo_match and rh_match
```