id: 698a1a73ade5ac0e19180fa9
title: "Challenge 206: Playing Card Values"
challengeType: 29
dashedName: challenge-206
---

# --description--

Kada dobijete niz podataka igračkih karata, vratite novi niz podataka sa numeričkom vrednošću svake karte.

Card Values:

- As (`"A"`) ima vrednost 1.
- Karte sa brojevima (`"2"` - `"10"`) imaju vrednost jednaku broju koji se nalazi ispred njih: 2 - 10, po redu.
- Face karte: Džek (`"J"`), Kvins (`"Q"`) i Kralj (`"K"`) svaka ima vrednost 10.

Suits:

- Svaka karta ima špilu: Pik (`"S"`), Trefl (`"C"`), Karo (`"D"`) ili Srca (`"H"`).

Card Format:

- Svaka karta je predstavljena kao niz slova: `"vrednostŠpila"`. Primer: `"AS"` je As od Pika, `"10H"` je Deset Srca, a `"QC"` je Kvins od Trefla.

# --hints--

`card_values(["3H", "4D", "5S"])` should return `[3, 4, 5]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["3H", "4D", "5S"]), [3, 4, 5])`)
}})
```

`card_values(["AS", "10S", "10H", "6D", "7D"])` should return `[1, 10, 10, 6, 7]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7])`)
}})
```

`card_values(["8D", "QS", "2H", "JC", "9C"])` should return `[8, 10, 2, 10, 9]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9])`)
}})
```

`card_values(["AS", "KS"])` should return `[1, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["AS", "KS"]), [1, 10])`)
}})
```

`card_values(["10H", "JH", "QH", "KH", "AH"])` should return `[10, 10, 10, 10, 1]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(card_values(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1])`)
}})
```

# --seed--

## --seed-contents--

```py
def card_values(cards):

    return cards
```

# --solutions--

```py
def card_values(cards):
    values = []
    for card in cards:
        value_str = card[:-1]

        if value_str == "A":
            value = 1
        elif value_str in ["J", "Q", "K"]:
            value = 10
        else:
            value = int(value_str)

        values.append(value)
    return values
```