---
id: 699c8e045ee7cb94ed2322dd
title: "Challenge 222: Equinox Shadows"
challengeType: 29
dashedName: challenge-222
---

# --description--

Danas je dan ekvinocija, kada je sunce direktno iznad ekvatora Zemlje i najviše u ponoć. Ako dobijete određeno vreme, odredite senku koju proizvodi vertikalna kolona visine 4 stopa.

- Vreme će biti niz znakova u formatu satova 24h `"HH:MM"` (primer: `"15:00"` je 3 popodne).
- Biće vam dato vreme samo u intervalima od 30 minuta.

Rules:

- Sunce zauskuva u 6 ujutru direktno "isto", i zalazi u 6 popodne direktno "zapad".
- Senka uvek pokazuje suprotnom smeru od sunca.
- Dužina senke (u stopama) je broj sati od ponoći, podignut na treći stepen.
- Nema senke pre nego što sunce izađe (pre 6 ujutru), nakon što sunce zađe (nakon 6 popodne), ili u ponoć.

Return:

- Ako postoji senka, vratite `"(dužina)ft (smer)"`. Primer: `"8ft west"`.
- Inače, vratite `"No shadow"`.

Na primer, ako dobijete `"10:00"`, vratite `"8ft west"` jer je 10 ujutru 2 sata od ponoći, pa je $2^3$ = 8 stopa, a senka pokazuje na zapad jer je sunce istočno u 10 ujutro.

# --hints--

`get_shadow("10:00")` should return `"8ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("10:00"), "8ft west")`)
}})
```

`get_shadow("15:00")` should return `"27ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("15:00"), "27ft east")`)
}})
```

`get_shadow("12:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("12:00"), "No shadow")`)
}})
```

`get_shadow("17:30")` should return `"166.375ft east"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("17:30"), "166.375ft east")`)
}})
```

`get_shadow("05:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("05:00"), "No shadow")`)
}})
```

`get_shadow("06:00")` should return `"216ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("06:00"), "216ft west")`)
}})
```

`get_shadow("18:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("18:00"), "No shadow")`)
}})
```

`get_shadow("07:30")` should return `"91.125ft west"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("07:30"), "91.125ft west")`)
}})
```

`get_shadow("00:00")` should return `"No shadow"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_shadow("00:00"), "No shadow")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_shadow(time):

    return time
```

# --solutions--

```py
def get_shadow(time):
    hour_str, minute_str = time.split(":")
    hours = int(hour_str)
    minutes = int(minute_str)
    time_in_hours = hours + minutes / 60

    if time_in_hours < 6 or time_in_hours >= 18 or time_in_hours == 12:
        return "No shadow"

    hours_from_noon = abs(12 - time_in_hours)
    length = hours_from_noon ** 3
    direction = "west" if time_in_hours < 12 else "east"

    return f"{length:g}ft {direction}"
```