id: 699c8e045ee7cb94ed2322d4
title: "Challenge 213: Word Length Converter"
challengeType: 29
dashedName: challenge-213
---

# --description--

Kada dobijete niz slova reči, vratite novi niz gde je svaka reč zamenjena svojom dužinom.

- Reči u datom nizu razdvojene su jednim razmakom između njih.
- Sačuvajte te razmake u vraćeni niz.

Na primer, ako dobijete `"hello world"`, vratite `"5 5"`.

# --hints--

`convert_words("hello world")` treba da vrati `"5 5"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("hello world"), "5 5")`)
}})
```

`convert_words("Thanks and happy coding")` treba da vrati `"6 3 5 6"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("Thanks and happy coding"), "6 3 5 6")`)
}})
```

`convert_words("The quick brown fox jumps over the lazy dog")` treba da vrati `"3 5 5 3 5 4 3 4 3"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("The quick brown fox jumps over the lazy dog"), "3 5 5 3 5 4 3 4 3")`)
}})
```

`convert_words("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl")` treba da vrati `"5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl"), "5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_words(s):

    return s
```

# --solutions--

```py
def convert_words(s):

    return " ".join(str(len(word)) for word in s.split(" "))
```