id: 69a890af247de743333bd4cd
title: "Challenge 224: Detektor za prženja kafe"
challengeType: 29
dashedName: challenge-224
---

# --description--

Dato je niz niz znakova koji predstavljaju zrna koja se koriste za pravljenje šolje kafe; odredite nivo prženja te šolje.

- Dat niz znakova će sadržati sljedeće karaktere, svaki od kojih predstavlja vrstu zrna:

- Apostrofski znak (`'`) je svijetlo prženo zrno sa vrijednošću od 1 poen svako.
  - Znak crtica (`-`) je srednje prženo zrno sa vrijednošću od 2 poena svako.
  - Tačka (`.`) je tamno prženo zrno sa vrijednošću od 3 poena svako.

- Nivo prženja određuje se prosjekom svih zrna.

Return:

- `"Light"` ako je prosjek manji od 1.75.
- `"Medium"` ako je prosjek između 1.75 i 2.5.
- `"Dark"` ako je prosjek veći od 2.5.

# --hints--

`detect_roast("''-''''''-'-''--''''")` treba da vrati `"Light"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("''-''''''-'-''--''''"), "Light")`)
}})
```

`detect_roast(".'-''-''..'''.-.-''-")` treba da vrati `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".'-''-''..'''.-.-''-"), "Medium")`)
}})
```

`detect_roast("--.''--'-''.--..-.--")` treba da vrati `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("--.''--'-''.--..-.--"), "Medium")`)
}})
```

`detect_roast("-...'-......-..-...-")` treba da vrati `"Dark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-...'-......-..-...-"), "Dark")`)
}})
```

`detect_roast(".--.-..-......----.'")` treba da vrati `"Medium"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast(".--.-..-......----.'"), "Medium")`)
}})
```

`detect_roast("..-..-..-..-....-.-.")` treba da vrati `"Dark"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("..-..-..-..-....-.-."), "Dark")`)
}})
```

`detect_roast("-'-''''''..-'.''-'.'")` treba da vrati `"Light"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(detect_roast("-'-''''''..-'.''-'.'"), "Light")`)
}})
```

# --seed--

## --seed-contents--

```py
def detect_roast(beans):

    return beans
```

# --solutions--

```py
def detect_roast(beans):
    total = 0

    for bean in beans:
        if bean == "'":
            total += 1
        elif bean == "-":
            total += 2
        elif bean == ".":
            total += 3

    avg = total / len(beans)

    if avg < 1.75:
        return "Light"
    elif avg <= 2.5:
        return "Medium"
    else:
        return "Dark"
```