id: 69b1028d6e265413d0198a2c
title: "Izazov 233: Budilnica"
challengeType: 29
dashedName: challenge-233
---

# --description--

Dato je niz niz string koji predstavlja vreme kada ste postavili alarm i niz znakova koji predstavlja vreme kada ste se probudili, utvrdite da li ste se probudili rano, na vreme ili kasno.

- Oba vremena će biti data u 24-satnom formatu `"HH:MM"`.

Vraćanje (Return):

- `"early"` ako ste se probudili pre vremena alarma.
- `"on time"` ako ste se probudili tačno na vreme alarma, ili unutar vremenskog promeka od 10 minuta kašnjenja nakon alarma.
- `"late"` ako ste se probudili više od 10 minuta nakon alarma.

Oba vremena su u istom danu.

# --hints--

`alarm_check("07:00", "06:45")` bi trebalo da vrati `"early"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("07:00", "06:45"), "early")`)
}})
```

`alarm_check("06:30", "06:30")` bi trebalo da vrati `"on time"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("06:30", "06:30"), "on time")`)
}})
```

`alarm_check("08:10", "08:15")` bi trebalo da vrati `"on time"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("08:10", "08:15"), "on time")`)
}})
```

`alarm_check("09:30", "09:45")` bi trebalo da vrati `"late"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("09:30", "09:45"), "late")`)
}})
```

`alarm_check("08:15", "08:25")` bi trebalo da vrati `"on time"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("08:15", "08:25"), "on time")`)
}})
```

`alarm_check("05:45", "05:56")` bi trebalo da vrati `"late"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("05:45", "05:56"), "late")`)
}})
```

`alarm_check("04:30", "04:00")` bi trebalo da vrati `"early"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(alarm_check("04:30", "04:00"), "early")`)
}})
```

# --seed--

## --seed-contents--

```py
def alarm_check(alarm_time, wake_time):

    return alarm_time
```

# --solutions--

```py
def alarm_check(alarm_time, wake_time):
    alarm_h, alarm_m = map(int, alarm_time.split(":"))
    wake_h, wake_m = map(int, wake_time.split(":"))

    alarm_minutes = alarm_h * 60 + alarm_m
    wake_minutes = wake_h * 60 + wake_m
    diff = wake_minutes - alarm_minutes

    if diff < 0:
        return "early"
    if diff <= 10:
        return "on time"
    return "late"
```