id: 69b1028d6e265413d0198a2d
title: "Izazov 234: Šaljiva Brojka"
challengeType: 29
dashedName: challenge-234
---

# --description--

Ako vam je dat niz brojeva gde svi elementi, osim jednog, prate određeni obrazac, vratite novi niz podataka sa ispravljenim jedinim elementom.

Obrazac će biti jedan od sledećih:

- Brojevi rastu od jednog do drugog za fiksni korak (dodavanje).
- Brojevi opadaju od jednog do drugog za fiksni korak (oduzimanje).

Na primer, ako dobijete `[2, 4, 7, 8, 10]`, vratite `[2, 4, 6, 8, 10]`.

# --hints--

`fix_prank_number([2, 4, 7, 8, 10])` bi trebalo da vrati `[2, 4, 6, 8, 10]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_prank_number([2, 4, 7, 8, 10]), [2, 4, 6, 8, 10])`)
}})
```

`fix_prank_number([10, 10, 8, 7, 6])` bi trebalo da vrati `[10, 9, 8, 7, 6]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_prank_number([10, 10, 8, 7, 6]), [10, 9, 8, 7, 6])`)
}})
```

`fix_prank_number([12, 24, 36, 48, 61, 72, 84, 96])` bi trebalo da vrati `[12, 24, 36, 48, 60, 72, 84, 96]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_prank_number([12, 24, 36, 48, 61, 72, 84, 96]), [12, 24, 36, 48, 60, 72, 84, 96])`)
}})
```

`fix_prank_number([4, 1, -2, -5, -8, -5])` bi trebalo da vrati `[4, 1, -2, -5, -8, -11]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_prank_number([4, 1, -2, -5, -8, -5]), [4, 1, -2, -5, -8, -11])`)
}})
```

`fix_prank_number([0, 100, 200, 300, 150, 500])` bi trebalo da vrati `[0, 100, 200, 300, 400, 500]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_prank_number([0, 100, 200, 300, 150, 500]), [0, 100, 200, 300, 400, 500])`)
}})
```

`fix_prank_number([400, 425, 400, 375, 350, 325, 300])` bi trebalo da vrati `[450, 425, 400, 375, 350, 325, 300]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_prank_number([400, 425, 400, 375, 350, 325, 300]), [450, 425, 400, 375, 350, 325, 300])`)
}})
```

`fix_prank_number([-5, 5, 10, 15, 20])` bi trebalo da vrati `[0, 5, 10, 15, 20]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_prank_number([-5, 5, 10, 15, 20]), [0, 5, 10, 15, 20])`)
}})
```

# --seed--

## --seed-contents--

```py
def fix_prank_number(arr):

    return arr
```

# --solutions--

```py
def fix_prank_number(arr):
    diffs = [arr[i + 1] - arr[i] for i in range(len(arr) - 1)]

    step = diffs[0]
    for i in range(len(diffs) - 1):
        if diffs[i] == diffs[i + 1]:
            step = diffs[i]
            break

    result = arr[:]

    for i in range(len(result) - 1):
        if result[i + 1] - result[i] != step:
            if i == 0 and result[2] - result[1] == step:
                result[0] = result[1] - step
            else:
                result[i + 1] = result[i] + step
            break

    return result
```