id: 69bc6cb30c1d112a2e110a04
title: "Izazov 247: Posledva slova"
challengeType: 29
dashedName: challenge-247
---

# --description--

Ako vam je dat niz slova, vratite slovo koje se pojavljuje najkasnije u alfabetu.

- Ako više od dva slova izgledaju kao da su poslednja u alfabetu, vratite prvo pomenuto.
- Ignorišite sva nealfabetička slova.

# --hints--

`get_last_letter("world")` should return `"w"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_last_letter("world"), "w")`)
}})
```

`get_last_letter("Hello World")` should return `"W"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_last_letter("Hello World"), "W")`)
}})
```

`get_last_letter("The quick brown fox jumped over the lazy dog.")` should return `"z"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_last_letter("The quick brown fox jumped over the lazy dog."), "z")`)
}})
```

`get_last_letter("HeLl0")` should return `"L"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_last_letter("HeLl0"), "L")`)
}})
```

`get_last_letter("!#$ er@R asd fT.,> 2t0e9")` should return `"T"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_last_letter("!#$ er@R asd fT.,> 2t0e9"), "T")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_last_letter(s):

    return s
```

# --solutions--

```py
def get_last_letter(s):
    letters = [c for c in s if c.isalpha()]
    return max(letters, key=lambda c: c.lower())
```