id: 69bc6cb30c1d112a2e110a07
title: "Izazov 250: Skriveni ključ"
challengeType: 29
dashedName: challenge-250
---

# --description--

Dobrodošli u dnevni izazov 250!

Dekriptujte sekvencu šifrovatih slova koristeći ključ za dekripciju i vratite rezultat.

Da biste pronašli ključ:

- Pogledajte sve dnevne izazove do danas čiji je broj izazova deo 25 (uključujući ovaj).
- Uzmite prvo slovo iz naslova svakog od tih izazova i spojite ih u sekvencu slova. Ako naslov počinje sa ne-slovom, pronađite njegovo prvo slovo.

Da biste dekriptovali poruku, prođite kroz svako slovo u šifrovanoj poruci i:

1. Pronađite odgovarajuće slovo u ključu (ponavljajte ključ ako je poruka duža od ključa).
  2. Pretvorite slovo iz ključa u odgovarajući broj: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26.
  3. Pomerite šifrovano slovo unazad u alfabet za taj broj.
  4. Ako pomeranje ide pre od slova `"A"`, zavojte nazad na `"Z"`.

Na primer, ako šifrovovana poruka počinje sa `"Y"` i prvo slovo ključa je `"V"` (22), pomerite `"Y"` unazad za 22 mesta da biste dobili `"C"`. Ponovite ovaj proces za svako slovo kako biste dekriptovali celu poruku.

- Šifruju se samo slova, razmaci se vraćaju kako su bili.
- Sva izvlačena i vraćena slova su velika slova.

# --hints--

`decode("YAVJYNXE")` should return `"CONGRATS"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YAVJYNXE"), "CONGRATS")`)
}})
```

`decode("YALLUT PQUMJP")` should return `"CODING LEGEND"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("YALLUT PQUMJP"), "CODING LEGEND")`)
}})
```

`decode("UAC DYR EISAKYM")` should return `"YOU ARE AWESOME"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("UAC DYR EISAKYM"), "YOU ARE AWESOME")`)
}})
```

`decode("GQMS NBMZU")` should return `"KEEP GOING"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("GQMS NBMZU"), "KEEP GOING")`)
}})
```

`decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB")` should return `"A WINNER IS A DREAMER WHO NEVER GIVES UP"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB"), "A WINNER IS A DREAMER WHO NEVER GIVES UP")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message):

    return message
```

# --solutions--

```py
def decode(message):
    key = "VLHCGMDLNH"
    result = []
    key_index = 0
    for char in message:
        if char == ' ':
            result.append(char)
        else:
            shift = ord(key[key_index % len(key)]) - 64
            key_index += 1
            result.append(chr((ord(char) - 65 - shift) % 26 + 65))
    return ''.join(result)
```