id: 69c5f3d787b1725d5f00c8bd
title: "Izazov 259: Eksplozija FizzBuzz-a"
challengeType: 29
dashedName: challenge-259
---

# --description--

Ako dobijete celo broj, vratite broj koraka potrebnih da se reč `"fizzbuzz"` transformiše u niz slova koji sadrži najmanje određeni broj `"z"` koristeći sledeća pravila:

- Počnite sa nizu slova `"fizzbuzz"`.
- Svaki korak, koristite standardna FizzBuzz pravila na osnovu pozicije slova u nizu (prvo slovo `"f"` je na poziciji 1).
  - Ako se pozicija slova deli sa 3, zamenite to slovo sa `"fizz"`
  - Ako se deli sa 5, zamenite to slovo sa `"buzz"`
  - Ako se deli i sa 3 i sa 5, zamenite to slovo sa `"fizzbuzz"`

Tako da nakon prvog koraka, `"fizzbuzz"` postaje `"fifizzzbuzzfizzzz"`, koji sadrži devet `"z"`.

# --hints--

`explode_fizzbuzz(9)` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(9), 1)`)
}})
```

`explode_fizzbuzz(15)` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(15), 2)`)
}})
```

`explode_fizzbuzz(51)` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(51), 3)`)
}})
```

`explode_fizzbuzz(52)` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(52), 4)`)
}})
```

`explode_fizzbuzz(359)` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(359), 5)`)
}})
```

`explode_fizzbuzz(789)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(789), 6)`)
}})
```

`explode_fizzbuzz(54482)` should return `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(54482), 11)`)
}})
```

`explode_fizzbuzz(1000000)` should return `14`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(explode_fizzbuzz(1000000), 14)`)
}})
```

# --seed--

## --seed-contents--

```py
def explode_fizzbuzz(target_z_count):

    return target_z_count
```

# --solutions--

```py
def explode_fizzbuzz(target_z_count):
    s = "fizzbuzz"
    steps = 0
    while s.count('z') < target_z_count:
        next_s = ""
        for i, c in enumerate(s):
            pos = i + 1
            if pos % 15 == 0:
                next_s += "fizzbuzz"
            elif pos % 3 == 0:
                next_s += "fizz"
            elif pos % 5 == 0:
                next_s += "buzz"
            else:
                next_s += c
        s = next_s
        steps += 1
    return steps
```