id: 69c6ff713a52713463aa7928
title: "Izazov 262: Parser URL upita"
challengeType: 29
dashedName: challenge-262
---

# --description--

Ako URL sadrži niz znakova za upitni parametar, razbijte taj niz znakova u objekat (ili rječnik) parova ključ-vrijednost.

- Niz znakova upitnog parametra počinje nakon `"?"`,
- svaki parametar je razdvojen sa `"&"`,
- svaki par ključ/vrijednost je razdvojen sa `"="`

Na primjer, za `"https://example.com/search?name=Alice&age=30"`, vratite:

```json
{
  "name": "Alice",
  "age": "30"
}
```

Sve vrijednosti treba da se vrate kao niz znakova.

# --hints--

`parse_url_query("https://example.com/search?name=Alice&age=30")` should return `{"name": "Alice", "age": "30"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://freecodecamp.org/learn?skill=programming&language=python"), {"skill": "programming", "language": "python"})`)
}})
```

`parse_url_query("https://freecodecamp.org/items?category=books&sort=asc&page=2")` should return `{"category": "books", "sort": "asc", "page": "2"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://freecodecamp.org/items?category=books&sort=asc&page=2"), {"category": "books", "sort": "asc", "page": "2"})`)
}})
```

`parse_url_query("https://example.com?redirect=freecodecamp.org/learn&when=now")` should return `{"redirect": "freecodecamp.org/learn", "when": "now"}`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(parse_url_query("https://example.com?redirect=freecodecamp.org/learn&when=now"), {"redirect": "freecodecamp.org/learn", "when": "now"})`)
}})
```

# --seed--

## --seed-contents--

```py
def parse_url_query(url):

    return url
```

# --solutions--

```py
def parse_url_query(url):
    query_string = url.split('?')[1]
    return dict(param.split('=') for param in query_string.split('&'))
```