id: 69cfca90e8a0a6d4d6871c53
title: "Challenge 269: Allergen Friendly Meals"
challengeType: 29
dashedName: challenge-269
---

# --description--

Dato listu obroka i listu alergena koje treba izbegavati, vratite imena svih obroka koji ne sadrže nijedan naveden alergen.

- Svaki obrok je u formatu `[obrok, alergeni]`, gde je `obrok` ime obroka, a `alergeni` lista alergena sadržanih u tom obroku. Primer: `["pasta", ["wheat", "milk"]]`.
- Alergeni koje treba izbegavati će biti lista stringova.

Vratite imena sigurnih obroka u istom redosledu kako su dati. Ako nema sigurnih obroka, vratite praznu listu.

# --hints--

`get_allergen_friendly_meals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"])` bi trebalo da vrati `["salad"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_allergen_friendly_meals([["pasta", ["wheat", "milk"]], ["salad", ["nuts"]]], ["milk"]), ["salad"])`)
}})
```

`get_allergen_friendly_meals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"])` bi trebalo da vrati `["fried rice", "chicken parmesan"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_allergen_friendly_meals([["steak", ["soy"]], ["fried rice", []], ["fish tacos", ["fish", "wheat"]], ["chicken parmesan", ["wheat", "milk"]]], ["soy", "fish"]), ["fried rice", "chicken parmesan"])`)
}})
```

`get_allergen_friendly_meals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"])` bi trebalo da vrati `["oatmeal", "granola", "toast"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_allergen_friendly_meals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["eggs", "milk"]), ["oatmeal", "granola", "toast"])`)
}})
```

`get_allergen_friendly_meals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"])` bi trebalo da vrati `["granola", "yogurt", "eggs"]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_allergen_friendly_meals([["oatmeal", ["nuts"]], ["pancakes", ["wheat", "milk"]], ["granola", []], ["yogurt", ["milk"]], ["eggs", ["eggs", "milk"]], ["toast", ["wheat"]]], ["wheat", "nuts"]), ["granola", "yogurt", "eggs"])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_allergen_friendly_meals(meals, allergens):

    return meals
```

# --solutions--

```py
def get_allergen_friendly_meals(meals, allergens):
    return [name for name, meal_allergens in meals if not any(a in meal_allergens for a in allergens)]
```