id: 69cfca90e8a0a6d4d6871c51
title: "Izazov 267: Konverter parsec-a"
challengeType: 29
dashedName: challenge-267
---

# --description--

U dalekoj galaksiji, parsec se koristi za merenje vremena i udaljenosti. Ako dobijete celoban broj parsec-a, vratite njegov ekvivalent u vremenu ili udaljenosti.

- Ako je celobrojni broj neparno, predstavlja vreme. Ako je parno, predstavlja udaljenost.

Koristite ove konverzioni standarde:

| Parsecs | Vreme/Udaljenost |
| - | - |
| 1 | 2 sata |
| 2 | 6 svetlosnih godina |

Vratite konvertovanu vrednost kao celobrojni broj.

# --hints--

`convert_parsecs(1)` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(1), 2)`)
}})
```

`convert_parsecs(2)` should return `6`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(2), 6)`)
}})
```

`convert_parsecs(31)` should return `62`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(31), 62)`)
}})
```

`convert_parsecs(88)` should return `264`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(88), 264)`)
}})
```

`convert_parsecs(17)` should return `34`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(17), 34)`)
}})
```

`convert_parsecs(14)` should return `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_parsecs(14), 42)`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_parsecs(parsecs):

    return parsecs
```

# --solutions--

```py
def convert_parsecs(parsecs):
    if parsecs % 2 != 0:
        return parsecs * 2
    else:
        return parsecs * 3
```