id: 69cfca90e8a0a6d4d6871c52
title: "Challenge 268: Narcissistic Number"
challengeType: 29
dashedName: challenge-268
---

# --description--

Ako je dat pozitivan cijeli broj, utvrdite da li je narcisistički broj.

- Broj je narcisistički ako je zbir svih njegovih cifara podignut na stepen ukupnog broja cifara jednak samom broju.

Na primjer, 153 ima 3 cifre, i $1^3 + 5^3 + 3^3 = 153$, pa je narcisistički.

# --hints--

`is_narcissistic(153)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(153), True)`)
}})
```

`is_narcissistic(154)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(154), False)`)
}})
```

`is_narcissistic(371)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(371), True)`)
}})
```

`is_narcissistic(512)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(512), False)`)
}})
```

`is_narcissistic(9)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(9), True)`)
}})
```

`is_narcissistic(11)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(11), False)`)
}})
```

`is_narcissistic(9474)` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(9474), True)`)
}})
```

`is_narcissistic(6549)` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_narcissistic(6549), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_narcissistic(n):

    return n
```

# --solutions--

```py
def is_narcissistic(n):
    digits = str(n)
    power = len(digits)
    return sum(int(d) ** power for d in digits) == n
```