id: 69e2383af7832c8032603b93
title: "Izazov 277: Ogledalni odraz"
challengeType: 29
dashedName: challenge-277
---

# --description--

U razmatranju niza od dva slova, utvrdite da li je drugi niz ogledalni odraz prvog niza.

Ogledalni odraz se stvara menjanjem niza slova i uzimanjem svakog slova koje se menja sa svojim ekvivalentom u ogledalu.

- Slova koja izgledaju isto u ogledalu su:

`W`, `T`, `Y`, `U`, `I`, `O`, `H`, `A`, `X`, `V`, `M`, `w`, `o`, `x`, `v`, `0`, `8`, `=`, `+`, `:`, `|`, `-`, `_`, `*`, `^`, `!`, `.`, i razmak (` `).

- Parovi koji se menjaju u ogledalu menjaju se međusobno:

| Slovo | Menja se sa |
| - | - |
| `[` | `]` |
| `{` | `}` |
| `<` | `>` |
| `b` | `d` |
| `p` | `q` |
| `(` | `)` |

Ako bilo koji niz sadrži slovo koje nije na gore navedenoj listi, on nema ogledalni odraz koji se može stvoriti iz tih slova.

Na primer, ogledalni odraz `"[HOW]"` je `"[WOH]"`.

# --hints--

`is_mirror_image("[HOW]", "[WOH]")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("[HOW]", "[WOH]"), True)`)
}})
```

`is_mirror_image("MOM", "MOM")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("MOM", "MOM"), True)`)
}})
```

`is_mirror_image("vow", "wov")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("vow", "wov"), True)`)
}})
```

`is_mirror_image("TIM", "TIM")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("TIM", "TIM"), False)`)
}})
```

`is_mirror_image("{WOW}", "}WOW{")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("{WOW}", "}WOW{"), False)`)
}})
```

`is_mirror_image("XXVII", "IIV%X")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("XXVII", "IIV%X"), False)`)
}})
```

`is_mirror_image("><(((*>", "<*)))><")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("><(((*>", "<*)))><"), True)`)
}})
```

`is_mirror_image("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_mirror_image("WTYUIOHAXVMwoxv08=+:|-_*^!.[]{}<>bdpq()", "()pqbd<>{}[].!^*_-|:+=80vxowMVXAHOIUYTW"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_mirror_image(s1, s2):

    return s1
```

# --solutions--

```py
def is_mirror_image(s1, s2):
    symmetric = set('WTYUIOHA XVMwoxv08=+:|-_*^!.')
    pairs = {'[': ']', ']': '[', '{': '}', '}': '{', '<': '>', '>': '<',
             'b': 'd', 'd': 'b', 'p': 'q', 'q': 'p', '(': ')', ')': '('}
    mirrored = []
    for c in reversed(s1):
        if c in symmetric:
            mirrored.append(c)
        elif c in pairs:
            mirrored.append(pairs[c])
        else:
            return False
    return ''.join(mirrored) == s2
```