id: 69f35a5bb823ed620fcb7cbb
title: "Challenge 282: Dugina spjava"
challengeType: 29
dashedName: challenge-282
---

# --description--

Uzimajući u obzir niz podataka sati koje ste spavali svake noći do danas i ciljani broj sati po noći, vratite koliko sati sna vam je potrebno večeras da sanirate dug spavanja.

- Uključite sate odnoć do danas u ukupan period potreban za postizanje ciljeva.
- Ako ste spavali dovoljno da dostignete ili premašite cilj za večeras, vratite `0`.

# --hints--

`sleep_debt([6, 6, 6, 6, 6, 6], 8)` should return `20`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([6, 6, 6, 6, 6, 6], 8), 20)`)
}})
```

`sleep_debt([6, 7, 8, 4, 8, 6], 7)` should return `10`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([6, 7, 8, 4, 8, 6], 7), 10)`)
}})
```

`sleep_debt([10, 10, 9, 10, 9, 11], 9)` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([10, 10, 9, 10, 9, 11], 9), 4)`)
}})
```

`sleep_debt([8, 7, 6, 7, 6, 8], 6)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([8, 7, 6, 7, 6, 8], 6), 0)`)
}})
```

`sleep_debt([8, 9, 10, 9, 10, 7], 7)` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sleep_debt([8, 9, 10, 9, 10, 7], 7), 0)`)
}})
```

# --seed--

## --seed-contents--

```py
def sleep_debt(hours_slept, target_hours):

    return hours_slept
```

# --solutions--

```py
def sleep_debt(hours_slept, target_hours):
    debt = (len(hours_slept) + 1) * target_hours - sum(hours_slept)
    return max(0, debt)
```