id: 69f35a5bb823ed620fcb7cb9
title: "Challenge 280: Mongo ID Date"
challengeType: 29
dashedName: challenge-280
---

# --description--

Izvuqanjem MongoDB ID sekvence karaktera, vratite njegovo vreme kreiranja kao ISO 8601 sekvencu karaktera.

- MongoDB ID je 24-karakterna heksadecimalna sekvenca karaktera. Prvih osam karaktera predstavlja Unix vremenski stamp (u sekundama) predstavljen kao broj u bazi 16.

Na primer, `"6a094b50bcf86cd799439011"` ima vremenski stamp `"6a094b50"` u heksadecimalnom formatu, što je `1778994000` u decimalnom formatu, predstavljajući vreme kreiranja `"2026-05-17T05:00:00.000Z"`.

# --hints--

`mongo_id_to_date("6a094b50bcf86cd799439011")` bi trebalo da vrati `"2026-05-17T05:00:00.000Z"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mongo_id_to_date("6a094b50bcf86cd799439011"), "2026-05-17T05:00:00.000Z")`)
}})
```

`mongo_id_to_date("695344eb1f4a4c1123042128")` bi trebalo da vrati `"2025-12-30T03:20:11.000Z"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mongo_id_to_date("695344eb1f4a4c1123042128"), "2025-12-30T03:20:11.000Z")`)
}})
```

`mongo_id_to_date("386da62df34123ac54617e56")` bi trebalo da vrati `"2000-01-01T07:01:01.000Z"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mongo_id_to_date("386da62df34123ac54617e56"), "2000-01-01T07:01:01.000Z")`)
}})
```

`mongo_id_to_date("69f571c3d7711807afd3dd55")` bi trebalo da vrati `"2026-05-02T03:38:43.000Z"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mongo_id_to_date("69f571c3d7711807afd3dd55"), "2026-05-02T03:38:43.000Z")`)
}})
```

`mongo_id_to_date("68adce01c0e1144d0a90295a")` bi trebalo da vrati `"2025-08-26T15:08:49.000Z"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(mongo_id_to_date("68adce01c0e1144d0a90295a"), "2025-08-26T15:08:49.000Z")`)
}})
```

# --seed--

## --seed-contents--

```py
def mongo_id_to_date(s):

    return s
```

# --solutions--

```py
from datetime import datetime, timezone

def mongo_id_to_date(s):
    timestamp = int(s[:8], 16)
    dt = datetime.fromtimestamp(timestamp, tz=timezone.utc)
    return dt.strftime("%Y-%m-%dT%H:%M:%S.000Z")
```