id: 69f8c998d78ad3171a0713ba
title: "Vežba 287: Ispravljanje romanskih brojeva"
challengeType: 29
dashedName: challenge-287
---

# --description--

Ako dobijete niz slova koji predstavljaju roman brojeve i koji nisu pravilno korišćeni, vratite njihovu vrednost koristeći standardnu notaciju romanskih brojeva.

Ulaz će koristiti samo sabirnu notaciju, pa svaki simbol dodaje svoju vrednost ukupnom iznosu. Kao podsetnik, ovo su simbole i njihove vrednosti:

| Simbol | Vrednost |
| ------ | ------- |
| `"I"`  | 1       |
| `"V"`  | 5       |
| `"X"`  | 10      |
| `"L"`  | 50      |
| `"C"`  | 100     |
| `"D"`  | 500     |
| `"M"`  | 1000    |

Kada ponovo kodirate, koristite najveći mogući simbol na svakom koraku, koristeći oduzimajuće parove (subtractive pairs) (`"IV"`, `"IX"`, `"XL"`, `"XC"`, `"CD"`, `"CM"`) kada je potrebno.

# --hints--

`fix_numerals("XIIIII")` bi trebalo da vrati `"XV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("XIIIII"), "XV")`)
}})
```

`fix_numerals("IIIILX")` bi trebalo da vrati `"LXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("IIIILX"), "LXIV")`)
}})
```

`fix_numerals("XXVVVIIIII")` bi trebalo da vrati `"XL"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("XXVVVIIIII"), "XL")`)
}})
```

`fix_numerals("MDCCLXXXXVIIII")` bi trebalo da vrati `"MDCCXCIX"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("MDCCLXXXXVIIII"), "MDCCXCIX")`)
}})
```

`fix_numerals("IIIIVVVVXXXXLLLLCCDD")` bi trebalo da vrati `"MCDLXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("IIIIVVVVXXXXLLLLCCDD"), "MCDLXIV")`)
}})
```

`fix_numerals("ILCDMIVDIIXLCVCXDL")` bi trebalo da vrati `"MMCMLXXXIV"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(fix_numerals("ILCDMIVDIIXLCVCXDL"), "MMCMLXXXIV")`)
}})
```

# --seed--

## --seed-contents--

```py
def fix_numerals(s):

    return s
```

# --solutions--

```py
def fix_numerals(s):
    values = {"I": 1, "V": 5, "X": 10, "L": 50, "C": 100, "D": 500, "M": 1000}
    encoding = [
        (1000, "M"), (900, "CM"), (500, "D"), (400, "CD"),
        (100, "C"), (90, "XC"), (50, "L"), (40, "XL"),
        (10, "X"), (9, "IX"), (5, "V"), (4, "IV"), (1, "I")
    ]

    total = sum(values[c] for c in s)

    result = ""
    for value, symbol in encoding:
        while total >= value:
            result += symbol
            total -= value
    return result
```