id: 69f8c998d78ad3171a0713bb
title: "Vežba 288: Tajna brojka"
challengeType: 29
dashedName: challenge-288
---

# --description--

Uzimajući u obzir tajnu brojku i pogađanje, utvrdite da li je pogađanje tačno.

Vratite:

- `"higher"` ako je tajna brojka veća od pogađanja.
- `"lower"` ako je tajna brojka manja od pogađanja.
- `"you got it!"` ako je pogađanje tačno.

# --hints--

`guess_number(50, 30)` bi trebalo da vrati `"higher"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(50, 30), "higher")`)
}})
```

`guess_number(85, 99)` bi trebalo da vrati `"lower"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(85, 99), "lower")`)
}})
```

`guess_number(2026, 2026)` bi trebalo da vrati `"you got it!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(2026, 2026), "you got it!")`)
}})
```

`guess_number(92904, 11283)` bi trebalo da vrati `"higher"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(92904, 11283), "higher")`)
}})
```

`guess_number(230495, 423920)` bi trebalo da vrati `"lower"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(230495, 423920), "lower")`)
}})
```

`guess_number(120349, 120349)` bi trebalo da vrati `"you got it!"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(guess_number(120349, 120349), "you got it!")`)
}})
```

# --seed--

## --seed-contents--

```py
def guess_number(secret, guess):

    return secret
```

# --solutions--

```py
def guess_number(secret, guess):
    if guess == secret:
        return "you got it!"
    return "higher" if guess < secret else "lower"
```