id: 69f8c998d78ad3171a0713b9
title: "Izazov 286: Otvoreni problemi"
challengeType: 29
dashedName: challenge-286
---

# --description--

Uzimajući u obzir niz brojeva problema i drugi niz brojeva zahteva za povlačenje promena (PR), vratite niz problema koji su još uvek otvoreni nakon što su svi PR zahtevi spojisani.

- Zahtev za PR zatvara problem ako su njihovi brojevi rotacija jedni drugima. Na primer, problem 123 bi bio zatvoren sa PR 231 ili 312.
- Zahtev za PR ne zatvara problem koji ima potpuno isti broj. Na primer, PR 123 ne zatvara problem 123. Stoga, problem sa istim brojevima ne može biti zatvoren.
- Bilo koji broj može imati uklonjene vodeće nule. Na primer, PR 201 bi zatvorio problem 12 (012, rotacija od 201). Slično tome, problem 201 bi bio zatvoren sa PR 12.

Vratite otvorene probleme koji su ostali u redosledu datih.

# --hints--

`get_open_issues([123, 234], [231])` should return `[234]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 234], [231]), [234])`)
}})
```

`get_open_issues([123, 345, 16], [345, 231])` should return `[345, 16]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([123, 345, 16], [345, 231]), [345, 16])`)
}})
```

`get_open_issues([456, 332, 12, 15], [201, 945, 180])` should return `[456, 332, 15]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15])`)
}})
```

`get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` should return `[115, 296, 24]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24])`)
}})
```

`get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` should return `[95, 395, 754, 296, 709, 237, 1802]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_open_issues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_open_issues(issues, prs):

    return issues
```

# --solutions--

```py
def get_open_issues(issues, prs):
    def get_rotations(n):
        s = str(n)
        rotations = set()
        for i in range(len(s)):
            rotated = s[i:] + s[:i]
            rotations.add(int(rotated))
        return rotations

    def are_rotations(a, b):
        if a == b:
            return False
        return b in get_rotations(a)

    return [issue for issue in issues if not any(are_rotations(pr, issue) or are_rotations(issue, pr) for pr in prs)]
```