id: 69f8c998d78ad3171a0713bf
title: "Vežba 292: Široki omjer aspekta"
challengeType: 29
dashedName: challenge-292
---

# --description--

Ako dobijete dva niza sa različitim dimenzijama slika, vratite omjer aspekta slike koja ima veći omjer širine na visinu.

- Omjer aspekta koji je dat će biti u formatu `"ŠxV"`, na primer, `"1920x1080"`.
- Omjer aspekta je odnos širine na visinu, sveden na najmanji cijeli broj. Na primjer, `"1920x1080"` se svodi na `"16:9"`.
- Vratite niz u formatu `"Š:V"`, na primjer, `"16:9"`.

# --hints--

`get_wider_aspect_ratio("1920x1080", "800x600")` treba da vrati `"16:9"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1920x1080", "800x600"), "16:9")`)
}})
```

`get_wider_aspect_ratio("1080x1350", "2048x1536")` treba da vrati `"4:3"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("1080x1350", "2048x1536"), "4:3")`)
}})
```

`get_wider_aspect_ratio("640x480", "2440x1220")` treba da vrati `"2:1"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("640x480", "2440x1220"), "2:1")`)
}})
```

`get_wider_aspect_ratio("360x640", "1080x1920")` treba da vrati `"9:16"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("360x640", "1080x1920"), "9:16")`)
}})
```

`get_wider_aspect_ratio("3440x1440", "2048x858")` treba da vrati `"43:18"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("3440x1440", "2048x858"), "43:18")`)
}})
```

`get_wider_aspect_ratio("12345x61234", "12534x51234")` treba da vrati `"2089:8539"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_wider_aspect_ratio("12345x61234", "12534x51234"), "2089:8539")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_wider_aspect_ratio(a, b):

    return a
```

# --solutions--

```py
from math import gcd

def get_wider_aspect_ratio(a, b):
    def parse(s):
        w, h = map(int, s.split("x"))
        g = gcd(w, h)
        return w / h, f"{w // g}:{h // g}"

    ratio_a, simplified_a = parse(a)
    ratio_b, simplified_b = parse(b)

    return simplified_a if ratio_a >= ratio_b else simplified_b
```