id: 69f90c1329a94b37e2a2086d
title: "Vežba 294: Kombinacija običnih zagradi"
challengeType: 29
dashedName: challenge-294
---

# --description--

Ako vam je dat ceo broj, `n`, vratite broj validnih kombinacija za par `n` običnih zagradi.

- Validna kombinacija je sekvenca karaktera gde svaka otvarajuća obična zagrada ima odgovarajuću zatvarajuću običnu zagradu, i nijedna zatvarajuća obična zagrada ne nastaje pre odgovarajuće otvarajuće obične zagrade.

Na primer, ako vam je dat `2`, postoji 2 validne kombinacije:

```json
(())
()()
```

# --hints--

`get_combinations(2)` bi trebalo da vrati `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(2), 2)`)
}})
```

`get_combinations(3)` bi trebalo da vratiti `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(3), 5)`)
}})
```

`get_combinations(5)` bi trebalo da vratiti `42`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(5), 42)`)
}})
```

`get_combinations(8)` bi trebalo da vratiti `1430`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(8), 1430)`)
}})
```

`get_combinations(13)` bi trebalo da vratiti `742900`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_combinations(13), 742900)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_combinations(n):

    return n
```

# --solutions--

```py
def get_combinations(n):
    count = 0

    def generate(op, cl):
        nonlocal count
        if op == 0 and cl == 0:
            count += 1
            return
        if op > 0:
            generate(op - 1, cl)
        if cl > op:
            generate(op, cl - 1)

    generate(n, n)
    return count
```