id: 6a0dcc730cb92a616f86f0bf
title: "Vežba 295: Proveri strukturu deo 1"
challengeType: 29
dashedName: challenge-295
---

# --description--

Ako vam je dat objekat (JavaScript) ili rečnik (Python), utvrdite da li odgovara sledećoj strukturi:

```json
{
  username: string
}
```

- Dopušteni su dodatni ključevi

# --hints--

`is_valid_schema({"username": "bob"})` bi trebalo da vrati `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "bob"}), True)`)
}})
```

`is_valid_schema({"username": "jen", "posts": 30})` bi trebalo da vrati `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "jen", "posts": 30}), True)`)
}})
```

`is_valid_schema({"username": ""})` bi trebalo da vrati `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": ""}), True)`)
}})
```

`is_valid_schema({"username": 7})` bi trebalo da vrati `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": 7}), False)`)
}})
```

`is_valid_schema({"posts": 25})` bi trebalo da vrati `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"posts": 25}), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_schema(obj):

    return obj
```

# --solutions--

```py
def is_valid_schema(obj):
    return isinstance(obj.get("username"), str)
```