id: 6a0dcd03ee4e68698080ef6b
title: "Vežba 307: Pravila gradskog planiranja"
challengeType: 29
dashedName: challenge-307
---

# --description--

Ako vam je data 2D mreža (niz niz podataka) koja predstavlja raspored gradskih objekata, vratite koordinate svih objekata koji krše pravila gradskog planiranja.

Svaka ćelija u mrežama ima jednu oznaku iz tabele ispod. Objekat krši pravilo ako je bilo koji od njegovih suseda (do četiri), vertikalno ili horizontalno, tip koji ne sme biti sused.

| Oznaka | Tip | Ne može biti susjed sa |
| - | - | - |
| `"i"` | industrija | `"R"`, `"I"` |
| `"A"` | poljoprivreda | `"C"` |
| `"R"` | stambeno | `"i"`, `"C"` |
| `"I"` | institucije | `"i"` |
| `"C"` | komercijalno | `"R"`, `"A"` |
| `""` (niz praznih karaktera) | neizgrađeno | nema ograničenja |

Vratite koordinate svih ćelija koje krše kao niz parova `[red, kol]`, u bilo kom redosledu. Ako ne postoji nijedan prekršaj, vratite prazan niz.

# --hints--

`get_zone_violations([["R", "C"], ["", "C"]])` bi trebalo da vrati `[[0, 0], [0, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "C"], ["", "C"]])), [[0, 0], [0, 1]])`)
}})
```

`get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])` bi trebalo da vrati `[[0, 1], [1, 1]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["", "i"], ["", "R"], ["R", "I"]])), [[0, 1], [1, 1]])`)
}})
```

`get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` bi trebalo da vrati `[]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])), [])`)
}})
```

`get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` bi trebalo da vrati `[[0, 1], [0, 2], [0, 3]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])), [[0, 1], [0, 2], [0, 3]])`)
}})
```

`get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` bi trebalo da vrati `[[2, 3], [2, 4], [3, 1], [3, 2]]`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(sorted(get_zone_violations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])), [[2, 3], [2, 4], [3, 1], [3, 2]])`)
}})
```

# --seed--

## --seed-contents--

```py
def get_zone_violations(grid):

    return grid
```

# --solutions--

```py
def get_zone_violations(grid):
    rules = {
        "i": ["R", "I"],
        "A": ["C"],
        "R": ["i", "C"],
        "I": ["i"],
        "C": ["R", "A"],
        "": []
    }

    directions = [[-1, 0], [0, -1], [0, 1], [1, 0]]
    violations = []

    for row in range(len(grid)):
        for col in range(len(grid[row])):
            cell = grid[row][col]
            forbidden = rules[cell]

            for dr, dc in directions:
                nr = row + dr
                nc = col + dc
                if 0 <= nr < len(grid) and 0 <= nc < len(grid[row]):
                    if grid[nr][nc] in forbidden:
                        violations.append([row, col])
                        break

    return violations
```