id: 6a0dcc730cb92a616f86f0c5
title: "Vežba 301: Poslednje opterećenje"
challengeType: 29
dashedName: challenge-301
---

# --description--

Uzimajući u obzir broj kašika deterdženta koje imate i niz podataka o kašikama koje ste koristili tokom proteklih dana, vratite ukupan broj punih dana deterdženta koji vam ostaje.

Izračunajte prosečnu potrošnju po danu iz istorije upotrebe i pretpostavite da će ta potrošnja trajati svakog dana.

# --hints--

`last_load_date(10, [2, 2, 2, 2, 2, 2, 2])` bi trebalo da vrati `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(10, [2, 2, 2, 2, 2, 2, 2]), 5)`)
}})
```

`last_load_date(16, [2, 3, 0, 3, 4, 2, 1])` bi trebalo da vrati `7`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(16, [2, 3, 0, 3, 4, 2, 1]), 7)`)
}})
```

`last_load_date(33, [5, 0, 4, 3, 3, 2])` bi trebalo da vrati `11`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(33, [5, 0, 4, 3, 3, 2]), 11)`)
}})
```

`last_load_date(50, [2, 0, 2, 9, 12, 0, 2])` bi trebalo da vrati `12`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(50, [2, 0, 2, 9, 12, 0, 2]), 12)`)
}})
```

`last_load_date(20, [13, 9, 12, 10, 8])` bi trebalo da vrati `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(last_load_date(20, [13, 9, 12, 10, 8]), 1)`)
}})
```

# --seed--

## --seed-contents--

```py
def last_load_date(scoops, usage):

    return scoops
```

# --solutions--

```py
def last_load_date(scoops, usage):
    avg = sum(usage) / len(usage)
    return int(scoops / avg)
```