id: 6a0dcd03ee4e68698080ef68
title: "Vežba 304: Raspored putovanja"
challengeType: 29
dashedName: challenge-304
---

# --description--

Ako se izvuče niz podataka koji sadrži najmanje dve opcione stanice za dnevni put, vratite broj validnih rasporeda.

Raspored uvek uključuje "breakfast", "lunch" i "dinner"; oni se ne primaju kao argumenti. Opcijalne stanice mogu biti postavljene bilo gde u rasporedu, prateći sledeća pravila:

- "breakfast" je uvek prvi, sa najmanje jednom stanicom pre od "lunch".
- "lunch" mora biti prikazan pre od "dinner", sa najmanje jednom stanicom između njih.
- Pre svega, jedna opcionalna stanica može se pojaviti nakon od "dinner".

Vratite broj validnih rasporeda.

# --hints--

`get_itinerary_count(["library", "park"])` bi trebalo da vrati `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park"]), 2)`)
}})
```

`get_itinerary_count(["library", "park", "arcade"])` bi trebalo da vrati `18`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade"]), 18)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store"])` bi trebalo da vrati `120`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store"]), 120)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe"])` bi trebalo da vrati `840`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe"]), 840)`)
}})
```

`get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"])` bi trebalo da vrati `55440`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_itinerary_count(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440)`)
}})
```

# --seed--

## --seed-contents--

```py
def get_itinerary_count(stops):

    return stops
```

# --solutions--

```py
from math import factorial

def get_itinerary_count(stops):
    n = len(stops)
    return (2 * n - 3) * factorial(n)
```