id: 6a0dcc730cb92a616f86f0c0
title: "Vežba 296: Validacija šeme deo 2"
challengeType: 29
dashedName: challenge-296
---

# --description--

Ako vam je dat objekat (JavaScript) ili rečnik (Python), utvrdite da li odgovara sledećoj šemi:

```json
{
  username: string,
  posts: number,
  verified: boolean
}
```

- Dodatni ključevi su dozvoljeni

# --hints--

`is_valid_schema({"username": "alice", "posts": 10, "verified": False})` treba da vrati `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "alice", "posts": 10, "verified": False}), True)`)
}})
```

`is_valid_schema({"username": "carol", "posts": 15, "verified": True, "followers": 25})` treba da vrati `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "carol", "posts": 15, "verified": True, "followers": 25}), True)`)
}})
```

`is_valid_schema({"username": "frank", "posts": "21", "verified": True})` treba da vrati `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "frank", "posts": "21", "verified": True}), False)`)
}})
```

`is_valid_schema({"username": "sam", "posts": 17, "verified": "false"})` treba da vrati `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "sam", "posts": 17, "verified": "false"}), False)`)
}})
```

`is_valid_schema({"username": "bill", "verified": True})` treba da vrati `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "bill", "verified": True}), False)`)
}})
```

`is_valid_schema({"username": "fred", "verified": True})` treba da vrati `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": "fred", "verified": True}), False)`)
}})
```

`is_valid_schema({"username": 5, "posts": 10, "verified": True})` treba da vrati `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_schema({"username": 5, "posts": 10, "verified": True}), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_schema(obj):

    return obj
```

# --solutions--

```py
def is_valid_schema(obj):
    return (
        isinstance(obj.get("username"), str) and
        isinstance(obj.get("posts"), int) and not isinstance(obj.get("posts"), bool) and
        isinstance(obj.get("verified"), bool)
    )
```