id: 6a151d32d772271dd2448c2d
title: "Vežba 311: Čarobnica reči"
challengeType: 29
dashedName: challenge-311
---

# --description--

Kada se koristi sekvenca rečnih kodova, izračunajte ukupne poene.

Svaka slovo u sekvenci predstavlja reč:

| Kod | Reč | Kategorija | Osnovni bodovi |
| - | - | - | - |
| `"f"` | Vatra | Uništenje | 3 |
| `"l"` | Električnost | Uništenje | 3 |
| `"i"` | Led | Kontrola | 2 |
| `"w"` | Vetar | Kontrola | 2 |
| `"h"` | Lekovi | Obnavljanje | 1 |
| `"s"` | Stražar | Obnavljanje | 1 |

Množnik kombinacije se određuje na osnovu koliko uzastopnih reči je korišćeno iz različitih kategorija:

- Prva reč uvek dobija osnovne bodove.
- Svaka sledeća reč iz različite kategorije od prethodne povećava množilac za 1.
- Korišćenje reči iz iste kategorije kao i prethodna vraća množilac na 1.
- Bodovi za svaku reč su njeni osnovni bodovi pomnoženi trenutnim množilacem.

Vratite ukupan broj poena za sekvencu reči.

# --hints--

`cast("fihwl")` bi trebalo da vrati `33`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("fihwl"), 33)`)
}})
```

`cast("lwswfi")` bi trebalo da vrati `45`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("lwswfi"), 45)`)
}})
```

`cast("wislhfl")` bi trebalo da vrati `37`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wislhfl"), 37)`)
}})
```

`cast("sihwlih")` bi trebalo da vrati `50`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("sihwlih"), 50)`)
}})
```

`cast("wishlfihwslwifihl")` bi trebalo da vrati `101`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(cast("wishlfihwslwifihl"), 101)`)
}})
```

# --seed--

## --seed-contents--

```py
def cast(spells):

    return spells
```

# --solutions--

```py
def cast(spells):
    lookup = {
        "f": ("destruction", 3),
        "l": ("destruction", 3),
        "i": ("control", 2),
        "w": ("control", 2),
        "h": ("restoration", 1),
        "s": ("restoration", 1)
    }

    total = 0
    multiplier = 1
    prev_category = None

    for spell in spells:
        category, score = lookup[spell]

        if prev_category and category != prev_category:
            multiplier += 1
        else:
            multiplier = 1

        total += score * multiplier
        prev_category = category

    return total
```