---
id: 5cdafbc32913098997531680
title: Završi obećanje koristeći resolve i reject
challengeType: 1
forumTopicId: 301196
dashedName: complete-a-promise-with-resolve-and-reject
---

# --description--

Obećanje ima tri stanja: `pending`, `fulfilled`, na `rejected`. Obećanje koje ste kreirali u prethodnom izazovu zarobljeno je zauvek u stanju `pending` jer niste naveli način za ispunjavanje tog obećanja. Parametri `resolve` i `reject` koji se pružaju kao argumenti obećanja koriste se za to. `resolve` se koristi kada želite da vaše obećanje uspe, a `reject` se koristi kada želite da ne uspe. Ovo su metode koje primaju argumente, kao što je prikazano ispod.

```js
const myPromise = new Promise((resolve, reject) => {
  if(condition here) {
    resolve("Promise was fulfilled");
  } else {
    reject("Promise was rejected");
  }
});
```

Primer gore koristi sekvencu karaktera kao argumente za ove metode, ali zapravo može biti bilo šta. Često, to može biti nešto iz čega biste koristili podatke da postavite na vaš sajt ili negde drugde.

# --instructions--

Fanya ahadi kushughulikia mafanikio na kushindwa. Ikiwa `responseFromServer` ni `true`, fanya wito wa njia ya `resolve` kumaliza ahadi kwa mafanikio. Pitia `resolve` mfuatano wa herufi wenye thamani ya `We got the data`. Ikiwa `responseFromServer` ni `false`, tumia njia ya `reject` badala yake na upitie mfuatano wa herufi: `Data not received`.

# --hints--

`resolve` mora biti pozvan sa očekivanim nizom znakova kada je stanje `if` na `true`.

```js
assert(
  __helpers.removeJSComments(code).match(/if\s*\(\s*responseFromServer\s*\)\s*{\s*resolve\s*\(\s*('|"|`)We got the data\1\s*\)(\s*|\s*;\s*)}/g)
);
```

`reject` treba da bude nazvan očekivanom sekvencijom karaktera kada je stanje `if` `false`.

```js
assert(
  __helpers.removeJSComments(code).match(/}\s*else\s*{\s*reject\s*\(\s*('|"|`)Data not received\1\s*\)(\s*|\s*;\s*)}/g)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer represents a response from a server
  let responseFromServer;
    
  if(responseFromServer) {
    // Change this line
  } else {  
    // Change this line
  }
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer represents a response from a server
  let responseFromServer;

  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});
```
